# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t104_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:![X4]:![X5]:~((r2_hidden(X1,k3_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X4,X5)))&![X6]:![X7]:~((X1=k4_tarski(X6,X7)&(r2_hidden(X6,k3_xboole_0(X2,X4))&r2_hidden(X7,k3_xboole_0(X3,X5)))))))), file('out/t104_zfmisc_1__2_premises__rerun__rerun.E_input', t104_zfmisc_1)).
fof(t85_zfmisc_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:~((r2_hidden(X1,k3_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X4,X5)))&![X6]:![X7]:~((X1=k4_tarski(X6,X7)&(r2_hidden(X6,k3_xboole_0(X2,X4))&r2_hidden(X7,k3_xboole_0(X3,X5)))))))), file('out/t104_zfmisc_1__2_premises__rerun__rerun.E_input', t85_zfmisc_1)).
fof(c_0_2, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:![X5]:~((r2_hidden(X1,k3_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X4,X5)))&![X6]:![X7]:~((X1=k4_tarski(X6,X7)&(r2_hidden(X6,k3_xboole_0(X2,X4))&r2_hidden(X7,k3_xboole_0(X3,X5))))))))), inference(assume_negation,[status(cth)],[t104_zfmisc_1])).
fof(c_0_3, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:((X8=k4_tarski(esk6_5(X8,X9,X10,X11,X12),esk7_5(X8,X9,X10,X11,X12))|~r2_hidden(X8,k3_xboole_0(k2_zfmisc_1(X9,X10),k2_zfmisc_1(X11,X12))))&((r2_hidden(esk6_5(X8,X9,X10,X11,X12),k3_xboole_0(X9,X11))|~r2_hidden(X8,k3_xboole_0(k2_zfmisc_1(X9,X10),k2_zfmisc_1(X11,X12))))&(r2_hidden(esk7_5(X8,X9,X10,X11,X12),k3_xboole_0(X10,X12))|~r2_hidden(X8,k3_xboole_0(k2_zfmisc_1(X9,X10),k2_zfmisc_1(X11,X12))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_zfmisc_1])])])])])])).
fof(c_0_4, negated_conjecture, (![X13]:![X14]:(r2_hidden(esk1_0,k3_xboole_0(k2_zfmisc_1(esk2_0,esk3_0),k2_zfmisc_1(esk4_0,esk5_0)))&(esk1_0!=k4_tarski(X13,X14)|(~r2_hidden(X13,k3_xboole_0(esk2_0,esk4_0))|~r2_hidden(X14,k3_xboole_0(esk3_0,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
cnf(c_0_5,plain,(X1=k4_tarski(esk6_5(X1,X2,X3,X4,X5),esk7_5(X1,X2,X3,X4,X5))|~r2_hidden(X1,k3_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X4,X5)))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(r2_hidden(esk1_0,k3_xboole_0(k2_zfmisc_1(esk2_0,esk3_0),k2_zfmisc_1(esk4_0,esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(r2_hidden(esk6_5(X1,X2,X3,X4,X5),k3_xboole_0(X2,X4))|~r2_hidden(X1,k3_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X4,X5)))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_8,plain,(r2_hidden(esk7_5(X1,X2,X3,X4,X5),k3_xboole_0(X3,X5))|~r2_hidden(X1,k3_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X4,X5)))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_9,negated_conjecture,(~r2_hidden(X1,k3_xboole_0(esk3_0,esk5_0))|~r2_hidden(X2,k3_xboole_0(esk2_0,esk4_0))|esk1_0!=k4_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(k4_tarski(esk6_5(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),esk7_5(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0))=esk1_0), inference(spm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_11,negated_conjecture,(r2_hidden(esk6_5(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k3_xboole_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_7, c_0_6])).
cnf(c_0_12,negated_conjecture,(r2_hidden(esk7_5(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_8, c_0_6])).
cnf(c_0_13,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 14
# Proof object clause steps            : 9
# Proof object formula steps           : 5
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 3
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 13
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 13
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 4
# ...of the previous two non-trivial   : 3
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 8
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 4
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 583

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 45232 pages
