# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t46_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), file('out/t107_xboole_1__17_premises__rerun.E_input', t46_xboole_1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t107_xboole_1__17_premises__rerun.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t107_xboole_1__17_premises__rerun.E_input', t100_xboole_1)).
fof(t21_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), file('out/t107_xboole_1__17_premises__rerun.E_input', t21_xboole_1)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t107_xboole_1__17_premises__rerun.E_input', t3_boole)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t107_xboole_1__17_premises__rerun.E_input', t91_xboole_1)).
fof(t2_boole, axiom, (![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), file('out/t107_xboole_1__17_premises__rerun.E_input', t2_boole)).
fof(l97_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), file('out/t107_xboole_1__17_premises__rerun.E_input', l97_xboole_1)).
fof(t49_xboole_1, axiom, (![X1]:![X2]:![X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), file('out/t107_xboole_1__17_premises__rerun.E_input', t49_xboole_1)).
fof(d7_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0)), file('out/t107_xboole_1__17_premises__rerun.E_input', d7_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t107_xboole_1__17_premises__rerun.E_input', commutativity_k3_xboole_0)).
fof(l32_xboole_1, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('out/t107_xboole_1__17_premises__rerun.E_input', l32_xboole_1)).
fof(t107_xboole_1, conjecture, (![X1]:![X2]:![X3]:(r1_tarski(X1,k5_xboole_0(X2,X3))<=>(r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,k3_xboole_0(X2,X3))))), file('out/t107_xboole_1__17_premises__rerun.E_input', t107_xboole_1)).
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t107_xboole_1__17_premises__rerun.E_input', t4_xboole_0)).
fof(t106_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3)))), file('out/t107_xboole_1__17_premises__rerun.E_input', t106_xboole_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t107_xboole_1__17_premises__rerun.E_input', d3_tarski)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t107_xboole_1__17_premises__rerun.E_input', t48_xboole_1)).
fof(t10_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2)))), file('out/t107_xboole_1__17_premises__rerun.E_input', t10_xboole_1)).
fof(c_0_18, plain, (![X3]:![X4]:k4_xboole_0(X3,k2_xboole_0(X3,X4))=k1_xboole_0), inference(variable_rename,[status(thm)],[t46_xboole_1])).
fof(c_0_19, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_20, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_21, plain, (![X3]:![X4]:k3_xboole_0(X3,k2_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t21_xboole_1])).
cnf(c_0_22,plain,(k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_27, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_28,plain,(k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_24])).
cnf(c_0_29,plain,(k3_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_23]), c_0_24])).
cnf(c_0_30,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, (![X2]:k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_32,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34,plain,(k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_30, c_0_24])).
cnf(c_0_35,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=k5_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_38, plain, (![X3]:![X4]:r1_xboole_0(k3_xboole_0(X3,X4),k5_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[l97_xboole_1])).
cnf(c_0_39,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_33]), c_0_37])).
cnf(c_0_40,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_41,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[c_0_36, c_0_39])).
fof(c_0_42, plain, (![X4]:![X5]:![X6]:k3_xboole_0(X4,k4_xboole_0(X5,X6))=k4_xboole_0(k3_xboole_0(X4,X5),X6)), inference(variable_rename,[status(thm)],[t49_xboole_1])).
fof(c_0_43, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k3_xboole_0(X3,X4)=k1_xboole_0)&(k3_xboole_0(X3,X4)!=k1_xboole_0|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])])])).
cnf(c_0_44,plain,(r1_xboole_0(k3_xboole_0(X1,k5_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45,plain,(k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_47, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_48, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|k4_xboole_0(X3,X4)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])])])).
fof(c_0_49, negated_conjecture, (~(![X1]:![X2]:![X3]:(r1_tarski(X1,k5_xboole_0(X2,X3))<=>(r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,k3_xboole_0(X2,X3)))))), inference(assume_negation,[status(cth)],[t107_xboole_1])).
fof(c_0_50, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk4_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
cnf(c_0_51,plain,(r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_44, c_0_29])).
cnf(c_0_52,plain,(k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))=k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(k3_xboole_0(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_24]), c_0_24])).
cnf(c_0_53,plain,(k3_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_40])).
cnf(c_0_54,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_55,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_56, negated_conjecture, (((~r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))|(~r1_tarski(esk1_0,k2_xboole_0(esk2_0,esk3_0))|~r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))))&((r1_tarski(esk1_0,k2_xboole_0(esk2_0,esk3_0))|r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0)))&(r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))|r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])])).
cnf(c_0_57,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_58, plain, (![X4]:![X5]:![X6]:((r1_tarski(X4,X5)|~r1_tarski(X4,k4_xboole_0(X5,X6)))&(r1_xboole_0(X4,X6)|~r1_tarski(X4,k4_xboole_0(X5,X6))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
fof(c_0_59, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_60,plain,(r1_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,X1))))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61,plain,(k3_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X2,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_41]), c_0_54])).
cnf(c_0_62,plain,(k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_55, c_0_24])).
cnf(c_0_63,negated_conjecture,(r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))|r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64,plain,(~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_57, c_0_54])).
cnf(c_0_65,plain,(r2_hidden(esk4_2(X1,X2),k3_xboole_0(X1,X2))|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_66,plain,(r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_67,plain,(r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_69, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_70,plain,(r1_xboole_0(k3_xboole_0(X1,k5_xboole_0(X1,X2)),k3_xboole_0(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_37])).
cnf(c_0_71,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_62]), c_0_37])).
cnf(c_0_72,negated_conjecture,(k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))=k1_xboole_0|r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_46, c_0_63])).
cnf(c_0_73,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74,plain,(r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_75,plain,(r1_tarski(X1,X2)|~r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_66, c_0_24])).
cnf(c_0_76,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_77,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78,negated_conjecture,(r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))|r1_tarski(esk1_0,k2_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_79,negated_conjecture,(~r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))|~r1_tarski(esk1_0,k2_xboole_0(esk2_0,esk3_0))|~r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_80,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,k5_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_70, c_0_41])).
cnf(c_0_81,negated_conjecture,(k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))=k1_xboole_0|k3_xboole_0(esk1_0,k5_xboole_0(esk2_0,esk3_0))=esk1_0), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_82,plain,(r1_xboole_0(X1,X2)|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_74])).
fof(c_0_83, plain, (![X4]:![X5]:![X6]:(~r1_tarski(X4,X5)|r1_tarski(X4,k2_xboole_0(X6,X5)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])])])).
cnf(c_0_84,plain,(r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_85,plain,(k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_24]), c_0_24])).
cnf(c_0_86,negated_conjecture,(r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))|r1_tarski(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_23]), c_0_24])).
cnf(c_0_87,negated_conjecture,(~r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))|~r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))|~r1_tarski(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_23]), c_0_24])).
cnf(c_0_88,negated_conjecture,(r1_xboole_0(k3_xboole_0(esk2_0,esk3_0),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82])).
cnf(c_0_89,plain,(r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_90,plain,(k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,X2)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_91,plain,(r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92,negated_conjecture,(r1_tarski(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk2_0,esk3_0))))|r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(rw,[status(thm)],[c_0_86, c_0_54])).
cnf(c_0_93,negated_conjecture,(~r1_tarski(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk2_0,esk3_0))))|~r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))|~r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))), inference(rw,[status(thm)],[c_0_87, c_0_54])).
cnf(c_0_94,negated_conjecture,(r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_73, c_0_88])).
cnf(c_0_95,plain,(r1_tarski(X1,k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,X2))))|~r1_tarski(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_23]), c_0_24])).
cnf(c_0_96,plain,(k3_xboole_0(k5_xboole_0(X1,X2),k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X1,X2))))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_53]), c_0_37]), c_0_32])).
cnf(c_0_97,plain,(k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_90]), c_0_37])).
cnf(c_0_98,plain,(r1_tarski(k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3),X1)), inference(spm,[status(thm)],[c_0_75, c_0_91])).
cnf(c_0_99,negated_conjecture,(k3_xboole_0(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk2_0,esk3_0))))=esk1_0|r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_71, c_0_92])).
cnf(c_0_100,negated_conjecture,(~r1_tarski(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk2_0,esk3_0))))|~r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_94])])).
cnf(c_0_101,plain,(r1_tarski(X1,k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X2,X3))))|~r1_tarski(X1,k5_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_32]), c_0_97]), c_0_33]), c_0_32]), c_0_32]), c_0_37])).
cnf(c_0_102,plain,(r1_tarski(k3_xboole_0(k3_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_98, c_0_85])).
cnf(c_0_103,negated_conjecture,(k3_xboole_0(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k5_xboole_0(k3_xboole_0(esk2_0,esk3_0),k3_xboole_0(k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk2_0,esk3_0))),X1)))))=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,X1))|r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_99]), c_0_32]), c_0_32])).
cnf(c_0_104,negated_conjecture,(~r1_tarski(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_105,plain,(r1_tarski(k3_xboole_0(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_102, c_0_54])).
cnf(c_0_106,plain,(k3_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X1,X2))))=X1), inference(spm,[status(thm)],[c_0_29, c_0_54])).
cnf(c_0_107,plain,(r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_108,negated_conjecture,(k3_xboole_0(esk1_0,k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k5_xboole_0(k3_xboole_0(esk2_0,esk3_0),k3_xboole_0(k5_xboole_0(esk2_0,k5_xboole_0(esk3_0,k3_xboole_0(esk2_0,esk3_0))),X1)))))=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,X1))), inference(sr,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_109,plain,(k5_xboole_0(X1,k3_xboole_0(X2,X1))=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_62, c_0_54])).
cnf(c_0_110,negated_conjecture,(k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_88]), c_0_54])).
cnf(c_0_111,plain,(r1_tarski(k3_xboole_0(X1,X2),k5_xboole_0(X1,k5_xboole_0(X3,k3_xboole_0(X1,X3))))), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_112,plain,(r1_tarski(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_107, c_0_24])).
cnf(c_0_113,negated_conjecture,(k3_xboole_0(esk1_0,k5_xboole_0(esk2_0,esk3_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_37]), c_0_110]), c_0_37]), c_0_111])])).
cnf(c_0_114,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_33])]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 115
# Proof object clause steps            : 78
# Proof object formula steps           : 37
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 24
# Processed clauses                    : 11895
# ...of these trivial                  : 3038
# ...subsumed                          : 7217
# ...remaining for further processing  : 1640
# Other redundant clauses eliminated   : 279
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 93
# Backward-rewritten                   : 59
# Generated clauses                    : 208178
# ...of the previous two non-trivial   : 97616
# Contextual simplify-reflections      : 385
# Paramodulations                      : 207877
# Factorizations                       : 0
# Equation resolutions                 : 279
# Current number of processed clauses  : 1440
#    Positive orientable unit clauses  : 917
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 44
#    Non-unit-clauses                  : 474
# Current number of unprocessed clauses: 75353
# ...number of literals in the above   : 99930
# Current number of archived formulas  : 0
# Current number of archived clauses   : 200
# Clause-clause subsumption calls (NU) : 79659
# Rec. Clause-clause subsumption calls : 74640
# Non-unit clause-clause subsumptions  : 3366
# Unit Clause-clause subsumption calls : 4920
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6817
# BW rewrite match successes           : 194
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2591655

# -------------------------------------------------
# User time                : 2.048 s
# System time              : 0.088 s
# Total time               : 2.136 s
# Maximum resident set size: 90688 pages
