# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_waybel_2, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&l1_waybel_0(X3,X1))=>![X4]:((v6_waybel_0(X4,X1)&l1_waybel_0(X4,X1))=>(X4=k3_waybel_2(X1,X2,X3)<=>(g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))&![X5]:(m1_subset_1(X5,u1_struct_0(X4))=>?[X6]:(m1_subset_1(X6,u1_struct_0(X1))&(X6=k1_funct_1(u1_waybel_0(X1,X3),X5)&k1_funct_1(u1_waybel_0(X1,X4),X5)=k11_lattice3(X1,X2,X6))))))))))), file('out/t10_waybel_9__7_premises__rerun.E_input', d3_waybel_2)).
fof(dt_k3_waybel_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&(m1_subset_1(X2,u1_struct_0(X1))&(~(v2_struct_0(X3))&l1_waybel_0(X3,X1))))=>(v6_waybel_0(k3_waybel_2(X1,X2,X3),X1)&l1_waybel_0(k3_waybel_2(X1,X2,X3),X1)))), file('out/t10_waybel_9__7_premises__rerun.E_input', dt_k3_waybel_2)).
fof(dt_l1_waybel_0, axiom, (![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2)))), file('out/t10_waybel_9__7_premises__rerun.E_input', dt_l1_waybel_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t10_waybel_9__7_premises__rerun.E_input', dt_l1_orders_2)).
fof(t10_waybel_9, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>(~(v2_struct_0(k3_waybel_2(X1,X2,X3)))&(v4_orders_2(k3_waybel_2(X1,X2,X3))&(v7_waybel_0(k3_waybel_2(X1,X2,X3))&l1_waybel_0(k3_waybel_2(X1,X2,X3),X1)))))))), file('out/t10_waybel_9__7_premises__rerun.E_input', t10_waybel_9)).
fof(l7_yellow_6, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(v4_orders_2(X1)<=>v4_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))))), file('out/t10_waybel_9__7_premises__rerun.E_input', l7_yellow_6)).
fof(fc5_waybel_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&(m1_subset_1(X2,u1_struct_0(X1))&(~(v2_struct_0(X3))&l1_waybel_0(X3,X1))))=>(~(v2_struct_0(k3_waybel_2(X1,X2,X3)))&v6_waybel_0(k3_waybel_2(X1,X2,X3),X1)))), file('out/t10_waybel_9__7_premises__rerun.E_input', fc5_waybel_2)).
fof(fc6_waybel_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&(m1_subset_1(X2,u1_struct_0(X1))&(~(v2_struct_0(X3))&(v7_waybel_0(X3)&l1_waybel_0(X3,X1)))))=>(v6_waybel_0(k3_waybel_2(X1,X2,X3),X1)&v7_waybel_0(k3_waybel_2(X1,X2,X3))))), file('out/t10_waybel_9__7_premises__rerun.E_input', fc6_waybel_2)).
fof(c_0_8, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X14]:(((((((g1_orders_2(u1_struct_0(X10),u1_orders_2(X10))=g1_orders_2(u1_struct_0(X9),u1_orders_2(X9))|X10!=k3_waybel_2(X7,X8,X9))|(~v6_waybel_0(X10,X7)|~l1_waybel_0(X10,X7)))|(v2_struct_0(X9)|~l1_waybel_0(X9,X7)))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7)))&(((((((m1_subset_1(esk1_5(X7,X8,X9,X10,X11),u1_struct_0(X7))|~m1_subset_1(X11,u1_struct_0(X10)))|X10!=k3_waybel_2(X7,X8,X9))|(~v6_waybel_0(X10,X7)|~l1_waybel_0(X10,X7)))|(v2_struct_0(X9)|~l1_waybel_0(X9,X7)))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7)))&(((((((esk1_5(X7,X8,X9,X10,X11)=k1_funct_1(u1_waybel_0(X7,X9),X11)|~m1_subset_1(X11,u1_struct_0(X10)))|X10!=k3_waybel_2(X7,X8,X9))|(~v6_waybel_0(X10,X7)|~l1_waybel_0(X10,X7)))|(v2_struct_0(X9)|~l1_waybel_0(X9,X7)))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7)))&((((((k1_funct_1(u1_waybel_0(X7,X10),X11)=k11_lattice3(X7,X8,esk1_5(X7,X8,X9,X10,X11))|~m1_subset_1(X11,u1_struct_0(X10)))|X10!=k3_waybel_2(X7,X8,X9))|(~v6_waybel_0(X10,X7)|~l1_waybel_0(X10,X7)))|(v2_struct_0(X9)|~l1_waybel_0(X9,X7)))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7))))))&(((((((m1_subset_1(esk2_4(X7,X8,X9,X10),u1_struct_0(X10))|g1_orders_2(u1_struct_0(X10),u1_orders_2(X10))!=g1_orders_2(u1_struct_0(X9),u1_orders_2(X9)))|X10=k3_waybel_2(X7,X8,X9))|(~v6_waybel_0(X10,X7)|~l1_waybel_0(X10,X7)))|(v2_struct_0(X9)|~l1_waybel_0(X9,X7)))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7)))&(((((((~m1_subset_1(X14,u1_struct_0(X7))|(X14!=k1_funct_1(u1_waybel_0(X7,X9),esk2_4(X7,X8,X9,X10))|k1_funct_1(u1_waybel_0(X7,X10),esk2_4(X7,X8,X9,X10))!=k11_lattice3(X7,X8,X14)))|g1_orders_2(u1_struct_0(X10),u1_orders_2(X10))!=g1_orders_2(u1_struct_0(X9),u1_orders_2(X9)))|X10=k3_waybel_2(X7,X8,X9))|(~v6_waybel_0(X10,X7)|~l1_waybel_0(X10,X7)))|(v2_struct_0(X9)|~l1_waybel_0(X9,X7)))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_waybel_2])])])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:((v6_waybel_0(k3_waybel_2(X4,X5,X6),X4)|((v2_struct_0(X4)|~l1_orders_2(X4))|(~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X6)|~l1_waybel_0(X6,X4)))))&(l1_waybel_0(k3_waybel_2(X4,X5,X6),X4)|((v2_struct_0(X4)|~l1_orders_2(X4))|(~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X6)|~l1_waybel_0(X6,X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_waybel_2])])])])).
fof(c_0_10, plain, (![X3]:![X4]:(~l1_struct_0(X3)|(~l1_waybel_0(X4,X3)|l1_orders_2(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])])])).
fof(c_0_11, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_12, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>(~(v2_struct_0(k3_waybel_2(X1,X2,X3)))&(v4_orders_2(k3_waybel_2(X1,X2,X3))&(v7_waybel_0(k3_waybel_2(X1,X2,X3))&l1_waybel_0(k3_waybel_2(X1,X2,X3),X1))))))))), inference(assume_negation,[status(cth)],[t10_waybel_9])).
fof(c_0_13, plain, (![X2]:(((~v4_orders_2(X2)|v4_orders_2(g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))))|(v2_struct_0(X2)|~l1_orders_2(X2)))&((~v4_orders_2(g1_orders_2(u1_struct_0(X2),u1_orders_2(X2)))|v4_orders_2(X2))|(v2_struct_0(X2)|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l7_yellow_6])])])])).
cnf(c_0_14,plain,(v2_struct_0(X1)|v2_struct_0(X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~l1_waybel_0(X4,X1)|~v6_waybel_0(X4,X1)|X4!=k3_waybel_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v2_struct_0(X2)|l1_waybel_0(k3_waybel_2(X2,X3,X1),X2)|~l1_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v6_waybel_0(k3_waybel_2(X2,X3,X1),X2)|~l1_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,plain,(l1_orders_2(X1)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:((~v2_struct_0(k3_waybel_2(X4,X5,X6))|((v2_struct_0(X4)|~l1_orders_2(X4))|(~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X6)|~l1_waybel_0(X6,X4)))))&(v6_waybel_0(k3_waybel_2(X4,X5,X6),X4)|((v2_struct_0(X4)|~l1_orders_2(X4))|(~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X6)|~l1_waybel_0(X6,X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_waybel_2])])])])).
fof(c_0_20, negated_conjecture, (((~v2_struct_0(esk3_0)&l1_orders_2(esk3_0))&(m1_subset_1(esk4_0,u1_struct_0(esk3_0))&((~v2_struct_0(esk5_0)&(v4_orders_2(esk5_0)&(v7_waybel_0(esk5_0)&l1_waybel_0(esk5_0,esk3_0))))&(v2_struct_0(k3_waybel_2(esk3_0,esk4_0,esk5_0))|(~v4_orders_2(k3_waybel_2(esk3_0,esk4_0,esk5_0))|(~v7_waybel_0(k3_waybel_2(esk3_0,esk4_0,esk5_0))|~l1_waybel_0(k3_waybel_2(esk3_0,esk4_0,esk5_0),esk3_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:((v6_waybel_0(k3_waybel_2(X4,X5,X6),X4)|((v2_struct_0(X4)|~l1_orders_2(X4))|(~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X4))))))&(v7_waybel_0(k3_waybel_2(X4,X5,X6))|((v2_struct_0(X4)|~l1_orders_2(X4))|(~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X4)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_waybel_2])])])])).
cnf(c_0_22,plain,(v2_struct_0(X1)|v4_orders_2(X1)|~l1_orders_2(X1)|~v4_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,plain,(g1_orders_2(u1_struct_0(k3_waybel_2(X1,X2,X3)),u1_orders_2(k3_waybel_2(X1,X2,X3)))=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|v2_struct_0(X3)|v2_struct_0(X1)|~l1_waybel_0(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15]), c_0_16])).
cnf(c_0_24,plain,(l1_orders_2(k3_waybel_2(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X1)|~l1_waybel_0(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_15]), c_0_18])).
cnf(c_0_25,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_struct_0(k3_waybel_2(X2,X3,X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,negated_conjecture,(l1_waybel_0(esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,negated_conjecture,(v2_struct_0(k3_waybel_2(esk3_0,esk4_0,esk5_0))|~l1_waybel_0(k3_waybel_2(esk3_0,esk4_0,esk5_0),esk3_0)|~v7_waybel_0(k3_waybel_2(esk3_0,esk4_0,esk5_0))|~v4_orders_2(k3_waybel_2(esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v7_waybel_0(k3_waybel_2(X2,X3,X1))|~l1_waybel_0(X1,X2)|~v7_waybel_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,negated_conjecture,(v7_waybel_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31,negated_conjecture,(l1_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33,negated_conjecture,(~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_34,plain,(v4_orders_2(k3_waybel_2(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X1)|~v4_orders_2(g1_orders_2(u1_struct_0(X3),u1_orders_2(X3)))|~l1_waybel_0(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_35,plain,(v2_struct_0(X1)|v4_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36,negated_conjecture,(l1_orders_2(esk5_0)|~l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_17, c_0_26])).
cnf(c_0_37,negated_conjecture,(v2_struct_0(k3_waybel_2(esk3_0,esk4_0,esk5_0))|~v4_orders_2(k3_waybel_2(esk3_0,esk4_0,esk5_0))|~l1_waybel_0(k3_waybel_2(esk3_0,esk4_0,esk5_0),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_26]), c_0_30]), c_0_31])]), c_0_32]), c_0_33])).
cnf(c_0_38,plain,(v4_orders_2(k3_waybel_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X3)|~v4_orders_2(X3)|~l1_waybel_0(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39,negated_conjecture,(v4_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_40,negated_conjecture,(l1_orders_2(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_18]), c_0_31])])).
cnf(c_0_41,negated_conjecture,(v2_struct_0(k3_waybel_2(esk3_0,esk4_0,esk5_0))|~l1_waybel_0(k3_waybel_2(esk3_0,esk4_0,esk5_0),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_26]), c_0_30]), c_0_31]), c_0_40])]), c_0_32]), c_0_33])).
cnf(c_0_42,negated_conjecture,(v2_struct_0(k3_waybel_2(esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_15]), c_0_26]), c_0_30]), c_0_31])]), c_0_32]), c_0_33])).
cnf(c_0_43,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_42]), c_0_26]), c_0_30]), c_0_31])]), c_0_32]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 27
# Proof object formula steps           : 17
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 57
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 55
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 2
# Generated clauses                    : 21
# ...of the previous two non-trivial   : 19
# Contextual simplify-reflections      : 14
# Paramodulations                      : 15
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 15
# Current number of unprocessed clauses: 8
# ...number of literals in the above   : 79
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 920
# Rec. Clause-clause subsumption calls : 48
# Non-unit clause-clause subsumptions  : 18
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3376

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 43108 pages
