# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t114_xboole_1, conjecture, (![X1]:![X2]:![X3]:![X4]:((r1_xboole_0(X1,X4)&(r1_xboole_0(X2,X4)&r1_xboole_0(X3,X4)))=>r1_xboole_0(k2_xboole_0(k2_xboole_0(X1,X2),X3),X4))), file('out/t114_xboole_1__8_premises__rerun.E_input', t114_xboole_1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t114_xboole_1__8_premises__rerun.E_input', t98_xboole_1)).
fof(t4_xboole_1, axiom, (![X1]:![X2]:![X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t114_xboole_1__8_premises__rerun.E_input', t4_xboole_1)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t114_xboole_1__8_premises__rerun.E_input', t91_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1))), file('out/t114_xboole_1__8_premises__rerun.E_input', symmetry_r1_xboole_0)).
fof(t83_xboole_1, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1)), file('out/t114_xboole_1__8_premises__rerun.E_input', t83_xboole_1)).
fof(t53_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), file('out/t114_xboole_1__8_premises__rerun.E_input', t53_xboole_1)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t114_xboole_1__8_premises__rerun.E_input', t48_xboole_1)).
fof(t47_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), file('out/t114_xboole_1__8_premises__rerun.E_input', t47_xboole_1)).
fof(c_0_9, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((r1_xboole_0(X1,X4)&(r1_xboole_0(X2,X4)&r1_xboole_0(X3,X4)))=>r1_xboole_0(k2_xboole_0(k2_xboole_0(X1,X2),X3),X4)))), inference(assume_negation,[status(cth)],[t114_xboole_1])).
fof(c_0_10, negated_conjecture, (((r1_xboole_0(esk1_0,esk4_0)&(r1_xboole_0(esk2_0,esk4_0)&r1_xboole_0(esk3_0,esk4_0)))&~r1_xboole_0(k2_xboole_0(k2_xboole_0(esk1_0,esk2_0),esk3_0),esk4_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
fof(c_0_11, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:k2_xboole_0(k2_xboole_0(X4,X5),X6)=k2_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t4_xboole_1])).
cnf(c_0_13,negated_conjecture,(~r1_xboole_0(k2_xboole_0(k2_xboole_0(esk1_0,esk2_0),esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_16,plain,(k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|r1_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_18,negated_conjecture,(~r1_xboole_0(k5_xboole_0(k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)),k4_xboole_0(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_14])).
cnf(c_0_19,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,plain,(k5_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k4_xboole_0(X3,k5_xboole_0(X1,k4_xboole_0(X2,X1))))=k5_xboole_0(X1,k4_xboole_0(k5_xboole_0(X2,k4_xboole_0(X3,X2)),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_14]), c_0_14]), c_0_14]), c_0_14])).
fof(c_0_21, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k4_xboole_0(X3,X4)=X3)&(k4_xboole_0(X3,X4)!=X3|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])])])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:k4_xboole_0(X4,k2_xboole_0(X5,X6))=k3_xboole_0(k4_xboole_0(X4,X5),k4_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[t53_xboole_1])).
fof(c_0_23, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_24,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,negated_conjecture,(r1_xboole_0(esk1_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_26, plain, (![X3]:![X4]:k4_xboole_0(X3,k3_xboole_0(X3,X4))=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_27,negated_conjecture,(r1_xboole_0(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_28,negated_conjecture,(~r1_xboole_0(k5_xboole_0(esk1_0,k5_xboole_0(k4_xboole_0(esk2_0,esk1_0),k4_xboole_0(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0))))),esk4_0)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_29,plain,(k5_xboole_0(X1,k5_xboole_0(k4_xboole_0(X2,X1),k4_xboole_0(X3,k5_xboole_0(X1,k4_xboole_0(X2,X1)))))=k5_xboole_0(X1,k4_xboole_0(k5_xboole_0(X2,k4_xboole_0(X3,X2)),X1))), inference(rw,[status(thm)],[c_0_20, c_0_19])).
cnf(c_0_30,plain,(r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31,plain,(k4_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33,plain,(k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34,negated_conjecture,(r1_xboole_0(esk4_0,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_35,plain,(k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36,negated_conjecture,(r1_xboole_0(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_27])).
cnf(c_0_37,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_38,negated_conjecture,(~r1_xboole_0(k5_xboole_0(esk1_0,k4_xboole_0(k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)),esk1_0)),esk4_0)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39,plain,(r1_xboole_0(X1,X2)|k4_xboole_0(X2,X1)!=X2), inference(spm,[status(thm)],[c_0_24, c_0_30])).
cnf(c_0_40,plain,(k4_xboole_0(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))=k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_14]), c_0_32])).
cnf(c_0_41,negated_conjecture,(k4_xboole_0(esk4_0,esk1_0)=esk4_0), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_42,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_35, c_0_32])).
cnf(c_0_43,negated_conjecture,(k4_xboole_0(esk4_0,esk2_0)=esk4_0), inference(spm,[status(thm)],[c_0_33, c_0_36])).
cnf(c_0_44,negated_conjecture,(r1_xboole_0(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_37])).
cnf(c_0_45,negated_conjecture,(k4_xboole_0(esk4_0,k5_xboole_0(esk1_0,k4_xboole_0(k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)),esk1_0)))!=esk4_0), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46,negated_conjecture,(k4_xboole_0(esk4_0,k5_xboole_0(esk1_0,k4_xboole_0(X1,esk1_0)))=k4_xboole_0(esk4_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_47,negated_conjecture,(k4_xboole_0(esk4_0,k5_xboole_0(esk2_0,k4_xboole_0(X1,esk2_0)))=k4_xboole_0(esk4_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_43]), c_0_42])).
cnf(c_0_48,negated_conjecture,(k4_xboole_0(esk4_0,esk3_0)=esk4_0), inference(spm,[status(thm)],[c_0_33, c_0_44])).
cnf(c_0_49,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46]), c_0_47]), c_0_48])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 31
# Proof object formula steps           : 19
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 11
# Processed clauses                    : 119
# ...of these trivial                  : 11
# ...subsumed                          : 1
# ...remaining for further processing  : 107
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 1345
# ...of the previous two non-trivial   : 1182
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1345
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 93
#    Positive orientable unit clauses  : 74
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 1085
# ...number of literals in the above   : 1268
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 15
# Rec. Clause-clause subsumption calls : 15
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 35878

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.004 s
# Total time               : 0.032 s
# Maximum resident set size: 45308 pages
