# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k7_relset_1, axiom, (![X1]:![X2]:![X3]:![X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2)))), file('out/t115_funct_2__12_premises__rerun.E_input', dt_k7_relset_1)).
fof(redefinition_k7_relset_1, axiom, (![X1]:![X2]:![X3]:![X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4))), file('out/t115_funct_2__12_premises__rerun.E_input', redefinition_k7_relset_1)).
fof(t115_funct_2, conjecture, (![X1]:![X2]:![X3]:![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:~((r2_hidden(X6,X1)&(r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6)))))))), file('out/t115_funct_2__12_premises__rerun.E_input', t115_funct_2)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t115_funct_2__12_premises__rerun.E_input', t4_subset)).
fof(t8_funct_1, axiom, (![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1))))), file('out/t115_funct_2__12_premises__rerun.E_input', t8_funct_1)).
fof(t143_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2)))))), file('out/t115_funct_2__12_premises__rerun.E_input', t143_relat_1)).
fof(cc1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('out/t115_funct_2__12_premises__rerun.E_input', cc1_relset_1)).
fof(redefinition_k1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3))), file('out/t115_funct_2__12_premises__rerun.E_input', redefinition_k1_relset_1)).
fof(t5_subset, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3))))), file('out/t115_funct_2__12_premises__rerun.E_input', t5_subset)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t115_funct_2__12_premises__rerun.E_input', t2_subset)).
fof(d1_funct_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0)))))), file('out/t115_funct_2__12_premises__rerun.E_input', d1_funct_2)).
fof(t7_ordinal1, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1)))), file('out/t115_funct_2__12_premises__rerun.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t115_funct_2__12_premises__rerun.E_input', t2_xboole_1)).
fof(c_0_13, plain, (![X5]:![X6]:![X7]:![X8]:(~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|m1_subset_1(k7_relset_1(X5,X6,X7,X8),k1_zfmisc_1(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])])])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X8]:(~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|k7_relset_1(X5,X6,X7,X8)=k9_relat_1(X7,X8))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])])])).
fof(c_0_15, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:~((r2_hidden(X6,X1)&(r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6))))))))), inference(assume_negation,[status(cth)],[t115_funct_2])).
cnf(c_0_16,plain,(m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, negated_conjecture, (![X12]:((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,esk2_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))&(~r2_hidden(X12,esk1_0)|(~r2_hidden(X12,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X12)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_20,plain,(m1_subset_1(k9_relat_1(X1,X2),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(k9_relat_1(esk4_0,X1),k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24,negated_conjecture,(r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:((((r2_hidden(X4,k1_relat_1(X6))|~r2_hidden(k4_tarski(X4,X5),X6))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((X5=k1_funct_1(X6,X4)|~r2_hidden(k4_tarski(X4,X5),X6))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&(((~r2_hidden(X4,k1_relat_1(X6))|X5!=k1_funct_1(X6,X4))|r2_hidden(k4_tarski(X4,X5),X6))|(~v1_relat_1(X6)|~v1_funct_1(X6))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_26, plain, (![X5]:![X6]:![X7]:![X9]:((((r2_hidden(esk6_3(X5,X6,X7),k1_relat_1(X7))|~r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7))&(((r2_hidden(k4_tarski(esk6_3(X5,X6,X7),X5),X7)|~r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7))&((r2_hidden(esk6_3(X5,X6,X7),X6)|~r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7))))&(((~r2_hidden(X9,k1_relat_1(X7))|(~r2_hidden(k4_tarski(X9,X5),X7)|~r2_hidden(X9,X6)))|r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])])])).
fof(c_0_27, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_28, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|k1_relset_1(X4,X5,X6)=k1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_29, plain, (![X4]:![X5]:![X6]:(~r2_hidden(X4,X5)|(~m1_subset_1(X5,k1_zfmisc_1(X6))|~v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])])])).
fof(c_0_30, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_31,negated_conjecture,(m1_subset_1(X1,esk2_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_32,negated_conjecture,(r2_hidden(esk5_0,k9_relat_1(esk4_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_17]), c_0_21])])).
cnf(c_0_33,plain,(X3=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(r2_hidden(k4_tarski(esk6_3(X2,X3,X1),X2),X1)|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_36, plain, (![X4]:![X5]:![X6]:((((((~v1_funct_2(X6,X4,X5)|X4=k1_relset_1(X4,X5,X6))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X4!=k1_relset_1(X4,X5,X6)|v1_funct_2(X6,X4,X5))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))&((((~v1_funct_2(X6,X4,X5)|X4=k1_relset_1(X4,X5,X6))|X4!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X4!=k1_relset_1(X4,X5,X6)|v1_funct_2(X6,X4,X5))|X4!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))&(((((~v1_funct_2(X6,X4,X5)|X6=k1_xboole_0)|X4=k1_xboole_0)|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&((((X6!=k1_xboole_0|v1_funct_2(X6,X4,X5))|X4=k1_xboole_0)|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_37,plain,(k1_relset_1(X1,X2,X3)=k1_relat_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38,plain,(~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41,negated_conjecture,(esk5_0!=k1_funct_1(esk4_0,X1)|~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_42,plain,(k1_funct_1(X1,esk6_3(X2,X3,X1))=X2|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_43,negated_conjecture,(v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_35, c_0_21])).
cnf(c_0_44,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_45,plain,(X3=k1_xboole_0|X2=k1_relset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46,negated_conjecture,(k1_relset_1(esk1_0,esk2_0,esk4_0)=k1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_37, c_0_21])).
cnf(c_0_47,negated_conjecture,(v1_funct_2(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_48,negated_conjecture,(~v1_xboole_0(esk2_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(spm,[status(thm)],[c_0_38, c_0_23])).
cnf(c_0_49,negated_conjecture,(v1_xboole_0(esk2_0)|r2_hidden(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_50,negated_conjecture,(~r2_hidden(esk6_3(esk5_0,X1,esk4_0),esk1_0)|~r2_hidden(esk6_3(esk5_0,X1,esk4_0),esk3_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])])])).
cnf(c_0_51,plain,(r2_hidden(esk6_3(X2,X3,X1),X3)|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_52,plain,(r2_hidden(esk6_3(X2,X3,X1),k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_53,negated_conjecture,(k1_relat_1(esk4_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_21]), c_0_46]), c_0_47])])).
fof(c_0_54, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~r1_tarski(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_55, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_56,negated_conjecture,(r2_hidden(esk5_0,esk2_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57,negated_conjecture,(~r2_hidden(esk6_3(esk5_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_32]), c_0_43])])).
cnf(c_0_58,negated_conjecture,(esk2_0=k1_xboole_0|r2_hidden(esk6_3(X1,X2,esk4_0),esk1_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_43])])).
cnf(c_0_59,plain,(~r1_tarski(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_60,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61,negated_conjecture,(r2_hidden(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_56, c_0_32])).
cnf(c_0_62,negated_conjecture,(esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_32])])).
cnf(c_0_63,plain,(~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_64,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_62]), c_0_63]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 38
# Proof object formula steps           : 27
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 124
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 117
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 26
# Generated clauses                    : 148
# ...of the previous two non-trivial   : 147
# Contextual simplify-reflections      : 4
# Paramodulations                      : 142
# Factorizations                       : 0
# Equation resolutions                 : 7
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 59
# ...number of literals in the above   : 267
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 922
# Rec. Clause-clause subsumption calls : 289
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 36
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5280

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 44916 pages
