# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t117_funct_1__6_premises__rerun__rerun.E_input', d1_tarski)).
fof(d12_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:![X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(X5,k1_relat_1(X1))&(r2_hidden(X5,X2)&X4=k1_funct_1(X1,X5))))))), file('out/t117_funct_1__6_premises__rerun__rerun.E_input', d12_funct_1)).
fof(t117_funct_1, conjecture, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>k11_relat_1(X2,X1)=k1_tarski(k1_funct_1(X2,X1))))), file('out/t117_funct_1__6_premises__rerun__rerun.E_input', t117_funct_1)).
fof(d16_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2)))), file('out/t117_funct_1__6_premises__rerun__rerun.E_input', d16_relat_1)).
fof(c_0_4, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
fof(c_0_5, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:![X7]:![X8]:![X13]:((((((r2_hidden(esk2_4(X6,X7,X8,X9),k1_relat_1(X6))|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((r2_hidden(esk2_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&(((X9=k1_funct_1(X6,esk2_4(X6,X7,X8,X9))|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))))&((((~r2_hidden(X11,k1_relat_1(X6))|(~r2_hidden(X11,X7)|X9!=k1_funct_1(X6,X11)))|r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&((((~r2_hidden(esk3_3(X6,X7,X8),X8)|(~r2_hidden(X13,k1_relat_1(X6))|(~r2_hidden(X13,X7)|esk3_3(X6,X7,X8)!=k1_funct_1(X6,X13))))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((r2_hidden(esk4_3(X6,X7,X8),k1_relat_1(X6))|r2_hidden(esk3_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((r2_hidden(esk4_3(X6,X7,X8),X7)|r2_hidden(esk3_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&(((esk3_3(X6,X7,X8)=k1_funct_1(X6,esk4_3(X6,X7,X8))|r2_hidden(esk3_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_funct_1])])])])])])])).
cnf(c_0_6,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(r2_hidden(esk2_4(X1,X3,X2,X4),X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k9_relat_1(X1,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(X4=k1_funct_1(X1,esk2_4(X1,X3,X2,X4))|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k9_relat_1(X1,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(r2_hidden(esk2_4(X1,X2,k9_relat_1(X1,X2),X3),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,X2))), inference(er,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(k1_funct_1(X1,esk2_4(X1,X2,k9_relat_1(X1,X2),X3))=X3|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,X2))), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(esk2_4(X1,k1_tarski(X2),k9_relat_1(X1,k1_tarski(X2)),X3)=X2|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,k1_tarski(X2)))), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13,plain,(r2_hidden(X4,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k9_relat_1(X1,X3)|X4!=k1_funct_1(X1,X5)|~r2_hidden(X5,X3)|~r2_hidden(X5,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_14, negated_conjecture, (~(![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>k11_relat_1(X2,X1)=k1_tarski(k1_funct_1(X2,X1)))))), inference(assume_negation,[status(cth)],[t117_funct_1])).
cnf(c_0_15,plain,(k1_funct_1(X1,X2)=X3|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,k1_tarski(X2)))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16,plain,(r2_hidden(k1_funct_1(X1,X2),k9_relat_1(X1,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X2,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_13])])).
fof(c_0_17, negated_conjecture, (((v1_relat_1(esk6_0)&v1_funct_1(esk6_0))&(r2_hidden(esk5_0,k1_relat_1(esk6_0))&k11_relat_1(esk6_0,esk5_0)!=k1_tarski(k1_funct_1(esk6_0,esk5_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_18,plain,(X1=k1_tarski(X2)|esk1_2(X2,X1)=X2|r2_hidden(esk1_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_19,plain,(k1_funct_1(X1,X2)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k1_relat_1(X1))|~r2_hidden(X3,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20,negated_conjecture,(r2_hidden(esk5_0,k1_relat_1(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,negated_conjecture,(v1_relat_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(esk1_2(X1,k9_relat_1(X2,k1_tarski(X3)))=k1_funct_1(X2,X3)|esk1_2(X1,k9_relat_1(X2,k1_tarski(X3)))=X1|k9_relat_1(X2,k1_tarski(X3))=k1_tarski(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_15, c_0_18])).
cnf(c_0_24,negated_conjecture,(k1_funct_1(esk6_0,X1)=k1_funct_1(esk6_0,esk5_0)|~r2_hidden(esk5_0,k1_tarski(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])])).
cnf(c_0_25,negated_conjecture,(esk1_2(X1,k9_relat_1(esk6_0,k1_tarski(X2)))=k1_funct_1(esk6_0,X2)|esk1_2(X1,k9_relat_1(esk6_0,k1_tarski(X2)))=X1|k9_relat_1(esk6_0,k1_tarski(X2))=k1_tarski(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_21]), c_0_22])])).
cnf(c_0_26,negated_conjecture,(r2_hidden(k1_funct_1(esk6_0,X1),k9_relat_1(esk6_0,X2))|~r2_hidden(esk5_0,k1_tarski(X1))|~r2_hidden(esk5_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_24]), c_0_21]), c_0_22]), c_0_20])])).
cnf(c_0_27,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_28,plain,(X1=k1_tarski(X2)|esk1_2(X2,X1)!=X2|~r2_hidden(esk1_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_29,negated_conjecture,(esk1_2(k1_funct_1(esk6_0,X1),k9_relat_1(esk6_0,k1_tarski(X1)))=k1_funct_1(esk6_0,X1)|k1_tarski(k1_funct_1(esk6_0,X1))=k9_relat_1(esk6_0,k1_tarski(X1))), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_25])])).
cnf(c_0_30,negated_conjecture,(r2_hidden(k1_funct_1(esk6_0,esk5_0),k9_relat_1(esk6_0,X1))|~r2_hidden(esk5_0,k1_tarski(X2))|~r2_hidden(esk5_0,X1)), inference(spm,[status(thm)],[c_0_26, c_0_24])).
cnf(c_0_31,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_27])])).
cnf(c_0_32,negated_conjecture,(k1_tarski(k1_funct_1(esk6_0,X1))=k9_relat_1(esk6_0,k1_tarski(X1))|~r2_hidden(k1_funct_1(esk6_0,X1),k9_relat_1(esk6_0,k1_tarski(X1)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33,negated_conjecture,(r2_hidden(k1_funct_1(esk6_0,esk5_0),k9_relat_1(esk6_0,X1))|~r2_hidden(esk5_0,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,negated_conjecture,(k11_relat_1(esk6_0,esk5_0)!=k1_tarski(k1_funct_1(esk6_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35,negated_conjecture,(k1_tarski(k1_funct_1(esk6_0,esk5_0))=k9_relat_1(esk6_0,k1_tarski(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_31])])).
fof(c_0_36, plain, (![X3]:![X4]:(~v1_relat_1(X3)|k11_relat_1(X3,X4)=k9_relat_1(X3,k1_tarski(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])])])).
cnf(c_0_37,negated_conjecture,(k11_relat_1(esk6_0,esk5_0)!=k9_relat_1(esk6_0,k1_tarski(esk5_0))), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38,plain,(k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_39,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_22])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 31
# Proof object formula steps           : 9
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 1011
# ...of these trivial                  : 0
# ...subsumed                          : 672
# ...remaining for further processing  : 339
# Other redundant clauses eliminated   : 16
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 25
# Backward-rewritten                   : 12
# Generated clauses                    : 3878
# ...of the previous two non-trivial   : 3792
# Contextual simplify-reflections      : 638
# Paramodulations                      : 3860
# Factorizations                       : 4
# Equation resolutions                 : 16
# Current number of processed clauses  : 279
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 272
# Current number of unprocessed clauses: 2532
# ...number of literals in the above   : 17011
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 49622
# Rec. Clause-clause subsumption calls : 18222
# Non-unit clause-clause subsumptions  : 1267
# Unit Clause-clause subsumption calls : 123
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 117634

# -------------------------------------------------
# User time                : 0.140 s
# System time              : 0.000 s
# Total time               : 0.140 s
# Maximum resident set size: 45276 pages
