# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', t7_boole)).
fof(t110_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k7_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2)))))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', t110_relat_1)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', t2_subset)).
fof(existence_m1_subset_1, axiom, (![X1]:?[X2]:m1_subset_1(X2,X1)), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', existence_m1_subset_1)).
fof(t6_boole, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', t6_boole)).
fof(t119_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>~((~(X1=k1_xboole_0)&(r1_tarski(X1,k9_xtuple_0(X2))&k7_relat_1(X2,X1)=k1_xboole_0))))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', t119_relat_1)).
fof(d12_xtuple_0, axiom, (![X1]:![X2]:(X2=k9_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', d12_xtuple_0)).
fof(d5_tarski, axiom, (![X1]:![X2]:k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', d5_tarski)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(fc4_xtuple_0, axiom, (![X1]:(v1_xboole_0(X1)=>v1_xboole_0(k9_xtuple_0(X1)))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', fc4_xtuple_0)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', fc1_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t119_relat_1__50_premises__rerun__rerun.E_input', d3_tarski)).
fof(c_0_12, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_13, plain, (![X5]:![X6]:![X7]:![X9]:((((r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X7))|~r2_hidden(X5,k7_relat_1(X7,X6)))|~v1_relat_1(X7))&(((r2_hidden(k4_tarski(esk1_3(X5,X6,X7),X5),X7)|~r2_hidden(X5,k7_relat_1(X7,X6)))|~v1_relat_1(X7))&((r2_hidden(esk1_3(X5,X6,X7),X6)|~r2_hidden(X5,k7_relat_1(X7,X6)))|~v1_relat_1(X7))))&(((~r2_hidden(X9,k9_xtuple_0(X7))|(~r2_hidden(k4_tarski(X9,X5),X7)|~r2_hidden(X9,X6)))|r2_hidden(X5,k7_relat_1(X7,X6)))|~v1_relat_1(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t110_relat_1])])])])])])])).
fof(c_0_14, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_15, plain, (![X3]:m1_subset_1(esk2_1(X3),X3)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_16,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(r2_hidden(esk1_3(X2,X3,X1),X3)|~v1_relat_1(X1)|~r2_hidden(X2,k7_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(m1_subset_1(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_21,plain,(~v1_relat_1(X1)|~v1_xboole_0(X2)|~r2_hidden(X3,k7_relat_1(X1,X2))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22,plain,(v1_xboole_0(X1)|r2_hidden(esk2_1(X1),X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(v1_xboole_0(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25,plain,(k7_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_26,plain,(~v1_relat_1(X1)|~v1_xboole_0(X2)|~r2_hidden(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_21, c_0_25])).
fof(c_0_27, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>~((~(X1=k1_xboole_0)&(r1_tarski(X1,k9_xtuple_0(X2))&k7_relat_1(X2,X1)=k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t119_relat_1])).
fof(c_0_28, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk5_3(X5,X6,X7)),X5))|X6!=k9_xtuple_0(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k9_xtuple_0(X5)))&(((~r2_hidden(esk6_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk6_2(X5,X6),X11),X5))|X6=k9_xtuple_0(X5))&((r2_hidden(esk6_2(X5,X6),X6)|r2_hidden(k4_tarski(esk6_2(X5,X6),esk7_2(X5,X6)),X5))|X6=k9_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_xtuple_0])])])])])])])).
fof(c_0_29, plain, (![X3]:![X4]:k4_tarski(X3,X4)=k2_tarski(k2_tarski(X3,X4),k1_tarski(X3))), inference(variable_rename,[status(thm)],[d5_tarski])).
cnf(c_0_30,plain,(~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_xboole_0(X3)|~r2_hidden(X4,k7_relat_1(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_26, c_0_17])).
fof(c_0_31, negated_conjecture, ((v1_relat_1(esk4_0)&(esk3_0!=k1_xboole_0&(r1_tarski(esk3_0,k9_xtuple_0(esk4_0))&k7_relat_1(esk4_0,esk3_0)=k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_27])])])])).
cnf(c_0_32,plain,(r2_hidden(X3,X1)|X1!=k9_xtuple_0(X2)|~r2_hidden(k4_tarski(X3,X4),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33,plain,(k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34,plain,(v1_xboole_0(k7_relat_1(X1,k1_xboole_0))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_30, c_0_22])).
cnf(c_0_35,negated_conjecture,(v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36,plain,(r2_hidden(X2,k7_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X4,X3)|~r2_hidden(k4_tarski(X4,X2),X1)|~r2_hidden(X4,k9_xtuple_0(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_37,plain,(r2_hidden(X3,X1)|X1!=k9_xtuple_0(X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),X2)), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38,plain,(r2_hidden(k4_tarski(X3,esk5_3(X2,X1,X3)),X2)|X1!=k9_xtuple_0(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_39, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_40, plain, (![X2]:(~v1_xboole_0(X2)|v1_xboole_0(k9_xtuple_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_xtuple_0])])).
cnf(c_0_41,negated_conjecture,(v1_xboole_0(k7_relat_1(X1,k1_xboole_0))|~v1_relat_1(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_43,plain,(r2_hidden(X2,k7_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X4,X3)|~r2_hidden(X4,k9_xtuple_0(X1))|~r2_hidden(k2_tarski(k2_tarski(X4,X2),k1_tarski(X4)),X1)), inference(rw,[status(thm)],[c_0_36, c_0_33])).
cnf(c_0_44,plain,(r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),X2)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_45,plain,(r2_hidden(k2_tarski(k2_tarski(X3,esk5_3(X2,X1,X3)),k1_tarski(X3)),X2)|X1!=k9_xtuple_0(X2)|~r2_hidden(X3,X1)), inference(rw,[status(thm)],[c_0_38, c_0_33])).
cnf(c_0_46,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47,plain,(v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48,negated_conjecture,(v1_xboole_0(k7_relat_1(X1,k1_xboole_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49,plain,(X1=k9_xtuple_0(X2)|r2_hidden(k4_tarski(esk6_2(X2,X1),esk7_2(X2,X1)),X2)|r2_hidden(esk6_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_50,plain,(r2_hidden(X1,k7_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k2_tarski(k2_tarski(X4,X1),k1_tarski(X4)),X2)|~r2_hidden(X4,X3)), inference(csr,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51,plain,(r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk5_3(X2,k9_xtuple_0(X2),X1))),X2)|~r2_hidden(X1,k9_xtuple_0(X2))), inference(er,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_52,plain,(k9_xtuple_0(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_23, c_0_47])).
fof(c_0_53, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk8_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk8_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_54,negated_conjecture,(k7_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_48])).
cnf(c_0_55,plain,(X1=k9_xtuple_0(X2)|r2_hidden(esk6_2(X2,X1),X1)|r2_hidden(k2_tarski(k2_tarski(esk6_2(X2,X1),esk7_2(X2,X1)),k1_tarski(esk6_2(X2,X1))),X2)), inference(rw,[status(thm)],[c_0_49, c_0_33])).
cnf(c_0_56,plain,(r2_hidden(X1,k7_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k2_tarski(k1_tarski(X4),k2_tarski(X4,X1)),X2)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_50, c_0_46])).
cnf(c_0_57,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1))), inference(spm,[status(thm)],[c_0_16, c_0_51])).
cnf(c_0_58,plain,(k9_xtuple_0(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_52, c_0_42])).
cnf(c_0_59,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60,negated_conjecture,(r1_tarski(esk3_0,k9_xtuple_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_61,negated_conjecture,(~v1_relat_1(X1)|~r2_hidden(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_54]), c_0_42])])).
cnf(c_0_62,plain,(X1=k9_xtuple_0(X2)|r2_hidden(k2_tarski(k1_tarski(esk6_2(X2,X1)),k2_tarski(esk6_2(X2,X1),esk7_2(X2,X1))),X2)|r2_hidden(esk6_2(X2,X1),X1)), inference(rw,[status(thm)],[c_0_55, c_0_46])).
cnf(c_0_63,plain,(r2_hidden(esk5_3(X1,k9_xtuple_0(X1),X2),k7_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X2,k9_xtuple_0(X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_56, c_0_51])).
cnf(c_0_64,negated_conjecture,(k7_relat_1(esk4_0,esk3_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_65,plain,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_42])])).
cnf(c_0_66,negated_conjecture,(r2_hidden(X1,k9_xtuple_0(esk4_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_67,negated_conjecture,(X1=k1_xboole_0|r2_hidden(esk6_2(k1_xboole_0,X1),X1)|~v1_relat_1(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_58])).
cnf(c_0_68,negated_conjecture,(~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_35])]), c_0_65]), c_0_66])).
cnf(c_0_69,negated_conjecture,(X1=k1_xboole_0|r2_hidden(esk6_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_67, c_0_35])).
cnf(c_0_70,negated_conjecture,(esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_71,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 48
# Proof object formula steps           : 24
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 2771
# ...of these trivial                  : 0
# ...subsumed                          : 2430
# ...remaining for further processing  : 341
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 17
# Backward-rewritten                   : 1
# Generated clauses                    : 13732
# ...of the previous two non-trivial   : 11045
# Contextual simplify-reflections      : 1906
# Paramodulations                      : 13728
# Factorizations                       : 2
# Equation resolutions                 : 2
# Current number of processed clauses  : 299
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 287
# Current number of unprocessed clauses: 8041
# ...number of literals in the above   : 34151
# Current number of archived formulas  : 0
# Current number of archived clauses   : 41
# Clause-clause subsumption calls (NU) : 57889
# Rec. Clause-clause subsumption calls : 32175
# Non-unit clause-clause subsumptions  : 3583
# Unit Clause-clause subsumption calls : 150
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 205295

# -------------------------------------------------
# User time                : 0.168 s
# System time              : 0.000 s
# Total time               : 0.168 s
# Maximum resident set size: 45340 pages
