# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t11_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>(v10_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r1_orders_2(X1,k2_waybel_0(X1,X2,X3),k2_waybel_0(X1,X2,X5)))))))))), file('out/t11_waybel_0__2_premises__rerun.E_input', t11_waybel_0)).
fof(d13_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>(v10_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_waybel_0(X1,X2,a_3_0_waybel_0(X1,X2,X3))))))), file('out/t11_waybel_0__2_premises__rerun.E_input', d13_waybel_0)).
fof(s1_waybel_0__e1_34_1__waybel_0, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))&m1_subset_1(X3,u1_struct_0(X2))))=>(r1_waybel_0(X1,X2,a_3_0_waybel_0(X1,X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r1_orders_2(X1,k2_waybel_0(X1,X2,X3),k2_waybel_0(X1,X2,X5)))))))), file('out/t11_waybel_0__2_premises__rerun.E_input', s1_waybel_0__e1_34_1__waybel_0)).
fof(c_0_3, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>(v10_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r1_orders_2(X1,k2_waybel_0(X1,X2,X3),k2_waybel_0(X1,X2,X5))))))))))), inference(assume_negation,[status(cth)],[t11_waybel_0])).
fof(c_0_4, plain, (![X4]:![X5]:![X6]:((((~v10_waybel_0(X5,X4)|(~m1_subset_1(X6,u1_struct_0(X5))|r1_waybel_0(X4,X5,a_3_0_waybel_0(X4,X5,X6))))|(v2_struct_0(X5)|~l1_waybel_0(X5,X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((m1_subset_1(esk6_2(X4,X5),u1_struct_0(X5))|v10_waybel_0(X5,X4))|(v2_struct_0(X5)|~l1_waybel_0(X5,X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&(((~r1_waybel_0(X4,X5,a_3_0_waybel_0(X4,X5,esk6_2(X4,X5)))|v10_waybel_0(X5,X4))|(v2_struct_0(X5)|~l1_waybel_0(X5,X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_waybel_0])])])])])])])])).
fof(c_0_5, negated_conjecture, (![X9]:![X11]:![X13]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~v2_struct_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))&(((m1_subset_1(esk3_0,u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0))&(((m1_subset_1(esk4_1(X9),u1_struct_0(esk2_0))|~m1_subset_1(X9,u1_struct_0(esk2_0)))|~v10_waybel_0(esk2_0,esk1_0))&(((r1_orders_2(esk2_0,X9,esk4_1(X9))|~m1_subset_1(X9,u1_struct_0(esk2_0)))|~v10_waybel_0(esk2_0,esk1_0))&((~r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk3_0),k2_waybel_0(esk1_0,esk2_0,esk4_1(X9)))|~m1_subset_1(X9,u1_struct_0(esk2_0)))|~v10_waybel_0(esk2_0,esk1_0)))))&(((m1_subset_1(esk5_1(X11),u1_struct_0(esk2_0))|~m1_subset_1(X11,u1_struct_0(esk2_0)))|v10_waybel_0(esk2_0,esk1_0))&(((~m1_subset_1(X13,u1_struct_0(esk2_0))|(~r1_orders_2(esk2_0,esk5_1(X11),X13)|r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,X11),k2_waybel_0(esk1_0,esk2_0,X13))))|~m1_subset_1(X11,u1_struct_0(esk2_0)))|v10_waybel_0(esk2_0,esk1_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])])])).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X10]:![X11]:((((m1_subset_1(esk7_3(X6,X7,X8),u1_struct_0(X7))|~r1_waybel_0(X6,X7,a_3_0_waybel_0(X6,X7,X8)))|((v2_struct_0(X6)|~l1_orders_2(X6))|((v2_struct_0(X7)|~l1_waybel_0(X7,X6))|~m1_subset_1(X8,u1_struct_0(X7)))))&(((~m1_subset_1(X10,u1_struct_0(X7))|(~r1_orders_2(X7,esk7_3(X6,X7,X8),X10)|r1_orders_2(X6,k2_waybel_0(X6,X7,X8),k2_waybel_0(X6,X7,X10))))|~r1_waybel_0(X6,X7,a_3_0_waybel_0(X6,X7,X8)))|((v2_struct_0(X6)|~l1_orders_2(X6))|((v2_struct_0(X7)|~l1_waybel_0(X7,X6))|~m1_subset_1(X8,u1_struct_0(X7))))))&((((m1_subset_1(esk8_4(X6,X7,X8,X11),u1_struct_0(X7))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,a_3_0_waybel_0(X6,X7,X8)))|((v2_struct_0(X6)|~l1_orders_2(X6))|((v2_struct_0(X7)|~l1_waybel_0(X7,X6))|~m1_subset_1(X8,u1_struct_0(X7)))))&((((r1_orders_2(X7,X11,esk8_4(X6,X7,X8,X11))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,a_3_0_waybel_0(X6,X7,X8)))|((v2_struct_0(X6)|~l1_orders_2(X6))|((v2_struct_0(X7)|~l1_waybel_0(X7,X6))|~m1_subset_1(X8,u1_struct_0(X7)))))&(((~r1_orders_2(X6,k2_waybel_0(X6,X7,X8),k2_waybel_0(X6,X7,esk8_4(X6,X7,X8,X11)))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,a_3_0_waybel_0(X6,X7,X8)))|((v2_struct_0(X6)|~l1_orders_2(X6))|((v2_struct_0(X7)|~l1_waybel_0(X7,X6))|~m1_subset_1(X8,u1_struct_0(X7))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s1_waybel_0__e1_34_1__waybel_0])])])])])])])])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v10_waybel_0(X2,X1)|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X2))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,plain,(v2_struct_0(X2)|v2_struct_0(X3)|r1_waybel_0(X3,X2,a_3_0_waybel_0(X3,X2,X1))|r1_orders_2(X2,X4,esk8_4(X3,X2,X1,X4))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,negated_conjecture,(v10_waybel_0(esk2_0,esk1_0)|m1_subset_1(esk5_1(X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,negated_conjecture,(m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk2_0))|v10_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])]), c_0_10]), c_0_11])).
cnf(c_0_15,plain,(v2_struct_0(X2)|v2_struct_0(X3)|r1_waybel_0(X3,X2,a_3_0_waybel_0(X3,X2,X1))|m1_subset_1(esk8_4(X3,X2,X1,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,X1))|r1_orders_2(esk2_0,X2,esk8_4(esk1_0,esk2_0,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_8]), c_0_9])]), c_0_11]), c_0_10])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk5_1(esk6_2(esk1_0,esk2_0)),u1_struct_0(esk2_0))|v10_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,X1))|m1_subset_1(esk8_4(esk1_0,esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_8]), c_0_9])]), c_0_11]), c_0_10])).
cnf(c_0_19,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,X1))|r1_orders_2(esk2_0,esk5_1(esk6_2(esk1_0,esk2_0)),esk8_4(esk1_0,esk2_0,X1,esk5_1(esk6_2(esk1_0,esk2_0))))|v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,X1))|m1_subset_1(esk8_4(esk1_0,esk2_0,X1,esk5_1(esk6_2(esk1_0,esk2_0))),u1_struct_0(esk2_0))|v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_21,negated_conjecture,(v10_waybel_0(esk2_0,esk1_0)|r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,X1),k2_waybel_0(esk1_0,esk2_0,X2))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r1_orders_2(esk2_0,esk5_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0)))|r1_orders_2(esk2_0,esk5_1(esk6_2(esk1_0,esk2_0)),esk8_4(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0),esk5_1(esk6_2(esk1_0,esk2_0))))|v10_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_19, c_0_14])).
cnf(c_0_23,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0)))|m1_subset_1(esk8_4(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0),esk5_1(esk6_2(esk1_0,esk2_0))),u1_struct_0(esk2_0))|v10_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_20, c_0_14])).
cnf(c_0_24,plain,(v2_struct_0(X2)|v2_struct_0(X3)|r1_waybel_0(X3,X2,a_3_0_waybel_0(X3,X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X3,k2_waybel_0(X3,X2,X1),k2_waybel_0(X3,X2,esk8_4(X3,X2,X1,X4)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0)))|r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0)),k2_waybel_0(esk1_0,esk2_0,esk8_4(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0),esk5_1(esk6_2(esk1_0,esk2_0)))))|v10_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14]), c_0_23])).
cnf(c_0_26,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v10_waybel_0(X2,X1)|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,a_3_0_waybel_0(X1,X2,esk6_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_27,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0)))|v10_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_8]), c_0_9])]), c_0_11]), c_0_10]), c_0_14]), c_0_17])).
cnf(c_0_28,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,a_3_0_waybel_0(X1,X2,X3))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~v10_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_29,negated_conjecture,(v10_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_8]), c_0_9])]), c_0_10]), c_0_11])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_31,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_8]), c_0_9])]), c_0_10]), c_0_11])).
cnf(c_0_32,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_29])])).
cnf(c_0_33,negated_conjecture,(r1_orders_2(esk2_0,X1,esk4_1(X1))|~v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_34,plain,(v2_struct_0(X2)|v2_struct_0(X3)|m1_subset_1(esk7_3(X3,X2,X1),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~r1_waybel_0(X3,X2,a_3_0_waybel_0(X3,X2,X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_35,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(esk4_1(X1),u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_37,plain,(v2_struct_0(X2)|v2_struct_0(X3)|r1_orders_2(X3,k2_waybel_0(X3,X2,X1),k2_waybel_0(X3,X2,X4))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~r1_waybel_0(X3,X2,a_3_0_waybel_0(X3,X2,X1))|~r1_orders_2(X2,esk7_3(X3,X2,X1),X4)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_38,negated_conjecture,(r1_orders_2(esk2_0,X1,esk4_1(X1))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_29])])).
cnf(c_0_39,negated_conjecture,(m1_subset_1(esk7_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_32]), c_0_8]), c_0_9])]), c_0_11]), c_0_10])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk4_1(X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_29])])).
cnf(c_0_41,negated_conjecture,(~v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk3_0),k2_waybel_0(esk1_0,esk2_0,esk4_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_42,negated_conjecture,(r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk3_0),k2_waybel_0(esk1_0,esk2_0,X1))|~r1_orders_2(esk2_0,esk7_3(esk1_0,esk2_0,esk3_0),X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_35]), c_0_32]), c_0_8]), c_0_9])]), c_0_11]), c_0_10])).
cnf(c_0_43,negated_conjecture,(r1_orders_2(esk2_0,esk7_3(esk1_0,esk2_0,esk3_0),esk4_1(esk7_3(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk4_1(esk7_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_40, c_0_39])).
cnf(c_0_45,negated_conjecture,(~r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk3_0),k2_waybel_0(esk1_0,esk2_0,esk4_1(X1)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_29])])).
cnf(c_0_46,negated_conjecture,(r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk3_0),k2_waybel_0(esk1_0,esk2_0,esk4_1(esk7_3(esk1_0,esk2_0,esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])])).
cnf(c_0_47,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_39])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 41
# Proof object formula steps           : 7
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 55
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 502
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 502
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 115
# Backward-rewritten                   : 263
# Generated clauses                    : 24345
# ...of the previous two non-trivial   : 24348
# Contextual simplify-reflections      : 4
# Paramodulations                      : 24345
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 106
#    Positive orientable unit clauses  : 89
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 15
# Current number of unprocessed clauses: 114
# ...number of literals in the above   : 136
# Current number of archived formulas  : 0
# Current number of archived clauses   : 396
# Clause-clause subsumption calls (NU) : 146483
# Rec. Clause-clause subsumption calls : 14605
# Non-unit clause-clause subsumptions  : 119
# Unit Clause-clause subsumption calls : 69
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 790
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1884572

# -------------------------------------------------
# User time                : 0.548 s
# System time              : 0.024 s
# Total time               : 0.572 s
# Maximum resident set size: 43152 pages
