# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t11_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>(v10_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r1_orders_2(X1,k2_waybel_0(X1,X2,X3),k2_waybel_0(X1,X2,X5)))))))))), file('out/t11_waybel_0__4_premises__rerun.E_input', t11_waybel_0)).
fof(fraenkel_a_3_0_waybel_0, axiom, (![X1]:![X2]:![X3]:![X4]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&((~(v2_struct_0(X3))&l1_waybel_0(X3,X2))&m1_subset_1(X4,u1_struct_0(X3))))=>(r2_hidden(X1,a_3_0_waybel_0(X2,X3,X4))<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&(X1=k2_waybel_0(X2,X3,X5)&r1_orders_2(X2,k2_waybel_0(X2,X3,X4),k2_waybel_0(X2,X3,X5))))))), file('out/t11_waybel_0__4_premises__rerun.E_input', fraenkel_a_3_0_waybel_0)).
fof(d13_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>(v10_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_waybel_0(X1,X2,a_3_0_waybel_0(X1,X2,X3))))))), file('out/t11_waybel_0__4_premises__rerun.E_input', d13_waybel_0)).
fof(d11_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))))), file('out/t11_waybel_0__4_premises__rerun.E_input', d11_waybel_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t11_waybel_0__4_premises__rerun.E_input', dt_l1_orders_2)).
fof(c_0_5, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>(v10_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r1_orders_2(X1,k2_waybel_0(X1,X2,X3),k2_waybel_0(X1,X2,X5))))))))))), inference(assume_negation,[status(cth)],[t11_waybel_0])).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:((((m1_subset_1(esk6_4(X6,X7,X8,X9),u1_struct_0(X8))|~r2_hidden(X6,a_3_0_waybel_0(X7,X8,X9)))|((v2_struct_0(X7)|~l1_orders_2(X7))|((v2_struct_0(X8)|~l1_waybel_0(X8,X7))|~m1_subset_1(X9,u1_struct_0(X8)))))&(((X6=k2_waybel_0(X7,X8,esk6_4(X6,X7,X8,X9))|~r2_hidden(X6,a_3_0_waybel_0(X7,X8,X9)))|((v2_struct_0(X7)|~l1_orders_2(X7))|((v2_struct_0(X8)|~l1_waybel_0(X8,X7))|~m1_subset_1(X9,u1_struct_0(X8)))))&((r1_orders_2(X7,k2_waybel_0(X7,X8,X9),k2_waybel_0(X7,X8,esk6_4(X6,X7,X8,X9)))|~r2_hidden(X6,a_3_0_waybel_0(X7,X8,X9)))|((v2_struct_0(X7)|~l1_orders_2(X7))|((v2_struct_0(X8)|~l1_waybel_0(X8,X7))|~m1_subset_1(X9,u1_struct_0(X8)))))))&(((~m1_subset_1(X11,u1_struct_0(X8))|(X6!=k2_waybel_0(X7,X8,X11)|~r1_orders_2(X7,k2_waybel_0(X7,X8,X9),k2_waybel_0(X7,X8,X11))))|r2_hidden(X6,a_3_0_waybel_0(X7,X8,X9)))|((v2_struct_0(X7)|~l1_orders_2(X7))|((v2_struct_0(X8)|~l1_waybel_0(X8,X7))|~m1_subset_1(X9,u1_struct_0(X8))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_3_0_waybel_0])])])])])])])])).
fof(c_0_7, negated_conjecture, (![X9]:![X11]:![X13]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~v2_struct_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))&(((m1_subset_1(esk3_0,u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0))&(((m1_subset_1(esk4_1(X9),u1_struct_0(esk2_0))|~m1_subset_1(X9,u1_struct_0(esk2_0)))|~v10_waybel_0(esk2_0,esk1_0))&(((r1_orders_2(esk2_0,X9,esk4_1(X9))|~m1_subset_1(X9,u1_struct_0(esk2_0)))|~v10_waybel_0(esk2_0,esk1_0))&((~r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk3_0),k2_waybel_0(esk1_0,esk2_0,esk4_1(X9)))|~m1_subset_1(X9,u1_struct_0(esk2_0)))|~v10_waybel_0(esk2_0,esk1_0)))))&(((m1_subset_1(esk5_1(X11),u1_struct_0(esk2_0))|~m1_subset_1(X11,u1_struct_0(esk2_0)))|v10_waybel_0(esk2_0,esk1_0))&(((~m1_subset_1(X13,u1_struct_0(esk2_0))|(~r1_orders_2(esk2_0,esk5_1(X11),X13)|r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,X11),k2_waybel_0(esk1_0,esk2_0,X13))))|~m1_subset_1(X11,u1_struct_0(esk2_0)))|v10_waybel_0(esk2_0,esk1_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])])])).
cnf(c_0_8,plain,(v2_struct_0(X2)|v2_struct_0(X3)|X4=k2_waybel_0(X3,X2,esk6_4(X4,X3,X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~r2_hidden(X4,a_3_0_waybel_0(X3,X2,X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:((((~v10_waybel_0(X5,X4)|(~m1_subset_1(X6,u1_struct_0(X5))|r1_waybel_0(X4,X5,a_3_0_waybel_0(X4,X5,X6))))|(v2_struct_0(X5)|~l1_waybel_0(X5,X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((m1_subset_1(esk7_2(X4,X5),u1_struct_0(X5))|v10_waybel_0(X5,X4))|(v2_struct_0(X5)|~l1_waybel_0(X5,X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&(((~r1_waybel_0(X4,X5,a_3_0_waybel_0(X4,X5,esk7_2(X4,X5)))|v10_waybel_0(X5,X4))|(v2_struct_0(X5)|~l1_waybel_0(X5,X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_waybel_0])])])])])])])])).
cnf(c_0_12,plain,(v2_struct_0(X2)|v2_struct_0(X3)|r1_orders_2(X3,k2_waybel_0(X3,X2,X1),k2_waybel_0(X3,X2,esk6_4(X4,X3,X2,X1)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~r2_hidden(X4,a_3_0_waybel_0(X3,X2,X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,negated_conjecture,(k2_waybel_0(X1,esk2_0,esk6_4(X2,X1,esk2_0,esk3_0))=X2|v2_struct_0(X1)|~r2_hidden(X2,a_3_0_waybel_0(X1,esk2_0,esk3_0))|~v10_waybel_0(esk2_0,esk1_0)|~l1_waybel_0(esk2_0,X1)|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])).
fof(c_0_14, plain, (![X6]:![X7]:![X8]:![X10]:![X8]:![X11]:(((((m1_subset_1(esk8_3(X6,X7,X8),u1_struct_0(X7))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~m1_subset_1(X10,u1_struct_0(X7))|(~r1_orders_2(X7,esk8_3(X6,X7,X8),X10)|r2_hidden(k2_waybel_0(X6,X7,X10),X8)))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))&(((((m1_subset_1(esk9_4(X6,X7,X8,X11),u1_struct_0(X7))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X11,esk9_4(X6,X7,X8,X11))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~r2_hidden(k2_waybel_0(X6,X7,esk9_4(X6,X7,X8,X11)),X8)|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_waybel_0])])])])])])])])).
fof(c_0_15, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_16,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,a_3_0_waybel_0(X1,X2,X3))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~v10_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(~v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,esk3_0),k2_waybel_0(esk1_0,esk2_0,esk4_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_18,negated_conjecture,(r1_orders_2(X1,k2_waybel_0(X1,esk2_0,esk3_0),X2)|v2_struct_0(X1)|~r2_hidden(X2,a_3_0_waybel_0(X1,esk2_0,esk3_0))|~v10_waybel_0(esk2_0,esk1_0)|~l1_waybel_0(esk2_0,X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_10]), c_0_9])).
cnf(c_0_19,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(k2_waybel_0(X1,X2,X4),X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)|~r1_orders_2(X2,esk8_3(X1,X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,negated_conjecture,(r1_orders_2(esk2_0,X1,esk4_1(X1))|~v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk4_1(X1),u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26,negated_conjecture,(r1_waybel_0(X1,esk2_0,a_3_0_waybel_0(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~v10_waybel_0(esk2_0,esk1_0)|~v10_waybel_0(esk2_0,X1)|~l1_waybel_0(esk2_0,X1)|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_9]), c_0_10])).
cnf(c_0_27,negated_conjecture,(~r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk4_1(X1)),a_3_0_waybel_0(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_28,negated_conjecture,(r2_hidden(k2_waybel_0(X1,esk2_0,esk4_1(esk8_3(X1,esk2_0,X2))),X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~r1_waybel_0(X1,esk2_0,X2)|~m1_subset_1(esk8_3(X1,esk2_0,X2),u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0)|~l1_waybel_0(esk2_0,X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_10]), c_0_24])).
cnf(c_0_29,negated_conjecture,(l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_25, c_0_20])).
cnf(c_0_30,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk3_0))|~v10_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_20]), c_0_19])]), c_0_21])).
cnf(c_0_31,plain,(v2_struct_0(X2)|v2_struct_0(X3)|r2_hidden(X4,a_3_0_waybel_0(X3,X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_orders_2(X3)|~r1_orders_2(X3,k2_waybel_0(X3,X2,X1),k2_waybel_0(X3,X2,X5))|X4!=k2_waybel_0(X3,X2,X5)|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_32,negated_conjecture,(v10_waybel_0(esk2_0,esk1_0)|r1_orders_2(esk1_0,k2_waybel_0(esk1_0,esk2_0,X1),k2_waybel_0(esk1_0,esk2_0,X2))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r1_orders_2(esk2_0,esk5_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_33,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|r1_orders_2(X2,X4,esk9_4(X1,X2,X3,X4))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_34,negated_conjecture,(v10_waybel_0(esk2_0,esk1_0)|m1_subset_1(esk5_1(X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_35,negated_conjecture,(~m1_subset_1(esk8_3(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))|~v10_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_19])]), c_0_21]), c_0_30])).
cnf(c_0_36,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37,negated_conjecture,(r2_hidden(X1,a_3_0_waybel_0(esk1_0,esk2_0,X2))|v10_waybel_0(esk2_0,esk1_0)|X1!=k2_waybel_0(esk1_0,esk2_0,X3)|~r1_orders_2(esk2_0,esk5_1(X2),X3)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_19]), c_0_20])]), c_0_21]), c_0_10])).
cnf(c_0_38,negated_conjecture,(r1_waybel_0(X1,esk2_0,X2)|r1_orders_2(esk2_0,esk5_1(X3),esk9_4(X1,esk2_0,X2,esk5_1(X3)))|v10_waybel_0(esk2_0,esk1_0)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~l1_waybel_0(esk2_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_10])).
cnf(c_0_39,negated_conjecture,(~v10_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_29]), c_0_19])]), c_0_10]), c_0_21]), c_0_30])).
cnf(c_0_40,negated_conjecture,(r1_waybel_0(X1,esk2_0,X2)|r2_hidden(X3,a_3_0_waybel_0(esk1_0,esk2_0,X4))|v2_struct_0(X1)|X3!=k2_waybel_0(esk1_0,esk2_0,esk9_4(X1,esk2_0,X2,esk5_1(X4)))|~l1_struct_0(X1)|~m1_subset_1(esk9_4(X1,esk2_0,X2,esk5_1(X4)),u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~l1_waybel_0(esk2_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_41,negated_conjecture,(r1_waybel_0(X1,esk2_0,X2)|r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(X1,esk2_0,X2,esk5_1(X3))),a_3_0_waybel_0(esk1_0,esk2_0,X3))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(esk9_4(X1,esk2_0,X2,esk5_1(X3)),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~l1_waybel_0(esk2_0,X1)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_42,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|m1_subset_1(esk9_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_43,negated_conjecture,(r1_waybel_0(X1,esk2_0,X2)|r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(X1,esk2_0,X2,esk5_1(X3))),a_3_0_waybel_0(esk1_0,esk2_0,X3))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(esk5_1(X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~l1_waybel_0(esk2_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_10])).
cnf(c_0_44,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(k2_waybel_0(X1,X2,esk9_4(X1,X2,X3,X4)),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_45,negated_conjecture,(r1_waybel_0(X1,esk2_0,X2)|r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(X1,esk2_0,X2,esk5_1(X3))),a_3_0_waybel_0(esk1_0,esk2_0,X3))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~l1_waybel_0(esk2_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_34]), c_0_39])).
cnf(c_0_46,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,X1))|~m1_subset_1(esk5_1(X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_29]), c_0_19])]), c_0_10]), c_0_21])).
cnf(c_0_47,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_34]), c_0_39])).
cnf(c_0_48,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v10_waybel_0(X2,X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(X2))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_49,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk7_2(X1,esk2_0)))|v10_waybel_0(esk2_0,X1)|v2_struct_0(X1)|~l1_waybel_0(esk2_0,X1)|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_10])).
cnf(c_0_50,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v10_waybel_0(X2,X1)|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,a_3_0_waybel_0(X1,X2,esk7_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_51,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,a_3_0_waybel_0(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_20]), c_0_19])]), c_0_39]), c_0_21])).
cnf(c_0_52,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_19]), c_0_20])]), c_0_39]), c_0_10]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 42
# Proof object formula steps           : 11
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 123
# ...of these trivial                  : 2
# ...subsumed                          : 9
# ...remaining for further processing  : 112
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 193
# ...of the previous two non-trivial   : 164
# Contextual simplify-reflections      : 7
# Paramodulations                      : 187
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 85
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 78
# Current number of unprocessed clauses: 87
# ...number of literals in the above   : 1061
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 2417
# Rec. Clause-clause subsumption calls : 233
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 41
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11477

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.004 s
# Total time               : 0.028 s
# Maximum resident set size: 43152 pages
