# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d11_binop_1, axiom, (![X1]:![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>(r6_binop_1(X1,X2,X3)<=>(r4_binop_1(X1,X2,X3)&r5_binop_1(X1,X2,X3)))))), file('out/t11_filter_1__7_premises__rerun.E_input', d11_binop_1)).
fof(dt_k3_filter_1, axiom, (![X1]:![X2]:![X3]:((~(v1_xboole_0(X1))&((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))&(v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))))=>(v1_funct_1(k3_filter_1(X1,X2,X3))&(v1_funct_2(k3_filter_1(X1,X2,X3),k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))&m1_subset_1(k3_filter_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2)))))))), file('out/t11_filter_1__7_premises__rerun.E_input', dt_k3_filter_1)).
fof(t11_filter_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m2_filter_1(X3,X1,X2)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r6_binop_1(X1,X3,X4)=>r6_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4)))))))), file('out/t11_filter_1__7_premises__rerun.E_input', t11_filter_1)).
fof(t10_filter_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m2_filter_1(X3,X1,X2)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r5_binop_1(X1,X3,X4)=>r5_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4)))))))), file('out/t11_filter_1__7_premises__rerun.E_input', t10_filter_1)).
fof(cc2_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2)))), file('out/t11_filter_1__7_premises__rerun.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, (![X1]:![X2]:v1_relat_1(k2_zfmisc_1(X1,X2))), file('out/t11_filter_1__7_premises__rerun.E_input', fc6_relat_1)).
fof(dt_m2_filter_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&v1_relat_1(X2))=>![X3]:(m2_filter_1(X3,X1,X2)=>(v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))))), file('out/t11_filter_1__7_premises__rerun.E_input', dt_m2_filter_1)).
fof(t9_filter_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m2_filter_1(X3,X1,X2)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r4_binop_1(X1,X3,X4)=>r4_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4)))))))), file('out/t11_filter_1__7_premises__rerun.E_input', t9_filter_1)).
fof(c_0_8, plain, (![X4]:![X5]:![X6]:(((((r4_binop_1(X4,X5,X6)|~r6_binop_1(X4,X5,X6))|(~v1_funct_1(X6)|(~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4))))))|(~v1_funct_1(X5)|(~v1_funct_2(X5,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4))))))&(((r5_binop_1(X4,X5,X6)|~r6_binop_1(X4,X5,X6))|(~v1_funct_1(X6)|(~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4))))))|(~v1_funct_1(X5)|(~v1_funct_2(X5,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))))))&((((~r4_binop_1(X4,X5,X6)|~r5_binop_1(X4,X5,X6))|r6_binop_1(X4,X5,X6))|(~v1_funct_1(X6)|(~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4))))))|(~v1_funct_1(X5)|(~v1_funct_2(X5,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_binop_1])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:((v1_funct_1(k3_filter_1(X4,X5,X6))|(v1_xboole_0(X4)|((~v1_partfun1(X5,X4)|(~v3_relat_2(X5)|(~v8_relat_2(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4))))))|(~v1_funct_1(X6)|(~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4))))))))&((v1_funct_2(k3_filter_1(X4,X5,X6),k2_zfmisc_1(k8_eqrel_1(X4,X5),k8_eqrel_1(X4,X5)),k8_eqrel_1(X4,X5))|(v1_xboole_0(X4)|((~v1_partfun1(X5,X4)|(~v3_relat_2(X5)|(~v8_relat_2(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4))))))|(~v1_funct_1(X6)|(~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4))))))))&(m1_subset_1(k3_filter_1(X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X4,X5),k8_eqrel_1(X4,X5)),k8_eqrel_1(X4,X5))))|(v1_xboole_0(X4)|((~v1_partfun1(X5,X4)|(~v3_relat_2(X5)|(~v8_relat_2(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4))))))|(~v1_funct_1(X6)|(~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_filter_1])])])])).
cnf(c_0_10,plain,(r6_binop_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X3,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X3)|~r5_binop_1(X2,X1,X3)|~r4_binop_1(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11,plain,(v1_xboole_0(X2)|m1_subset_1(k3_filter_1(X2,X3,X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X2,X3),k8_eqrel_1(X2,X3)),k8_eqrel_1(X2,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v8_relat_2(X3)|~v3_relat_2(X3)|~v1_partfun1(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(v1_xboole_0(X2)|v1_funct_1(k3_filter_1(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v8_relat_2(X3)|~v3_relat_2(X3)|~v1_partfun1(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(v1_xboole_0(X2)|v1_funct_2(k3_filter_1(X2,X3,X1),k2_zfmisc_1(k8_eqrel_1(X2,X3),k8_eqrel_1(X2,X3)),k8_eqrel_1(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v8_relat_2(X3)|~v3_relat_2(X3)|~v1_partfun1(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m2_filter_1(X3,X1,X2)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r6_binop_1(X1,X3,X4)=>r6_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))))))))), inference(assume_negation,[status(cth)],[t11_filter_1])).
cnf(c_0_15,plain,(v1_xboole_0(X1)|r6_binop_1(k8_eqrel_1(X1,X2),X3,k3_filter_1(X1,X2,X4))|~v8_relat_2(X2)|~v3_relat_2(X2)|~v1_partfun1(X2,X1)|~r5_binop_1(k8_eqrel_1(X1,X2),X3,k3_filter_1(X1,X2,X4))|~r4_binop_1(k8_eqrel_1(X1,X2),X3,k3_filter_1(X1,X2,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_2(X3,k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X3)|~v1_funct_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])).
fof(c_0_16, plain, (![X5]:![X6]:![X7]:![X8]:(v1_xboole_0(X5)|((~v1_partfun1(X6,X5)|(~v3_relat_2(X6)|(~v8_relat_2(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5))))))|(~m2_filter_1(X7,X5,X6)|(~m2_filter_1(X8,X5,X6)|(~r5_binop_1(X5,X7,X8)|r5_binop_1(k8_eqrel_1(X5,X6),k3_filter_1(X5,X6,X7),k3_filter_1(X5,X6,X8)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_filter_1])])])])])])).
fof(c_0_17, plain, (![X3]:![X4]:(~v1_relat_1(X3)|(~m1_subset_1(X4,k1_zfmisc_1(X3))|v1_relat_1(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])])])).
fof(c_0_18, negated_conjecture, ((~v1_xboole_0(esk1_0)&((v1_partfun1(esk2_0,esk1_0)&(v3_relat_2(esk2_0)&(v8_relat_2(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0))))))&(m2_filter_1(esk3_0,esk1_0,esk2_0)&(m2_filter_1(esk4_0,esk1_0,esk2_0)&(r6_binop_1(esk1_0,esk3_0,esk4_0)&~r6_binop_1(k8_eqrel_1(esk1_0,esk2_0),k3_filter_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])])])).
fof(c_0_19, plain, (![X3]:![X4]:v1_relat_1(k2_zfmisc_1(X3,X4))), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(((v1_funct_1(X6)|~m2_filter_1(X6,X4,X5))|(v1_xboole_0(X4)|~v1_relat_1(X5)))&(((v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m2_filter_1(X6,X4,X5))|(v1_xboole_0(X4)|~v1_relat_1(X5)))&((m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))|~m2_filter_1(X6,X4,X5))|(v1_xboole_0(X4)|~v1_relat_1(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_filter_1])])])])])])])).
cnf(c_0_21,plain,(v1_xboole_0(X1)|r6_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~v8_relat_2(X2)|~v3_relat_2(X2)|~v1_partfun1(X2,X1)|~r5_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~r4_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X4)|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_11]), c_0_12]), c_0_13])).
cnf(c_0_22,plain,(r5_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|v1_xboole_0(X1)|~r5_binop_1(X1,X3,X4)|~m2_filter_1(X4,X1,X2)|~m2_filter_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v8_relat_2(X2)|~v3_relat_2(X2)|~v1_partfun1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, (![X5]:![X6]:![X7]:![X8]:(v1_xboole_0(X5)|((~v1_partfun1(X6,X5)|(~v3_relat_2(X6)|(~v8_relat_2(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5))))))|(~m2_filter_1(X7,X5,X6)|(~m2_filter_1(X8,X5,X6)|(~r4_binop_1(X5,X7,X8)|r4_binop_1(k8_eqrel_1(X5,X6),k3_filter_1(X5,X6,X7),k3_filter_1(X5,X6,X8)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_filter_1])])])])])])).
cnf(c_0_24,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(v1_xboole_0(X2)|v1_funct_1(X3)|~v1_relat_1(X1)|~m2_filter_1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,negated_conjecture,(m2_filter_1(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30,negated_conjecture,(m2_filter_1(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31,plain,(v1_xboole_0(X1)|r6_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~m2_filter_1(X4,X1,X2)|~m2_filter_1(X3,X1,X2)|~v8_relat_2(X2)|~v3_relat_2(X2)|~v1_partfun1(X2,X1)|~r5_binop_1(X1,X3,X4)|~r4_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X4)|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_32,plain,(r4_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|v1_xboole_0(X1)|~r4_binop_1(X1,X3,X4)|~m2_filter_1(X4,X1,X2)|~m2_filter_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v8_relat_2(X2)|~v3_relat_2(X2)|~v1_partfun1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33,plain,(v1_xboole_0(X2)|m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_relat_1(X1)|~m2_filter_1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_34,negated_conjecture,(v1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_35,plain,(v1_xboole_0(X2)|v1_funct_2(X3,k2_zfmisc_1(X2,X2),X2)|~v1_relat_1(X1)|~m2_filter_1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36,negated_conjecture,(v1_funct_1(esk4_0)|~v1_relat_1(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_37,negated_conjecture,(v1_funct_1(esk3_0)|~v1_relat_1(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_30]), c_0_29])).
cnf(c_0_38,negated_conjecture,(~r6_binop_1(k8_eqrel_1(esk1_0,esk2_0),k3_filter_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_39,plain,(v1_xboole_0(X1)|r6_binop_1(k8_eqrel_1(X1,X2),k3_filter_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~m2_filter_1(X4,X1,X2)|~m2_filter_1(X3,X1,X2)|~v8_relat_2(X2)|~v3_relat_2(X2)|~v1_partfun1(X2,X1)|~r5_binop_1(X1,X3,X4)|~r4_binop_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X4)|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40,negated_conjecture,(v8_relat_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_41,negated_conjecture,(v3_relat_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_42,negated_conjecture,(v1_partfun1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_43,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_28]), c_0_34])]), c_0_29])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_30]), c_0_34])]), c_0_29])).
cnf(c_0_45,negated_conjecture,(v1_funct_2(esk4_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_28]), c_0_34])]), c_0_29])).
cnf(c_0_46,negated_conjecture,(v1_funct_2(esk3_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_30]), c_0_34])]), c_0_29])).
cnf(c_0_47,negated_conjecture,(v1_funct_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_34])])).
cnf(c_0_48,negated_conjecture,(v1_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_34])])).
cnf(c_0_49,plain,(r5_binop_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X3,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X3)|~r6_binop_1(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_50,negated_conjecture,(~r5_binop_1(esk1_0,esk3_0,esk4_0)|~r4_binop_1(esk1_0,esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_28]), c_0_30]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_25]), c_0_45]), c_0_46]), c_0_47]), c_0_48])]), c_0_29])).
cnf(c_0_51,negated_conjecture,(r5_binop_1(esk1_0,X1,esk4_0)|~r6_binop_1(esk1_0,X1,esk4_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))|~v1_funct_2(X1,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_43]), c_0_45]), c_0_47])])).
cnf(c_0_52,negated_conjecture,(r6_binop_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_53,plain,(r4_binop_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X3,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X3)|~r6_binop_1(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_54,negated_conjecture,(~r4_binop_1(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_44]), c_0_46]), c_0_48])])).
cnf(c_0_55,negated_conjecture,(r4_binop_1(esk1_0,X1,esk4_0)|~r6_binop_1(esk1_0,X1,esk4_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))|~v1_funct_2(X1,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_43]), c_0_45]), c_0_47])])).
cnf(c_0_56,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_52]), c_0_44]), c_0_46]), c_0_48])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 40
# Proof object formula steps           : 17
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 54
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 82
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 43
# ...of the previous two non-trivial   : 40
# Contextual simplify-reflections      : 16
# Paramodulations                      : 43
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 38
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 42
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 1452
# Rec. Clause-clause subsumption calls : 152
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5738

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45332 pages
