# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t11_waybel_9, conjecture, (![X1]:(l1_orders_2(X1)=>g1_orders_2(u1_struct_0(k7_lattice3(X1)),u1_orders_2(k7_lattice3(X1)))=g1_orders_2(u1_struct_0(k3_waybel_9(X1)),u1_orders_2(k3_waybel_9(X1))))), file('out/t11_waybel_9__5_premises__rerun.E_input', t11_waybel_9)).
fof(abstractness_v1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))), file('out/t11_waybel_9__5_premises__rerun.E_input', abstractness_v1_orders_2)).
fof(d6_waybel_9, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:((v6_waybel_0(X2,X1)&l1_waybel_0(X2,X1))=>(X2=k3_waybel_9(X1)<=>(u1_struct_0(X2)=u1_struct_0(X1)&(u1_orders_2(X2)=k3_relset_1(u1_struct_0(X1),u1_struct_0(X1),u1_orders_2(X1))&u1_waybel_0(X1,X2)=k3_struct_0(X1))))))), file('out/t11_waybel_9__5_premises__rerun.E_input', d6_waybel_9)).
fof(dt_k3_waybel_9, axiom, (![X1]:(l1_orders_2(X1)=>(v6_waybel_0(k3_waybel_9(X1),X1)&l1_waybel_0(k3_waybel_9(X1),X1)))), file('out/t11_waybel_9__5_premises__rerun.E_input', dt_k3_waybel_9)).
fof(d5_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>k7_lattice3(X1)=g1_orders_2(u1_struct_0(X1),k3_relset_1(u1_struct_0(X1),u1_struct_0(X1),u1_orders_2(X1))))), file('out/t11_waybel_9__5_premises__rerun.E_input', d5_lattice3)).
fof(dt_k7_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(k7_lattice3(X1))&l1_orders_2(k7_lattice3(X1))))), file('out/t11_waybel_9__5_premises__rerun.E_input', dt_k7_lattice3)).
fof(c_0_6, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>g1_orders_2(u1_struct_0(k7_lattice3(X1)),u1_orders_2(k7_lattice3(X1)))=g1_orders_2(u1_struct_0(k3_waybel_9(X1)),u1_orders_2(k3_waybel_9(X1)))))), inference(assume_negation,[status(cth)],[t11_waybel_9])).
fof(c_0_7, negated_conjecture, ((l1_orders_2(esk1_0)&g1_orders_2(u1_struct_0(k7_lattice3(esk1_0)),u1_orders_2(k7_lattice3(esk1_0)))!=g1_orders_2(u1_struct_0(k3_waybel_9(esk1_0)),u1_orders_2(k3_waybel_9(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
fof(c_0_8, plain, (![X2]:(~l1_orders_2(X2)|(~v1_orders_2(X2)|X2=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
fof(c_0_9, plain, (![X3]:![X4]:(((((u1_struct_0(X4)=u1_struct_0(X3)|X4!=k3_waybel_9(X3))|(~v6_waybel_0(X4,X3)|~l1_waybel_0(X4,X3)))|~l1_orders_2(X3))&((((u1_orders_2(X4)=k3_relset_1(u1_struct_0(X3),u1_struct_0(X3),u1_orders_2(X3))|X4!=k3_waybel_9(X3))|(~v6_waybel_0(X4,X3)|~l1_waybel_0(X4,X3)))|~l1_orders_2(X3))&(((u1_waybel_0(X3,X4)=k3_struct_0(X3)|X4!=k3_waybel_9(X3))|(~v6_waybel_0(X4,X3)|~l1_waybel_0(X4,X3)))|~l1_orders_2(X3))))&((((u1_struct_0(X4)!=u1_struct_0(X3)|(u1_orders_2(X4)!=k3_relset_1(u1_struct_0(X3),u1_struct_0(X3),u1_orders_2(X3))|u1_waybel_0(X3,X4)!=k3_struct_0(X3)))|X4=k3_waybel_9(X3))|(~v6_waybel_0(X4,X3)|~l1_waybel_0(X4,X3)))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_waybel_9])])])])])])).
fof(c_0_10, plain, (![X2]:((v6_waybel_0(k3_waybel_9(X2),X2)|~l1_orders_2(X2))&(l1_waybel_0(k3_waybel_9(X2),X2)|~l1_orders_2(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_waybel_9])])])).
cnf(c_0_11,negated_conjecture,(g1_orders_2(u1_struct_0(k7_lattice3(esk1_0)),u1_orders_2(k7_lattice3(esk1_0)))!=g1_orders_2(u1_struct_0(k3_waybel_9(esk1_0)),u1_orders_2(k3_waybel_9(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(u1_struct_0(X2)=u1_struct_0(X1)|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~v6_waybel_0(X2,X1)|X2!=k3_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(v6_waybel_0(k3_waybel_9(X1),X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(l1_waybel_0(k3_waybel_9(X1),X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, plain, (![X2]:(~l1_orders_2(X2)|k7_lattice3(X2)=g1_orders_2(u1_struct_0(X2),k3_relset_1(u1_struct_0(X2),u1_struct_0(X2),u1_orders_2(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_lattice3])])).
cnf(c_0_17,plain,(u1_orders_2(X2)=k3_relset_1(u1_struct_0(X1),u1_struct_0(X1),u1_orders_2(X1))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~v6_waybel_0(X2,X1)|X2!=k3_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,negated_conjecture,(g1_orders_2(u1_struct_0(k3_waybel_9(esk1_0)),u1_orders_2(k3_waybel_9(esk1_0)))!=k7_lattice3(esk1_0)|~v1_orders_2(k7_lattice3(esk1_0))|~l1_orders_2(k7_lattice3(esk1_0))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_19,plain,(u1_struct_0(k3_waybel_9(X1))=u1_struct_0(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_14]), c_0_15])).
cnf(c_0_20,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21,plain,(k7_lattice3(X1)=g1_orders_2(u1_struct_0(X1),k3_relset_1(u1_struct_0(X1),u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(k3_relset_1(u1_struct_0(X1),u1_struct_0(X1),u1_orders_2(X1))=u1_orders_2(k3_waybel_9(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_17]), c_0_14]), c_0_15])).
cnf(c_0_23,negated_conjecture,(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(k3_waybel_9(esk1_0)))!=k7_lattice3(esk1_0)|~v1_orders_2(k7_lattice3(esk1_0))|~l1_orders_2(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
cnf(c_0_24,plain,(g1_orders_2(u1_struct_0(X1),u1_orders_2(k3_waybel_9(X1)))=k7_lattice3(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_25, plain, (![X2]:((v1_orders_2(k7_lattice3(X2))|~l1_orders_2(X2))&(l1_orders_2(k7_lattice3(X2))|~l1_orders_2(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_lattice3])])])).
cnf(c_0_26,negated_conjecture,(~v1_orders_2(k7_lattice3(esk1_0))|~l1_orders_2(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_20])])).
cnf(c_0_27,plain,(v1_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_28,negated_conjecture,(~l1_orders_2(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_20])])).
cnf(c_0_29,plain,(l1_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 18
# Proof object formula steps           : 13
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 34
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 34
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 21
# ...of the previous two non-trivial   : 20
# Contextual simplify-reflections      : 6
# Paramodulations                      : 18
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 16
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 45
# Current number of archived formulas  : 0
# Current number of archived clauses   : 15
# Clause-clause subsumption calls (NU) : 91
# Rec. Clause-clause subsumption calls : 41
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1218

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 45280 pages
