# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t11_funct_2, conjecture, (![X1]:![X2]:![X3]:![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>~((r2_hidden(X3,k2_relset_1(X2,X4))&![X5]:~((r2_hidden(X5,X1)&k1_funct_1(X4,X5)=X3)))))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', t11_funct_2)).
fof(cc2_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2)))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', cc1_relset_1)).
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', t7_boole)).
fof(redefinition_k2_relset_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', redefinition_k2_relset_1)).
fof(d3_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X1))&X3=k1_funct_1(X1,X4)))))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', d3_funct_1)).
fof(d1_funct_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((~(X2=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X3)))&(X2=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', d1_funct_2)).
fof(redefinition_k1_relset_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', redefinition_k1_relset_1)).
fof(fc11_relat_1, axiom, (![X1]:(v1_xboole_0(X1)=>v1_xboole_0(k10_xtuple_0(X1)))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', fc11_relat_1)).
fof(cc8_relat_1, axiom, (![X1]:(v1_xboole_0(X1)=>![X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>(v1_xboole_0(X2)&(v1_relat_1(X2)&v5_relat_1(X2,X1)))))), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', cc8_relat_1)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t11_funct_2__102_premises__rerun__rerun.E_input', fc1_xboole_0)).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>~((r2_hidden(X3,k2_relset_1(X2,X4))&![X5]:~((r2_hidden(X5,X1)&k1_funct_1(X4,X5)=X3))))))), inference(assume_negation,[status(cth)],[t11_funct_2])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:((v4_relat_1(X6,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(v5_relat_1(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_13, negated_conjecture, (![X10]:((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,esk2_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(r2_hidden(esk3_0,k2_relset_1(esk2_0,esk4_0))&(~r2_hidden(X10,esk1_0)|k1_funct_1(esk4_0,X10)!=esk3_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_15, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_16, plain, (![X3]:![X4]:((~v1_relat_1(X4)|~v5_relat_1(X4,X3))|k2_relset_1(X3,X4)=k10_xtuple_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_17,plain,(v5_relat_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X6]:![X11]:((((((r2_hidden(esk5_3(X5,X6,X7),k9_xtuple_0(X5))|~r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((X7=k1_funct_1(X5,esk5_3(X5,X6,X7))|~r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(X9,k9_xtuple_0(X5))|X7!=k1_funct_1(X5,X9))|r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(esk6_2(X5,X6),X6)|(~r2_hidden(X11,k9_xtuple_0(X5))|esk6_2(X5,X6)!=k1_funct_1(X5,X11)))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&((((r2_hidden(esk7_2(X5,X6),k9_xtuple_0(X5))|r2_hidden(esk6_2(X5,X6),X6))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((esk6_2(X5,X6)=k1_funct_1(X5,esk7_2(X5,X6))|r2_hidden(esk6_2(X5,X6),X6))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_1])])])])])])])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:(((((~v1_funct_2(X6,X4,X5)|X4=k1_relset_1(X4,X6))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X4!=k1_relset_1(X4,X6)|v1_funct_2(X6,X4,X5))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))&((((~v1_funct_2(X6,X4,X5)|X6=k1_xboole_0)|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X6!=k1_xboole_0|v1_funct_2(X6,X4,X5))|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_funct_2])])])])).
cnf(c_0_22,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,negated_conjecture,(r2_hidden(esk3_0,k2_relset_1(esk2_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,plain,(k2_relset_1(X1,X2)=k10_xtuple_0(X2)|~v5_relat_1(X2,X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,negated_conjecture,(v5_relat_1(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26,negated_conjecture,(v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_19, c_0_18])).
cnf(c_0_27,plain,(X3=k1_funct_1(X1,esk5_3(X1,X2,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, (![X3]:![X4]:((~v1_relat_1(X4)|~v4_relat_1(X4,X3))|k1_relset_1(X3,X4)=k9_xtuple_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_29,plain,(X3=k1_xboole_0|X2=k1_relset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,negated_conjecture,(v1_funct_2(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31,plain,(v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32,negated_conjecture,(~v1_xboole_0(k2_relset_1(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_33,negated_conjecture,(k2_relset_1(esk2_0,esk4_0)=k10_xtuple_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
fof(c_0_34, plain, (![X2]:(~v1_xboole_0(X2)|v1_xboole_0(k10_xtuple_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc11_relat_1])])).
cnf(c_0_35,negated_conjecture,(k1_funct_1(esk4_0,X1)!=esk3_0|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36,plain,(k1_funct_1(X1,esk5_3(X1,k10_xtuple_0(X1),X2))=X2|~v1_relat_1(X1)|~r2_hidden(X2,k10_xtuple_0(X1))|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_37,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38,plain,(r2_hidden(esk5_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_39,plain,(k1_relset_1(X1,X2)=k9_xtuple_0(X2)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40,negated_conjecture,(k1_relset_1(esk1_0,esk4_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_18]), c_0_30])])).
cnf(c_0_41,negated_conjecture,(v4_relat_1(esk4_0,esk1_0)), inference(spm,[status(thm)],[c_0_31, c_0_18])).
fof(c_0_42, plain, (![X3]:![X4]:(((v1_xboole_0(X4)|(~v1_relat_1(X4)|~v5_relat_1(X4,X3)))|~v1_xboole_0(X3))&(((v1_relat_1(X4)|(~v1_relat_1(X4)|~v5_relat_1(X4,X3)))|~v1_xboole_0(X3))&((v5_relat_1(X4,X3)|(~v1_relat_1(X4)|~v5_relat_1(X4,X3)))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_relat_1])])])])])])).
cnf(c_0_43,negated_conjecture,(~v1_xboole_0(k10_xtuple_0(esk4_0))), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_44,plain,(v1_xboole_0(k10_xtuple_0(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45,negated_conjecture,(~r2_hidden(esk5_3(esk4_0,k10_xtuple_0(esk4_0),esk3_0),esk1_0)|~r2_hidden(esk3_0,k10_xtuple_0(esk4_0))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_26]), c_0_37])])])).
cnf(c_0_46,negated_conjecture,(r2_hidden(esk3_0,k10_xtuple_0(esk4_0))), inference(rw,[status(thm)],[c_0_23, c_0_33])).
cnf(c_0_47,plain,(r2_hidden(esk5_3(X1,k10_xtuple_0(X1),X2),k9_xtuple_0(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k10_xtuple_0(X1))|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_48,negated_conjecture,(k9_xtuple_0(esk4_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_26]), c_0_41])])).
cnf(c_0_49,plain,(v1_xboole_0(X2)|~v1_xboole_0(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50,negated_conjecture,(~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51,negated_conjecture,(~r2_hidden(esk5_3(esk4_0,k10_xtuple_0(esk4_0),esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_52,negated_conjecture,(esk2_0=k1_xboole_0|r2_hidden(esk5_3(esk4_0,k10_xtuple_0(esk4_0),X1),esk1_0)|~r2_hidden(X1,k10_xtuple_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_26]), c_0_37])])).
cnf(c_0_53,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_25]), c_0_26])]), c_0_50])).
cnf(c_0_54,negated_conjecture,(esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_46])])).
cnf(c_0_55,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_56,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54]), c_0_55])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 35
# Proof object formula steps           : 22
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 24
# Processed clauses                    : 58
# ...of these trivial                  : 0
# ...subsumed                          : 8
# ...remaining for further processing  : 50
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 12
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 41
# Contextual simplify-reflections      : 4
# Paramodulations                      : 35
# Factorizations                       : 0
# Equation resolutions                 : 8
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 17
# Current number of archived formulas  : 0
# Current number of archived clauses   : 12
# Clause-clause subsumption calls (NU) : 236
# Rec. Clause-clause subsumption calls : 92
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 35
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2346

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 47340 pages
