# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t11_setfam_1__11_premises__rerun__rerun.E_input', d1_tarski)).
fof(d1_setfam_1, axiom, (![X1]:![X2]:((~(X1=k1_xboole_0)=>(X2=k1_setfam_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>![X4]:(r2_hidden(X4,X1)=>r2_hidden(X3,X4)))))&(X1=k1_xboole_0=>(X2=k1_setfam_1(X1)<=>X2=k1_xboole_0)))), file('out/t11_setfam_1__11_premises__rerun__rerun.E_input', d1_setfam_1)).
fof(t7_tarski, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3)))))))), file('out/t11_setfam_1__11_premises__rerun__rerun.E_input', t7_tarski)).
fof(t4_setfam_1, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1))), file('out/t11_setfam_1__11_premises__rerun__rerun.E_input', t4_setfam_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t11_setfam_1__11_premises__rerun__rerun.E_input', d3_tarski)).
fof(t11_setfam_1, conjecture, (![X1]:k1_setfam_1(k1_tarski(X1))=X1), file('out/t11_setfam_1__11_premises__rerun__rerun.E_input', t11_setfam_1)).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_7,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_8, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:![X6]:![X12]:![X5]:![X6]:![X6]:((((((~r2_hidden(X7,X6)|(~r2_hidden(X8,X5)|r2_hidden(X7,X8)))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&((((r2_hidden(esk2_3(X5,X6,X7),X5)|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(X7,esk2_3(X5,X6,X7))|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((((r2_hidden(esk4_2(X5,X6),X5)|~r2_hidden(esk3_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(esk3_2(X5,X6),esk4_2(X5,X6))|~r2_hidden(esk3_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0))&(((r2_hidden(esk3_2(X5,X6),X6)|(~r2_hidden(X12,X5)|r2_hidden(esk3_2(X5,X6),X12)))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((X6!=k1_setfam_1(X5)|X6=k1_xboole_0)|X5!=k1_xboole_0)&((X6!=k1_xboole_0|X6=k1_setfam_1(X5))|X5!=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_setfam_1])])])])])])])])).
cnf(c_0_9,plain,(r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(er,[status(thm)],[c_0_7])).
fof(c_0_10, plain, (![X5]:![X6]:![X8]:((r2_hidden(esk5_2(X5,X6),X6)|~r2_hidden(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,esk5_2(X5,X6)))|~r2_hidden(X5,X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])])])).
cnf(c_0_11,plain,(X1=k1_xboole_0|X2=k1_setfam_1(X1)|r2_hidden(esk3_2(X1,X2),X3)|r2_hidden(esk3_2(X1,X2),X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,plain,(r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|r1_tarski(k1_setfam_1(X4),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_setfam_1])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk6_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk6_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_17,plain,(X2=k1_xboole_0|X1!=k1_xboole_0|X2!=k1_setfam_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18,plain,(X1=k1_setfam_1(k1_tarski(X2))|k1_tarski(X2)=k1_xboole_0|r2_hidden(esk3_2(k1_tarski(X2),X1),X2)|r2_hidden(esk3_2(k1_tarski(X2),X1),X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_19,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(r2_hidden(esk5_2(X1,k1_tarski(X1)),k1_tarski(X1))), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_21,plain,(r1_tarski(k1_setfam_1(X1),X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(r1_tarski(X1,X2)|r2_hidden(esk6_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(k1_setfam_1(X1)=k1_xboole_0|X1!=k1_xboole_0), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_24,plain,(X1=k1_xboole_0|X2=k1_setfam_1(X1)|r2_hidden(esk4_2(X1,X2),X1)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25,plain,(k1_setfam_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0|r2_hidden(esk3_2(k1_tarski(X1),X1),X1)), inference(ef,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(~r2_hidden(X1,X2)|~r2_hidden(X3,esk5_2(X1,X2))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27,plain,(esk5_2(X1,k1_tarski(X1))=X1), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29,plain,(r1_tarski(k1_setfam_1(X1),esk6_2(X1,X2))|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30,plain,(k1_setfam_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_23])).
fof(c_0_31, negated_conjecture, (~(![X1]:k1_setfam_1(k1_tarski(X1))=X1)), inference(assume_negation,[status(cth)],[t11_setfam_1])).
cnf(c_0_32,plain,(k1_setfam_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0|r2_hidden(esk4_2(k1_tarski(X1),X1),k1_tarski(X1))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,plain,(~r2_hidden(X1,k1_tarski(X2))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_12])])).
cnf(c_0_34,plain,(r1_tarski(X1,X2)|r2_hidden(esk6_2(X1,X2),X3)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_28, c_0_22])).
cnf(c_0_35,plain,(r1_tarski(k1_xboole_0,esk6_2(k1_xboole_0,X1))|r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_36, negated_conjecture, (k1_setfam_1(k1_tarski(esk7_0))!=esk7_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
cnf(c_0_37,plain,(X1=k1_xboole_0|X2=k1_setfam_1(X1)|~r2_hidden(esk3_2(X1,X2),X2)|~r2_hidden(esk3_2(X1,X2),esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_38,plain,(esk4_2(k1_tarski(X1),X1)=X1|k1_setfam_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_19, c_0_32])).
cnf(c_0_39,plain,(~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_33, c_0_12])).
cnf(c_0_40,plain,(r1_tarski(k1_xboole_0,X1)|r1_tarski(k1_xboole_0,X2)|r2_hidden(esk6_2(k1_xboole_0,X2),esk6_2(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41,negated_conjecture,(k1_setfam_1(k1_tarski(esk7_0))!=esk7_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42,plain,(k1_setfam_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_25])).
cnf(c_0_43,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44,negated_conjecture,(k1_tarski(esk7_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_43])).
cnf(c_0_46,plain,(r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_28, c_0_12])).
cnf(c_0_47,negated_conjecture,(~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_44]), c_0_45])).
cnf(c_0_48,negated_conjecture,(r2_hidden(esk7_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_44]), c_0_43])])).
cnf(c_0_49,negated_conjecture,($false), inference(spm,[status(thm)],[c_0_47, c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 37
# Proof object formula steps           : 13
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 295
# ...of these trivial                  : 8
# ...subsumed                          : 85
# ...remaining for further processing  : 202
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 13
# Generated clauses                    : 1729
# ...of the previous two non-trivial   : 1473
# Contextual simplify-reflections      : 56
# Paramodulations                      : 1681
# Factorizations                       : 8
# Equation resolutions                 : 40
# Current number of processed clauses  : 173
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 159
# Current number of unprocessed clauses: 865
# ...number of literals in the above   : 3888
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 4875
# Rec. Clause-clause subsumption calls : 3230
# Non-unit clause-clause subsumptions  : 146
# Unit Clause-clause subsumption calls : 240
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 24
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 30872

# -------------------------------------------------
# User time                : 0.036 s
# System time              : 0.008 s
# Total time               : 0.044 s
# Maximum resident set size: 43052 pages
