# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t120_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k2_relat_1(X2))=>k2_relat_1(k8_relat_1(X1,X2))=X1))), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', t120_relat_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', d3_tarski)).
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', d5_xboole_0)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', t3_boole)).
fof(d2_tarski, axiom, (![X1]:![X2]:![X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2)))), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', d2_tarski)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t7_tarski, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3)))))))), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', t7_tarski)).
fof(t12_setfam_1, axiom, (![X1]:![X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', t12_setfam_1)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', t48_xboole_1)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(t119_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k2_relat_1(k8_relat_1(X1,X2))=k3_xboole_0(k2_relat_1(X2),X1))), file('out/t120_relat_1__11_premises__rerun__rerun.E_input', t119_relat_1)).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k2_relat_1(X2))=>k2_relat_1(k8_relat_1(X1,X2))=X1)))), inference(assume_negation,[status(cth)],[t120_relat_1])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk3_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk3_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_13, negated_conjecture, ((v1_relat_1(esk2_0)&(r1_tarski(esk1_0,k2_relat_1(esk2_0))&k2_relat_1(k8_relat_1(esk1_0,esk2_0))!=esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk4_3(X5,X6,X7),X7)|(~r2_hidden(esk4_3(X5,X6,X7),X5)|r2_hidden(esk4_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk4_3(X5,X6,X7),X5)|r2_hidden(esk4_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk4_3(X5,X6,X7),X6)|r2_hidden(esk4_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
cnf(c_0_15,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(r1_tarski(esk1_0,k2_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(r2_hidden(X4,X2)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18,plain,(r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(r2_hidden(X1,k2_relat_1(esk2_0))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21,plain,(r1_tarski(X1,X2)|r2_hidden(esk3_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_23, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_24, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(X8=X5|X8=X6))|X7!=k2_tarski(X5,X6))&(((X8!=X5|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))&((X8!=X6|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))))&((((esk5_3(X5,X6,X7)!=X5|~r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6))&((esk5_3(X5,X6,X7)!=X6|~r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6)))&((r2_hidden(esk5_3(X5,X6,X7),X7)|(esk5_3(X5,X6,X7)=X5|esk5_3(X5,X6,X7)=X6))|X7=k2_tarski(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])])).
fof(c_0_25, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_26,negated_conjecture,(r1_tarski(X1,k2_relat_1(esk2_0))|~r2_hidden(esk3_2(X1,k2_relat_1(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_27,plain,(r2_hidden(esk3_2(k4_xboole_0(X1,X2),X3),X1)|r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28,plain,(~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, (![X5]:![X6]:![X8]:((r2_hidden(esk6_2(X5,X6),X6)|~r2_hidden(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,esk6_2(X5,X6)))|~r2_hidden(X5,X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])])])).
cnf(c_0_31,plain,(r2_hidden(X4,X1)|X1!=k2_tarski(X2,X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_33, plain, (![X3]:![X4]:k1_setfam_1(k2_tarski(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t12_setfam_1])).
cnf(c_0_34,negated_conjecture,(r1_tarski(k4_xboole_0(esk1_0,X1),k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35,plain,(~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36,plain,(r2_hidden(esk6_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37,plain,(r2_hidden(X4,X1)|X3!=X4|X1!=k1_enumset1(X2,X2,X3)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_38, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_39,plain,(k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_40, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_41, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k2_relat_1(k8_relat_1(X3,X4))=k3_xboole_0(k2_relat_1(X4),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t119_relat_1])])).
cnf(c_0_42,negated_conjecture,(r2_hidden(X1,k2_relat_1(esk2_0))|~r2_hidden(X1,k4_xboole_0(esk1_0,X2))), inference(spm,[status(thm)],[c_0_15, c_0_34])).
cnf(c_0_43,plain,(~r2_hidden(esk6_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_44,plain,(r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_37])])).
cnf(c_0_45,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk4_3(X2,X3,X1),X1)|r2_hidden(esk4_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_46,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47,plain,(k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_39, c_0_32])).
cnf(c_0_48,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49,plain,(k2_relat_1(k8_relat_1(X1,X2))=k3_xboole_0(k2_relat_1(X2),X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~r2_hidden(esk3_2(k4_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_28, c_0_21])).
cnf(c_0_51,negated_conjecture,(r2_hidden(esk3_2(k4_xboole_0(esk1_0,X1),X2),k2_relat_1(esk2_0))|r1_tarski(k4_xboole_0(esk1_0,X1),X2)), inference(spm,[status(thm)],[c_0_42, c_0_21])).
cnf(c_0_52,plain,(~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_53,plain,(k4_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_45])).
cnf(c_0_54,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_setfam_1(k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55,plain,(k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_32]), c_0_32])).
cnf(c_0_56,plain,(k2_relat_1(k8_relat_1(X1,X2))=k1_setfam_1(k1_enumset1(k2_relat_1(X2),k2_relat_1(X2),X1))|~v1_relat_1(X2)), inference(rw,[status(thm)],[c_0_49, c_0_47])).
cnf(c_0_57,negated_conjecture,(r1_tarski(k4_xboole_0(esk1_0,k2_relat_1(esk2_0)),X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_59,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_54])).
cnf(c_0_60,plain,(k4_xboole_0(k2_relat_1(X1),k4_xboole_0(k2_relat_1(X1),X2))=k2_relat_1(k8_relat_1(X2,X1))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_56, c_0_54])).
cnf(c_0_61,negated_conjecture,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(esk1_0,k2_relat_1(esk2_0)))), inference(spm,[status(thm)],[c_0_15, c_0_57])).
cnf(c_0_62,plain,(X1=k1_xboole_0|r2_hidden(esk4_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_45]), c_0_58])).
cnf(c_0_63,negated_conjecture,(k2_relat_1(k8_relat_1(esk1_0,esk2_0))!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_64,plain,(k2_relat_1(k8_relat_1(X1,X2))=k4_xboole_0(X1,k4_xboole_0(X1,k2_relat_1(X2)))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_65,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_66,negated_conjecture,(k4_xboole_0(esk1_0,k2_relat_1(esk2_0))=k1_xboole_0|r2_hidden(esk4_3(k1_xboole_0,X1,k4_xboole_0(esk1_0,k2_relat_1(esk2_0))),X2)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67,negated_conjecture,(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,k2_relat_1(esk2_0)))!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])])).
cnf(c_0_68,negated_conjecture,(k4_xboole_0(esk1_0,k2_relat_1(esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_52, c_0_66])).
cnf(c_0_69,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68]), c_0_29])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 70
# Proof object clause steps            : 47
# Proof object formula steps           : 23
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 24
# Processed clauses                    : 366
# ...of these trivial                  : 15
# ...subsumed                          : 204
# ...remaining for further processing  : 147
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 11
# Generated clauses                    : 2404
# ...of the previous two non-trivial   : 2008
# Contextual simplify-reflections      : 47
# Paramodulations                      : 2394
# Factorizations                       : 2
# Equation resolutions                 : 10
# Current number of processed clauses  : 105
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 80
# Current number of unprocessed clauses: 1592
# ...number of literals in the above   : 5806
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 1860
# Rec. Clause-clause subsumption calls : 1650
# Non-unit clause-clause subsumptions  : 187
# Unit Clause-clause subsumption calls : 77
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 49
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 42197

# -------------------------------------------------
# User time                : 0.048 s
# System time              : 0.000 s
# Total time               : 0.048 s
# Maximum resident set size: 45500 pages
