# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_yellow_6, conjecture, (![X1]:(l1_struct_0(X1)=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:((~(v2_struct_0(X3))&(v2_yellow_6(X3,X1,X2)&m1_yellow_6(X3,X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X3))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X4=X6&(X5=X7&r1_orders_2(X2,X4,X5)))=>r1_orders_2(X3,X6,X7)))))))))), file('out/t12_yellow_6__7_premises__rerun.E_input', t12_yellow_6)).
fof(d7_yellow_6, axiom, (![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>![X3]:(m1_yellow_6(X3,X1,X2)=>(v2_yellow_6(X3,X1,X2)<=>(v4_yellow_0(X3,X2)&m1_yellow_0(X3,X2))))))), file('out/t12_yellow_6__7_premises__rerun.E_input', d7_yellow_6)).
fof(dt_l1_waybel_0, axiom, (![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2)))), file('out/t12_yellow_6__7_premises__rerun.E_input', dt_l1_waybel_0)).
fof(t60_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:((v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X5=X3&(X6=X4&(r1_orders_2(X1,X3,X4)&r2_hidden(X5,u1_struct_0(X2)))))=>r1_orders_2(X2,X5,X6))))))))), file('out/t12_yellow_6__7_premises__rerun.E_input', t60_yellow_0)).
fof(dt_m1_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2)))), file('out/t12_yellow_6__7_premises__rerun.E_input', dt_m1_yellow_0)).
fof(d1_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t12_yellow_6__7_premises__rerun.E_input', d1_subset_1)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t12_yellow_6__7_premises__rerun.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t12_yellow_6__7_premises__rerun.E_input', dt_l1_orders_2)).
fof(c_0_8, negated_conjecture, (~(![X1]:(l1_struct_0(X1)=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:((~(v2_struct_0(X3))&(v2_yellow_6(X3,X1,X2)&m1_yellow_6(X3,X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X3))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X4=X6&(X5=X7&r1_orders_2(X2,X4,X5)))=>r1_orders_2(X3,X6,X7))))))))))), inference(assume_negation,[status(cth)],[t12_yellow_6])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:((((((v4_yellow_0(X6,X5)|~v2_yellow_6(X6,X4,X5))|~m1_yellow_6(X6,X4,X5))|~l1_waybel_0(X5,X4))|~l1_struct_0(X4))&((((m1_yellow_0(X6,X5)|~v2_yellow_6(X6,X4,X5))|~m1_yellow_6(X6,X4,X5))|~l1_waybel_0(X5,X4))|~l1_struct_0(X4)))&(((((~v4_yellow_0(X6,X5)|~m1_yellow_0(X6,X5))|v2_yellow_6(X6,X4,X5))|~m1_yellow_6(X6,X4,X5))|~l1_waybel_0(X5,X4))|~l1_struct_0(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_6])])])])])])).
fof(c_0_10, negated_conjecture, ((l1_struct_0(esk1_0)&((~v2_struct_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&(v2_yellow_6(esk3_0,esk1_0,esk2_0)&m1_yellow_6(esk3_0,esk1_0,esk2_0)))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(m1_subset_1(esk6_0,u1_struct_0(esk3_0))&(m1_subset_1(esk7_0,u1_struct_0(esk3_0))&((esk4_0=esk6_0&(esk5_0=esk7_0&r1_orders_2(esk2_0,esk4_0,esk5_0)))&~r1_orders_2(esk3_0,esk6_0,esk7_0)))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_11, plain, (![X3]:![X4]:(~l1_struct_0(X3)|(~l1_waybel_0(X4,X3)|l1_orders_2(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])])])).
fof(c_0_12, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(~l1_orders_2(X7)|((~v4_yellow_0(X8,X7)|~m1_yellow_0(X8,X7))|(~m1_subset_1(X9,u1_struct_0(X7))|(~m1_subset_1(X10,u1_struct_0(X7))|(~m1_subset_1(X11,u1_struct_0(X8))|(~m1_subset_1(X12,u1_struct_0(X8))|((X11!=X9|(X12!=X10|(~r1_orders_2(X7,X9,X10)|~r2_hidden(X11,u1_struct_0(X8)))))|r1_orders_2(X8,X11,X12))))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_yellow_0])])])])])).
fof(c_0_13, plain, (![X3]:![X4]:(~l1_orders_2(X3)|(~m1_yellow_0(X4,X3)|l1_orders_2(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])])])).
cnf(c_0_14,plain,(m1_yellow_0(X3,X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_yellow_6(X3,X1,X2)|~v2_yellow_6(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,negated_conjecture,(m1_yellow_6(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,negated_conjecture,(v2_yellow_6(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,negated_conjecture,(l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,plain,(l1_orders_2(X1)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,plain,(r1_orders_2(X1,X2,X3)|~r2_hidden(X2,u1_struct_0(X1))|~r1_orders_2(X4,X5,X6)|X3!=X6|X2!=X5|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~l1_orders_2(X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_21, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_subset_1])])])])])])).
fof(c_0_22, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_23, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_24,plain,(l1_orders_2(X1)|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,negated_conjecture,(m1_yellow_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_26,negated_conjecture,(l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_17]), c_0_18])])).
cnf(c_0_27,plain,(r1_orders_2(X1,X2,X3)|~r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X4)|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~r1_orders_2(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_20])])).
cnf(c_0_28,plain,(v1_xboole_0(X1)|r2_hidden(X2,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,negated_conjecture,(r1_orders_2(esk2_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_30,negated_conjecture,(esk5_0=esk7_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_31,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_32,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_33,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35,negated_conjecture,(l1_orders_2(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_36,plain,(v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,X2,X3)|~l1_orders_2(X4)|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~r1_orders_2(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_37,negated_conjecture,(r1_orders_2(esk2_0,esk4_0,esk7_0)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38,negated_conjecture,(m1_subset_1(esk7_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_39,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40,plain,(v4_yellow_0(X3,X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_yellow_6(X3,X1,X2)|~v2_yellow_6(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41,negated_conjecture,(m1_subset_1(esk6_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_42,negated_conjecture,(esk4_0=esk6_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43,negated_conjecture,(~v1_xboole_0(u1_struct_0(esk3_0))|~l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_44,negated_conjecture,(l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45,negated_conjecture,(~r1_orders_2(esk3_0,esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_46,negated_conjecture,(v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk4_0,esk7_0)|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk7_0,u1_struct_0(X1))|~m1_subset_1(esk4_0,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_26]), c_0_38]), c_0_39])])).
cnf(c_0_47,negated_conjecture,(v4_yellow_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_15]), c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_48,negated_conjecture,(m1_subset_1(esk7_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_49,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50,negated_conjecture,(~v1_xboole_0(u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44])])).
cnf(c_0_51,negated_conjecture,(~r1_orders_2(esk3_0,esk4_0,esk7_0)), inference(rw,[status(thm)],[c_0_45, c_0_42])).
cnf(c_0_52,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_25]), c_0_48]), c_0_49])]), c_0_50]), c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 36
# Proof object formula steps           : 17
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 69
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 66
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 17
# Contextual simplify-reflections      : 0
# Paramodulations                      : 17
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 15
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 179
# Rec. Clause-clause subsumption calls : 26
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2226

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 46900 pages
