# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_orders_2, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3]:![X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4)))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', free_g1_orders_2)).
fof(dt_k1_yellow_1, axiom, (![X1]:(v1_relat_2(k1_yellow_1(X1))&(v4_relat_2(k1_yellow_1(X1))&(v8_relat_2(k1_yellow_1(X1))&(v1_partfun1(k1_yellow_1(X1),X1)&m1_subset_1(k1_yellow_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', dt_k1_yellow_1)).
fof(dt_k2_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1)))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', dt_k2_yellow_1)).
fof(d1_yellow_1, axiom, (![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', d1_yellow_1)).
fof(t3_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t3_yellow_1)).
fof(abstractness_v1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', abstractness_v1_orders_2)).
fof(fc5_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1)))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', fc6_yellow_1)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', redefinition_r3_orders_2)).
fof(t16_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t16_yellow_0)).
fof(t8_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t8_yellow_0)).
fof(t19_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X1,X3))=>r1_tarski(X1,k3_xboole_0(X2,X3)))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t19_xboole_1)).
fof(t9_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r2_hidden(k3_xboole_0(X2,X3),X1)=>k11_lattice3(k2_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3)))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t9_yellow_1)).
fof(t17_xboole_1, axiom, (![X1]:![X2]:r1_tarski(k3_xboole_0(X1,X2),X1)), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t21_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>(v2_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r2_yellow_0(X1,k2_tarski(X2,X3))))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t21_yellow_0)).
fof(dt_k11_lattice3, axiom, (![X1]:![X2]:![X3]:((l1_orders_2(X1)&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1)))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', dt_k11_lattice3)).
fof(t12_yellow_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>(![X2]:![X3]:((r2_hidden(X2,X1)&r2_hidden(X3,X1))=>r2_hidden(k3_xboole_0(X2,X3),X1))=>v2_lattice3(k2_yellow_1(X1))))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t12_yellow_1)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t12_yellow_1__19_premises__rerun__rerun.E_input', t2_subset)).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:(((X5=X7|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5))))&((X6=X8|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])])])).
fof(c_0_20, plain, (![X2]:![X2]:![X2]:![X2]:![X2]:(v1_relat_2(k1_yellow_1(X2))&(v4_relat_2(k1_yellow_1(X2))&(v8_relat_2(k1_yellow_1(X2))&(v1_partfun1(k1_yellow_1(X2),X2)&m1_subset_1(k1_yellow_1(X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k1_yellow_1])])])).
fof(c_0_21, plain, (![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&l1_orders_2(k2_yellow_1(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k2_yellow_1])])])).
fof(c_0_22, plain, (![X2]:k2_yellow_1(X2)=g1_orders_2(X2,k1_yellow_1(X2))), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:(((((~r3_orders_2(k2_yellow_1(X4),X5,X6)|r1_tarski(X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4))&((((~r1_tarski(X5,X6)|r3_orders_2(k2_yellow_1(X4),X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])])])).
cnf(c_0_24,plain,(X2=X3|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|g1_orders_2(X2,X1)!=g1_orders_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,plain,(m1_subset_1(k1_yellow_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, (![X2]:(~l1_orders_2(X2)|(~v1_orders_2(X2)|X2=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
cnf(c_0_27,plain,(l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28,plain,(k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(v1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_30, plain, (![X4]:![X3]:![X2]:![X1]:(epred1_4(X1,X2,X3,X4)<=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2))))))), introduced(definition)).
cnf(c_0_31,plain,(v1_xboole_0(X1)|r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(X1=X2|g1_orders_2(X1,k1_yellow_1(X1))!=g1_orders_2(X2,X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,plain,(X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,plain,(l1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35,plain,(v1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_29, c_0_28])).
fof(c_0_36, plain, (![X2]:![X2]:![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&(v3_orders_2(k2_yellow_1(X2))&(v4_orders_2(k2_yellow_1(X2))&v5_orders_2(k2_yellow_1(X2)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc5_yellow_1])])])).
fof(c_0_37, plain, (![X2]:((~v2_struct_0(k2_yellow_1(X2))|v1_xboole_0(X2))&(v1_orders_2(k2_yellow_1(X2))|v1_xboole_0(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
fof(c_0_38, plain, (![X4]:![X3]:![X2]:![X1]:(epred1_4(X1,X2,X3,X4)=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2))))))), inference(split_equiv,[status(thm)],[c_0_30])).
fof(c_0_39, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_40,plain,(v1_xboole_0(X1)|r3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|~r1_tarski(X2,X3)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_41,plain,(u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33])]), c_0_34]), c_0_35])])).
cnf(c_0_42,plain,(v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43,plain,(v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_44, plain, (![X5]:![X6]:![X8]:![X6]:![X9]:((((m1_subset_1(esk1_2(X5,X6),u1_struct_0(X5))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_lattice3(X5,X6,esk1_2(X5,X6))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X6,X8)|r1_orders_2(X5,X8,esk1_2(X5,X6))))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))&(((((m1_subset_1(esk2_3(X5,X6,X9),u1_struct_0(X5))|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((((r1_lattice3(X5,X6,esk2_3(X5,X6,X9))|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~r1_orders_2(X5,esk2_3(X5,X6,X9),X9)|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])])).
cnf(c_0_45,plain,(v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46,plain,(v1_xboole_0(X1)|r1_tarski(X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r3_orders_2(k2_yellow_1(X1),X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_47, plain, (![X5]:![X6]:![X7]:![X8]:((((r1_orders_2(X8,X7,X6)|~r1_lattice3(X8,k2_tarski(X6,X5),X7))|~epred1_4(X8,X7,X6,X5))&((r1_orders_2(X8,X7,X5)|~r1_lattice3(X8,k2_tarski(X6,X5),X7))|~epred1_4(X8,X7,X6,X5)))&((((~r1_orders_2(X8,X7,X6)|~r1_orders_2(X8,X7,X5))|r1_lattice3(X8,k2_tarski(X6,X5),X7))|~epred1_4(X8,X7,X6,X5))&((((r1_orders_2(X8,X6,X7)|~r2_lattice3(X8,k2_tarski(X6,X5),X7))|~epred1_4(X8,X7,X6,X5))&((r1_orders_2(X8,X5,X7)|~r2_lattice3(X8,k2_tarski(X6,X5),X7))|~epred1_4(X8,X7,X6,X5)))&(((~r1_orders_2(X8,X6,X7)|~r1_orders_2(X8,X5,X7))|r2_lattice3(X8,k2_tarski(X6,X5),X7))|~epred1_4(X8,X7,X6,X5)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])).
fof(c_0_48, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>epred1_4(X1,X2,X3,X4)))))), inference(apply_def,[status(thm)],[t8_yellow_0, c_0_30])).
cnf(c_0_49,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50,plain,(r3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_41])).
cnf(c_0_51,plain,(v3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_42, c_0_28])).
cnf(c_0_52,plain,(v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_43, c_0_28])).
cnf(c_0_53,plain,(r2_yellow_0(X1,X2)|m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54,plain,(v5_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_45, c_0_28])).
cnf(c_0_55,plain,(v1_xboole_0(X1)|r1_tarski(X2,X3)|~r3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_56,plain,(r1_orders_2(X1,X2,X3)|~epred1_4(X1,X2,X3,X4)|~r1_lattice3(X1,k2_tarski(X3,X4),X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57,plain,(r2_yellow_0(X1,X2)|r1_lattice3(X1,X2,esk2_3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_58, plain, (![X5]:![X6]:![X7]:![X8]:(~l1_orders_2(X5)|(~m1_subset_1(X6,u1_struct_0(X5))|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|epred1_4(X5,X6,X7,X8)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_48])])])])])).
cnf(c_0_59,plain,(r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_60,plain,(r1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_34]), c_0_41]), c_0_41])]), c_0_52])).
cnf(c_0_61,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),X2)|m1_subset_1(esk2_3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3),X1)|~r1_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|~m1_subset_1(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_41]), c_0_54]), c_0_34])])).
fof(c_0_62, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X4,X6))|r1_tarski(X4,k3_xboole_0(X5,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t19_xboole_1])])).
cnf(c_0_63,plain,(r1_tarski(X1,X2)|v1_xboole_0(X3)|~r3_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_41]), c_0_41])).
cnf(c_0_64,plain,(v2_struct_0(X2)|r3_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r1_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_65,plain,(r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk2_3(X1,k2_tarski(X2,X3),X4),X2)|~epred1_4(X1,esk2_3(X1,k2_tarski(X2,X3),X4),X2,X3)|~r1_lattice3(X1,k2_tarski(X2,X3),X4)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_66,plain,(epred1_4(X1,X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_67,plain,(r1_orders_2(X1,X2,X4)|~epred1_4(X1,X2,X3,X4)|~r1_lattice3(X1,k2_tarski(X3,X4),X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_68,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),X2)|v1_xboole_0(X1)|~r1_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|~r1_tarski(esk2_3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3),X3)|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_54]), c_0_34]), c_0_41])]), c_0_61])).
cnf(c_0_69,plain,(r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70,plain,(r1_tarski(X1,X2)|v1_xboole_0(X3)|~r1_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_51]), c_0_34]), c_0_41]), c_0_41])]), c_0_52])).
cnf(c_0_71,plain,(r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk2_3(X1,k2_tarski(X2,X3),X4),X2)|~r1_lattice3(X1,k2_tarski(X2,X3),X4)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_53])).
cnf(c_0_72,plain,(r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk2_3(X1,k2_tarski(X2,X3),X4),X3)|~epred1_4(X1,esk2_3(X1,k2_tarski(X2,X3),X4),X2,X3)|~r1_lattice3(X1,k2_tarski(X2,X3),X4)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_67, c_0_57])).
cnf(c_0_73,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),X2)|v1_xboole_0(X1)|~r1_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,k3_xboole_0(X3,X4))|~r1_tarski(esk2_3(g1_orders_2(X1,k1_yellow_1(X1)),X2,k3_xboole_0(X3,X4)),X4)|~r1_tarski(esk2_3(g1_orders_2(X1,k1_yellow_1(X1)),X2,k3_xboole_0(X3,X4)),X3)|~m1_subset_1(k3_xboole_0(X3,X4),X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_74,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|r1_tarski(esk2_3(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3),X4),X2)|v1_xboole_0(X1)|~r1_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3),X4)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_54]), c_0_34]), c_0_41]), c_0_41]), c_0_41])]), c_0_61])).
cnf(c_0_75,plain,(r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk2_3(X1,k2_tarski(X2,X3),X4),X3)|~r1_lattice3(X1,k2_tarski(X2,X3),X4)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_66]), c_0_53])).
cnf(c_0_76,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|v1_xboole_0(X1)|~r1_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3),k3_xboole_0(X4,X2))|~r1_tarski(esk2_3(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3),k3_xboole_0(X4,X2)),X4)|~m1_subset_1(k3_xboole_0(X4,X2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_77,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|r1_tarski(esk2_3(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3),X4),X3)|v1_xboole_0(X1)|~r1_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3),X4)|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_75]), c_0_54]), c_0_34]), c_0_41]), c_0_41]), c_0_41])]), c_0_61])).
cnf(c_0_78,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|v1_xboole_0(X1)|~r1_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3),k3_xboole_0(X3,X2))|~m1_subset_1(k3_xboole_0(X3,X2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_79,plain,(r1_lattice3(X1,k2_tarski(X3,X4),X2)|~epred1_4(X1,X2,X3,X4)|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_80, plain, (![X4]:![X5]:![X6]:(v1_xboole_0(X4)|(~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4)))|(~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4)))|(~r2_hidden(k3_xboole_0(X5,X6),X4)|k11_lattice3(k2_yellow_1(X4),X5,X6)=k3_xboole_0(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_yellow_1])])])])])])).
cnf(c_0_81,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|v1_xboole_0(X1)|~epred1_4(g1_orders_2(X1,k1_yellow_1(X1)),k3_xboole_0(X3,X2),X2,X3)|~r1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),k3_xboole_0(X3,X2),X3)|~r1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),k3_xboole_0(X3,X2),X2)|~m1_subset_1(k3_xboole_0(X3,X2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
fof(c_0_82, plain, (![X3]:![X4]:r1_tarski(k3_xboole_0(X3,X4),X3)), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_83, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_84,plain,(k11_lattice3(k2_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(k3_xboole_0(X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_85,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|v1_xboole_0(X1)|~r1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),k3_xboole_0(X3,X2),X3)|~r1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),k3_xboole_0(X3,X2),X2)|~m1_subset_1(k3_xboole_0(X3,X2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_66]), c_0_34]), c_0_41]), c_0_41]), c_0_41])])).
cnf(c_0_86,plain,(r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_87,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
fof(c_0_88, plain, (![X4]:![X5]:![X6]:(((~v2_lattice3(X4)|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|r2_yellow_0(X4,k2_tarski(X5,X6)))))|(~v5_orders_2(X4)|~l1_orders_2(X4)))&(((m1_subset_1(esk3_1(X4),u1_struct_0(X4))|v2_lattice3(X4))|(~v5_orders_2(X4)|~l1_orders_2(X4)))&(((m1_subset_1(esk4_1(X4),u1_struct_0(X4))|v2_lattice3(X4))|(~v5_orders_2(X4)|~l1_orders_2(X4)))&((~r2_yellow_0(X4,k2_tarski(esk3_1(X4),esk4_1(X4)))|v2_lattice3(X4))|(~v5_orders_2(X4)|~l1_orders_2(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_yellow_0])])])])])])])).
fof(c_0_89, plain, (![X4]:![X5]:![X6]:((~l1_orders_2(X4)|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k11_lattice3(X4,X5,X6),u1_struct_0(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k11_lattice3])])).
cnf(c_0_90,plain,(k11_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)=k3_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(k3_xboole_0(X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_28]), c_0_28]), c_0_28])).
fof(c_0_91, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>(![X2]:![X3]:((r2_hidden(X2,X1)&r2_hidden(X3,X1))=>r2_hidden(k3_xboole_0(X2,X3),X1))=>v2_lattice3(k2_yellow_1(X1)))))), inference(assume_negation,[status(cth)],[t12_yellow_1])).
cnf(c_0_92,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|v1_xboole_0(X1)|~r1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),k3_xboole_0(X3,X2),X2)|~m1_subset_1(k3_xboole_0(X3,X2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_60]), c_0_86])])).
cnf(c_0_93,plain,(r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_94,plain,(v2_lattice3(X1)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_95,plain,(v2_lattice3(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_96,plain,(m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_97,plain,(k11_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)=k3_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(k3_xboole_0(X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_41]), c_0_41])).
fof(c_0_98, negated_conjecture, (![X5]:![X6]:(~v1_xboole_0(esk5_0)&(((~r2_hidden(X5,esk5_0)|~r2_hidden(X6,esk5_0))|r2_hidden(k3_xboole_0(X5,X6),esk5_0))&~v2_lattice3(k2_yellow_1(esk5_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_91])])])])])])])).
cnf(c_0_99,plain,(v2_lattice3(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,k2_tarski(esk3_1(X1),esk4_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100,plain,(r2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)),k2_tarski(X2,X3))|v1_xboole_0(X1)|~m1_subset_1(k3_xboole_0(X3,X2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_60]), c_0_93])])).
cnf(c_0_101,plain,(v2_lattice3(g1_orders_2(X1,k1_yellow_1(X1)))|m1_subset_1(esk4_1(g1_orders_2(X1,k1_yellow_1(X1))),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_41]), c_0_54]), c_0_34])])).
cnf(c_0_102,plain,(v2_lattice3(g1_orders_2(X1,k1_yellow_1(X1)))|m1_subset_1(esk3_1(g1_orders_2(X1,k1_yellow_1(X1))),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_41]), c_0_54]), c_0_34])])).
cnf(c_0_103,plain,(v1_xboole_0(X1)|m1_subset_1(k3_xboole_0(X2,X3),X1)|~r2_hidden(k3_xboole_0(X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_41]), c_0_34]), c_0_41]), c_0_41])])).
cnf(c_0_104,negated_conjecture,(r2_hidden(k3_xboole_0(X1,X2),esk5_0)|~r2_hidden(X2,esk5_0)|~r2_hidden(X1,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_105,negated_conjecture,(~v1_xboole_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_106,negated_conjecture,(~v2_lattice3(k2_yellow_1(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_107,plain,(v2_lattice3(g1_orders_2(X1,k1_yellow_1(X1)))|v1_xboole_0(X1)|~m1_subset_1(k3_xboole_0(esk3_1(g1_orders_2(X1,k1_yellow_1(X1))),esk4_1(g1_orders_2(X1,k1_yellow_1(X1)))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_54]), c_0_34])]), c_0_87]), c_0_101]), c_0_102])).
cnf(c_0_108,negated_conjecture,(m1_subset_1(k3_xboole_0(X1,X2),esk5_0)|~r2_hidden(X2,esk5_0)|~r2_hidden(X1,esk5_0)|~m1_subset_1(X2,esk5_0)|~m1_subset_1(X1,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105])).
cnf(c_0_109,negated_conjecture,(~v2_lattice3(g1_orders_2(esk5_0,k1_yellow_1(esk5_0)))), inference(rw,[status(thm)],[c_0_106, c_0_28])).
fof(c_0_110, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_111,negated_conjecture,(~r2_hidden(esk4_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)|~r2_hidden(esk3_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)|~m1_subset_1(esk4_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)|~m1_subset_1(esk3_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109]), c_0_105])).
cnf(c_0_112,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_113,negated_conjecture,(~r2_hidden(esk3_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)|~m1_subset_1(esk4_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)|~m1_subset_1(esk3_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_105])).
cnf(c_0_114,negated_conjecture,(~m1_subset_1(esk4_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)|~m1_subset_1(esk3_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_112]), c_0_105])).
cnf(c_0_115,negated_conjecture,(~m1_subset_1(esk3_1(g1_orders_2(esk5_0,k1_yellow_1(esk5_0))),esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_101]), c_0_109])).
cnf(c_0_116,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_102]), c_0_109]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 117
# Proof object clause steps            : 74
# Proof object formula steps           : 43
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 94
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 46
# Processed clauses                    : 275
# ...of these trivial                  : 4
# ...subsumed                          : 76
# ...remaining for further processing  : 195
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 0
# Generated clauses                    : 314
# ...of the previous two non-trivial   : 274
# Contextual simplify-reflections      : 27
# Paramodulations                      : 290
# Factorizations                       : 0
# Equation resolutions                 : 24
# Current number of processed clauses  : 139
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 121
# Current number of unprocessed clauses: 76
# ...number of literals in the above   : 587
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 9064
# Rec. Clause-clause subsumption calls : 686
# Non-unit clause-clause subsumptions  : 115
# Unit Clause-clause subsumption calls : 29
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 18150

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.004 s
# Total time               : 0.036 s
# Maximum resident set size: 43328 pages
