# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t130_relat_1, conjecture, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(X1,X3)))), file('out/t130_relat_1__18_premises__rerun.E_input', t130_relat_1)).
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t130_relat_1__18_premises__rerun.E_input', t1_xboole_1)).
fof(t125_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(k2_relat_1(X2),X1)=>k8_relat_1(X1,X2)=X2))), file('out/t130_relat_1__18_premises__rerun.E_input', t125_relat_1)).
fof(t116_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>r1_tarski(k2_relat_1(k8_relat_1(X1,X2)),X1))), file('out/t130_relat_1__18_premises__rerun.E_input', t116_relat_1)).
fof(dt_k8_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>v1_relat_1(k8_relat_1(X1,X2)))), file('out/t130_relat_1__18_premises__rerun.E_input', dt_k8_relat_1)).
fof(t126_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>k8_relat_1(k2_relat_1(X1),X1)=X1)), file('out/t130_relat_1__18_premises__rerun.E_input', t126_relat_1)).
fof(t71_relat_1, axiom, (![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1)), file('out/t130_relat_1__18_premises__rerun.E_input', t71_relat_1)).
fof(dt_k6_relat_1, axiom, (![X1]:v1_relat_1(k6_relat_1(X1))), file('out/t130_relat_1__18_premises__rerun.E_input', dt_k6_relat_1)).
fof(t129_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>k8_relat_1(X2,k8_relat_1(X1,X3))=k8_relat_1(X1,X3)))), file('out/t130_relat_1__18_premises__rerun.E_input', t129_relat_1)).
fof(t12_setfam_1, axiom, (![X1]:![X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), file('out/t130_relat_1__18_premises__rerun.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t130_relat_1__18_premises__rerun.E_input', t70_enumset1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t130_relat_1__18_premises__rerun.E_input', t7_xboole_1)).
fof(t127_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(k3_xboole_0(X1,X2),X3))), file('out/t130_relat_1__18_premises__rerun.E_input', t127_relat_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t130_relat_1__18_premises__rerun.E_input', t71_enumset1)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t130_relat_1__18_premises__rerun.E_input', t48_xboole_1)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t130_relat_1__18_premises__rerun.E_input', d10_xboole_0)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t130_relat_1__18_premises__rerun.E_input', t2_xboole_1)).
fof(t43_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3))), file('out/t130_relat_1__18_premises__rerun.E_input', t43_xboole_1)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t130_relat_1__18_premises__rerun.E_input', t3_boole)).
fof(c_0_19, negated_conjecture, (~(![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(X1,X3))))), inference(assume_negation,[status(cth)],[t130_relat_1])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_21, negated_conjecture, ((v1_relat_1(esk3_0)&(r1_tarski(esk1_0,esk2_0)&k8_relat_1(esk1_0,k8_relat_1(esk2_0,esk3_0))!=k8_relat_1(esk1_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
fof(c_0_22, plain, (![X3]:![X4]:(~v1_relat_1(X4)|(~r1_tarski(k2_relat_1(X4),X3)|k8_relat_1(X3,X4)=X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t125_relat_1])])).
cnf(c_0_23,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,negated_conjecture,(r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25,plain,(k8_relat_1(X1,X2)=X2|~r1_tarski(k2_relat_1(X2),X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,negated_conjecture,(r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_27, plain, (![X3]:![X4]:(~v1_relat_1(X4)|r1_tarski(k2_relat_1(k8_relat_1(X3,X4)),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t116_relat_1])])).
fof(c_0_28, plain, (![X3]:![X4]:(~v1_relat_1(X4)|v1_relat_1(k8_relat_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relat_1])])).
fof(c_0_29, plain, (![X2]:(~v1_relat_1(X2)|k8_relat_1(k2_relat_1(X2),X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t126_relat_1])])).
fof(c_0_30, plain, (![X2]:![X2]:(k1_relat_1(k6_relat_1(X2))=X2&k2_relat_1(k6_relat_1(X2))=X2)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t71_relat_1])])])).
fof(c_0_31, plain, (![X2]:v1_relat_1(k6_relat_1(X2))), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_32,negated_conjecture,(k8_relat_1(esk2_0,X1)=X1|~r1_tarski(k2_relat_1(X1),esk1_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33,plain,(r1_tarski(k2_relat_1(k8_relat_1(X1,X2)),X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34,plain,(v1_relat_1(k8_relat_1(X1,X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35,plain,(k8_relat_1(k2_relat_1(X1),X1)=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36,plain,(k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37,plain,(v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|(~r1_tarski(X4,X5)|k8_relat_1(X5,k8_relat_1(X4,X6))=k8_relat_1(X4,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_relat_1])])).
cnf(c_0_39,negated_conjecture,(k8_relat_1(esk2_0,k8_relat_1(esk1_0,X1))=k8_relat_1(esk1_0,X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_40,plain,(k8_relat_1(X1,k6_relat_1(X1))=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
cnf(c_0_41,plain,(k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(X2,X3)|~r1_tarski(X2,X1)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42,negated_conjecture,(k8_relat_1(esk2_0,k6_relat_1(esk1_0))=k6_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_37])])).
fof(c_0_43, plain, (![X3]:![X4]:k1_setfam_1(k2_tarski(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_44, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_45,negated_conjecture,(k8_relat_1(X1,k6_relat_1(esk1_0))=k6_relat_1(esk1_0)|~r1_tarski(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_37])])).
fof(c_0_46, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_47, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k8_relat_1(X4,k8_relat_1(X5,X6))=k8_relat_1(k3_xboole_0(X4,X5),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t127_relat_1])])).
cnf(c_0_48,plain,(k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_50, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_51, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_52, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
fof(c_0_53, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_54, plain, (![X4]:![X5]:![X6]:(~r1_tarski(X4,k2_xboole_0(X5,X6))|r1_tarski(k4_xboole_0(X4,X5),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_55,negated_conjecture,(r1_tarski(esk1_0,X1)|~r1_tarski(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_45]), c_0_36]), c_0_37])])).
cnf(c_0_56,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57,plain,(k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(k3_xboole_0(X1,X2),X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58,plain,(k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_59,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64,negated_conjecture,(r1_tarski(esk1_0,k2_xboole_0(esk2_0,X1))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_65,plain,(k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(k1_setfam_1(k2_enumset1(X1,X1,X1,X2)),X3)|~v1_relat_1(X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58]), c_0_59])).
cnf(c_0_66,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_setfam_1(k2_enumset1(X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_58]), c_0_59])).
cnf(c_0_67,plain,(X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68,negated_conjecture,(r1_tarski(k4_xboole_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
fof(c_0_69, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_70,plain,(k8_relat_1(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)=k8_relat_1(X1,k8_relat_1(X2,X3))|~v1_relat_1(X3)), inference(rw,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_71,negated_conjecture,(k4_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_72,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_73,negated_conjecture,(k8_relat_1(esk1_0,k8_relat_1(esk2_0,esk3_0))!=k8_relat_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_74,negated_conjecture,(k8_relat_1(esk1_0,k8_relat_1(esk2_0,X1))=k8_relat_1(esk1_0,X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
cnf(c_0_75,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_76,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 77
# Proof object clause steps            : 38
# Proof object formula steps           : 39
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 21
# Processed clauses                    : 82
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 139
# ...of the previous two non-trivial   : 107
# Contextual simplify-reflections      : 1
# Paramodulations                      : 137
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 46
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 29
# Current number of unprocessed clauses: 50
# ...number of literals in the above   : 144
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 155
# Rec. Clause-clause subsumption calls : 138
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2863

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 43000 pages
