# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t131_relat_1, conjecture, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k8_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k10_xtuple_0(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2)))))), file('out/t131_relat_1__2_premises__rerun.E_input', t131_relat_1)).
fof(d13_xtuple_0, axiom, (![X1]:![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1)))), file('out/t131_relat_1__2_premises__rerun.E_input', d13_xtuple_0)).
fof(d14_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:![X3]:(X3=k8_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2)))))), file('out/t131_relat_1__2_premises__rerun.E_input', d14_relat_1)).
fof(c_0_3, negated_conjecture, (~(![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k8_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k10_xtuple_0(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))))), inference(assume_negation,[status(cth)],[t131_relat_1])).
fof(c_0_4, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(esk5_3(X5,X6,X7),X7),X5))|X6!=k10_xtuple_0(X5))&((~r2_hidden(k4_tarski(X9,X7),X5)|r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5)))&(((~r2_hidden(esk6_2(X5,X6),X6)|~r2_hidden(k4_tarski(X11,esk6_2(X5,X6)),X5))|X6=k10_xtuple_0(X5))&((r2_hidden(esk6_2(X5,X6),X6)|r2_hidden(k4_tarski(esk7_2(X5,X6),esk6_2(X5,X6)),X5))|X6=k10_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_xtuple_0])])])])])])])).
fof(c_0_5, negated_conjecture, (![X8]:(v1_relat_1(esk3_0)&((~r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))|(~r2_hidden(X8,k10_xtuple_0(esk3_0))|(~r2_hidden(k4_tarski(esk1_0,X8),esk3_0)|~r2_hidden(X8,esk2_0))))&((r2_hidden(esk4_0,k10_xtuple_0(esk3_0))|r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0)))&((r2_hidden(k4_tarski(esk1_0,esk4_0),esk3_0)|r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0)))&(r2_hidden(esk4_0,esk2_0)|r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])])])).
cnf(c_0_6,plain,(r2_hidden(X3,X1)|X1!=k10_xtuple_0(X2)|~r2_hidden(k4_tarski(X4,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:![X7]:![X8]:![X13]:((((((r2_hidden(k4_tarski(X9,esk8_4(X6,X7,X8,X9)),X6)|~r2_hidden(X9,X8))|X8!=k8_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk8_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|X8!=k8_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(k4_tarski(X9,X11),X6)|~r2_hidden(X11,X7))|r2_hidden(X9,X8))|X8!=k8_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(esk9_3(X6,X7,X8),X8)|(~r2_hidden(k4_tarski(esk9_3(X6,X7,X8),X13),X6)|~r2_hidden(X13,X7)))|X8=k8_relat_1(X6,X7))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk9_3(X6,X7,X8),esk10_3(X6,X7,X8)),X6)|r2_hidden(esk9_3(X6,X7,X8),X8))|X8=k8_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk10_3(X6,X7,X8),X7)|r2_hidden(esk9_3(X6,X7,X8),X8))|X8=k8_relat_1(X6,X7))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])])).
cnf(c_0_8,negated_conjecture,(~r2_hidden(X1,esk2_0)|~r2_hidden(k4_tarski(esk1_0,X1),esk3_0)|~r2_hidden(X1,k10_xtuple_0(esk3_0))|~r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(r2_hidden(X1,k10_xtuple_0(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(r2_hidden(k4_tarski(X4,esk8_4(X1,X3,X2,X4)),X1)|~v1_relat_1(X1)|X2!=k8_relat_1(X1,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,negated_conjecture,(~r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))|~r2_hidden(k4_tarski(esk1_0,X1),esk3_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12,plain,(r2_hidden(k4_tarski(X1,esk8_4(X2,X3,k8_relat_1(X2,X3),X1)),X2)|~r2_hidden(X1,k8_relat_1(X2,X3))|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_13,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,plain,(r2_hidden(esk8_4(X1,X3,X2,X4),X3)|~v1_relat_1(X1)|X2!=k8_relat_1(X1,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,plain,(r2_hidden(X4,X2)|~v1_relat_1(X1)|X2!=k8_relat_1(X1,X3)|~r2_hidden(X5,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16,negated_conjecture,(~r2_hidden(esk8_4(esk3_0,X1,k8_relat_1(esk3_0,X1),esk1_0),esk2_0)|~r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))|~r2_hidden(esk1_0,k8_relat_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_17,plain,(r2_hidden(esk8_4(X1,X2,k8_relat_1(X1,X2),X3),X2)|~r2_hidden(X3,k8_relat_1(X1,X2))|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_18,plain,(r2_hidden(X1,k8_relat_1(X2,X3))|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_19,negated_conjecture,(r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))|r2_hidden(k4_tarski(esk1_0,esk4_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_20,negated_conjecture,(r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))|r2_hidden(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_21,negated_conjecture,(~r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_13])])).
cnf(c_0_22,negated_conjecture,(r2_hidden(esk1_0,k8_relat_1(esk3_0,esk2_0))|r2_hidden(esk1_0,k8_relat_1(esk3_0,X1))|~r2_hidden(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_13])])).
cnf(c_0_23,negated_conjecture,(r2_hidden(esk4_0,esk2_0)), inference(sr,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 25
# Proof object clause steps            : 18
# Proof object formula steps           : 7
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 40
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 39
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 26
# ...of the previous two non-trivial   : 22
# Contextual simplify-reflections      : 1
# Paramodulations                      : 17
# Factorizations                       : 0
# Equation resolutions                 : 5
# Current number of processed clauses  : 15
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 12
# ...number of literals in the above   : 37
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 150
# Rec. Clause-clause subsumption calls : 98
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1834

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 42964 pages
