# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t134_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:![X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', t134_zfmisc_1)).
fof(t118_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,X2)=>(r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))))), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', t118_zfmisc_1)).
fof(t117_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3)))))), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', t117_zfmisc_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', d3_tarski)).
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', d4_xboole_0)).
fof(t123_zfmisc_1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', t123_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X1)=X1), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', idempotence_k3_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t134_zfmisc_1__14_premises__rerun__rerun__rerun__rerun.E_input', d10_xboole_0)).
fof(c_0_9, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:![X6]:((r1_tarski(k2_zfmisc_1(X4,X6),k2_zfmisc_1(X5,X6))|~r1_tarski(X4,X5))&(r1_tarski(k2_zfmisc_1(X6,X4),k2_zfmisc_1(X6,X5))|~r1_tarski(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t118_zfmisc_1])])])])])).
fof(c_0_11, negated_conjecture, ((k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(esk3_0!=esk5_0|esk4_0!=esk6_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:(((~r1_tarski(k2_zfmisc_1(X5,X4),k2_zfmisc_1(X6,X4))|r1_tarski(X5,X6))|X4=k1_xboole_0)&((~r1_tarski(k2_zfmisc_1(X4,X5),k2_zfmisc_1(X4,X6))|r1_tarski(X5,X6))|X4=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])])])).
cnf(c_0_13,plain,(r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,negated_conjecture,(k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk2_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk2_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_16,plain,(X1=k1_xboole_0|r1_tarski(X2,X3)|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(r1_tarski(k2_zfmisc_1(X1,esk6_0),k2_zfmisc_1(esk3_0,esk4_0))|~r1_tarski(X1,esk5_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18,negated_conjecture,(esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|~r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X8]:k2_zfmisc_1(k3_xboole_0(X5,X6),k3_xboole_0(X7,X8))=k3_xboole_0(k2_zfmisc_1(X5,X7),k2_zfmisc_1(X6,X8))), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_21,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,negated_conjecture,(r1_tarski(esk6_0,esk4_0)|~r1_tarski(esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
cnf(c_0_23,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|r2_hidden(esk1_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24,plain,(k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, (![X3]:k3_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_26, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_27,negated_conjecture,(r2_hidden(X1,esk4_0)|~r1_tarski(esk3_0,esk5_0)|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28,plain,(k3_xboole_0(X1,X2)=X1|r2_hidden(esk1_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_23])).
cnf(c_0_29,plain,(r2_hidden(X4,X3)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(X1,esk5_0),k3_xboole_0(X2,esk6_0))=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),k3_xboole_0(X2,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_14]), c_0_24])).
cnf(c_0_31,plain,(k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,plain,(X1=k3_xboole_0(X2,X3)|~r2_hidden(esk1_3(X2,X3,X1),X3)|~r2_hidden(esk1_3(X2,X3,X1),X2)|~r2_hidden(esk1_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34,negated_conjecture,(k3_xboole_0(esk6_0,X1)=esk6_0|r2_hidden(esk1_3(esk6_0,X1,esk6_0),esk4_0)|~r1_tarski(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_36,plain,(r1_tarski(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_37,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),k3_xboole_0(esk4_0,esk6_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_38,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk6_0|~r1_tarski(esk3_0,esk5_0)|~r2_hidden(esk1_3(esk6_0,esk4_0,esk6_0),esk6_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_32])).
cnf(c_0_39,negated_conjecture,(k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0))=k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),k3_xboole_0(X1,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_40,plain,(r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_41,plain,(r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_42, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_43,negated_conjecture,(k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,esk6_0))=k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk6_0)), inference(spm,[status(thm)],[c_0_37, c_0_31])).
cnf(c_0_44,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk6_0|~r1_tarski(esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_28]), c_0_32])])).
cnf(c_0_45,negated_conjecture,(esk5_0=k1_xboole_0|r1_tarski(k3_xboole_0(X1,esk6_0),X2)|~r1_tarski(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),k3_xboole_0(X1,esk4_0)),k2_zfmisc_1(esk5_0,X2))), inference(spm,[status(thm)],[c_0_16, c_0_39])).
cnf(c_0_46,plain,(r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48,negated_conjecture,(r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk6_0))|~r1_tarski(esk5_0,X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_49,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk6_0)=k2_zfmisc_1(esk3_0,esk6_0)|~r1_tarski(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)|~r1_tarski(esk3_0,esk5_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_44]), c_0_31])).
cnf(c_0_51,negated_conjecture,(esk5_0=k1_xboole_0|r1_tarski(k3_xboole_0(X1,esk6_0),k3_xboole_0(X1,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_13]), c_0_46])])).
cnf(c_0_52,negated_conjecture,(esk6_0=esk4_0|~r1_tarski(esk4_0,esk6_0)|~r1_tarski(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_47, c_0_22])).
cnf(c_0_53,negated_conjecture,(r1_tarski(esk4_0,esk6_0)|~r1_tarski(esk5_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_48]), c_0_18])).
cnf(c_0_54,negated_conjecture,(k2_zfmisc_1(esk3_0,esk6_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(esk3_0,esk5_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_31])).
cnf(c_0_55,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56,negated_conjecture,(esk5_0=k1_xboole_0|r1_tarski(esk6_0,k3_xboole_0(esk4_0,esk6_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_31]), c_0_32])).
cnf(c_0_57,plain,(X1=k1_xboole_0|r1_tarski(X2,X3)|~r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_58,negated_conjecture,(esk4_0!=esk6_0|esk3_0!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_59,negated_conjecture,(esk6_0=esk4_0|~r1_tarski(esk3_0,esk5_0)|~r1_tarski(esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_60,negated_conjecture,(r1_tarski(X1,esk6_0)|~r1_tarski(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(esk3_0,esk4_0))|~r1_tarski(esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_54]), c_0_18])).
cnf(c_0_61,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_62,plain,(r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_63,plain,(r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_46, c_0_32])).
cnf(c_0_64,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk6_0|esk5_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_56]), c_0_46])])).
cnf(c_0_65,negated_conjecture,(esk6_0=k1_xboole_0|r1_tarski(esk5_0,X1)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk6_0))), inference(spm,[status(thm)],[c_0_57, c_0_14])).
cnf(c_0_66,negated_conjecture,(~r1_tarski(esk3_0,esk5_0)|~r1_tarski(esk5_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_47])).
cnf(c_0_67,negated_conjecture,(r1_tarski(esk4_0,esk6_0)|~r1_tarski(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_68,negated_conjecture,(r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk5_0,X1))|~r1_tarski(esk6_0,X1)), inference(spm,[status(thm)],[c_0_62, c_0_14])).
cnf(c_0_69,negated_conjecture,(esk5_0=k1_xboole_0|r1_tarski(esk6_0,esk4_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_70,negated_conjecture,(esk6_0=k1_xboole_0|~r1_tarski(esk3_0,esk5_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_54]), c_0_61])]), c_0_66])).
cnf(c_0_71,negated_conjecture,(esk6_0=esk4_0|~r1_tarski(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_52, c_0_67])).
cnf(c_0_72,negated_conjecture,(esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_73,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk5_0,X1),k3_xboole_0(esk6_0,X2))=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),k3_xboole_0(esk4_0,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_14]), c_0_24])).
cnf(c_0_74,negated_conjecture,(esk5_0=k1_xboole_0|r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk5_0,esk4_0))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75,negated_conjecture,(~r1_tarski(esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
cnf(c_0_76,negated_conjecture,(k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,X1))=k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),k3_xboole_0(esk4_0,X1))), inference(spm,[status(thm)],[c_0_73, c_0_31])).
cnf(c_0_77,negated_conjecture,(esk5_0=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_74]), c_0_72]), c_0_75])).
cnf(c_0_78,negated_conjecture,(r1_tarski(X1,k3_xboole_0(esk4_0,esk6_0))|~r1_tarski(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk6_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_43]), c_0_18])).
cnf(c_0_79,negated_conjecture,(k2_zfmisc_1(k1_xboole_0,k3_xboole_0(esk6_0,X1))=k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),k3_xboole_0(esk4_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_77]), c_0_77]), c_0_32])).
cnf(c_0_80,negated_conjecture,(k2_zfmisc_1(k1_xboole_0,esk6_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_14, c_0_77])).
cnf(c_0_81,negated_conjecture,(r1_tarski(esk4_0,k3_xboole_0(esk4_0,esk6_0))|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_78, c_0_48])).
cnf(c_0_82,negated_conjecture,(r1_tarski(k3_xboole_0(esk4_0,esk6_0),X1)|~r1_tarski(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk6_0),k2_zfmisc_1(esk3_0,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_43]), c_0_18])).
cnf(c_0_83,plain,(r2_hidden(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k2_zfmisc_1(k3_xboole_0(X4,X2),k3_xboole_0(X5,X3)))), inference(spm,[status(thm)],[c_0_35, c_0_24])).
cnf(c_0_84,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),k3_xboole_0(esk4_0,esk6_0))=k2_zfmisc_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_31]), c_0_80])).
cnf(c_0_85,negated_conjecture,(k2_zfmisc_1(X1,esk6_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk6_0))|~r1_tarski(X1,esk5_0)), inference(spm,[status(thm)],[c_0_47, c_0_17])).
cnf(c_0_86,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk4_0|~r1_tarski(k3_xboole_0(esk4_0,esk6_0),esk4_0)|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_47, c_0_81])).
cnf(c_0_87,negated_conjecture,(r1_tarski(k3_xboole_0(esk4_0,esk6_0),esk4_0)|~r1_tarski(k3_xboole_0(esk3_0,esk5_0),esk5_0)), inference(spm,[status(thm)],[c_0_82, c_0_17])).
cnf(c_0_88,negated_conjecture,(r2_hidden(X1,k2_zfmisc_1(esk3_0,esk6_0))|~r2_hidden(X1,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_89,negated_conjecture,(k2_zfmisc_1(X1,esk6_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(X1,esk5_0)|~r1_tarski(esk5_0,X1)), inference(spm,[status(thm)],[c_0_85, c_0_48])).
cnf(c_0_90,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk4_0|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))|~r1_tarski(k3_xboole_0(esk3_0,esk5_0),esk5_0)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_91,negated_conjecture,(r1_tarski(X1,k2_zfmisc_1(esk3_0,esk6_0))|~r2_hidden(esk2_2(X1,k2_zfmisc_1(esk3_0,esk6_0)),k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_40, c_0_88])).
cnf(c_0_92,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk5_0,X1),k3_xboole_0(X2,esk6_0))=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),k3_xboole_0(X2,esk4_0))), inference(spm,[status(thm)],[c_0_30, c_0_32])).
cnf(c_0_93,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk3_0,X1),k3_xboole_0(esk4_0,X2))=k2_zfmisc_1(k3_xboole_0(X3,X1),k3_xboole_0(esk6_0,X2))|~r1_tarski(X3,esk5_0)|~r1_tarski(esk5_0,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_89]), c_0_24])).
cnf(c_0_94,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk4_0|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_46])])).
cnf(c_0_95,negated_conjecture,(r1_tarski(k2_zfmisc_1(esk5_0,X1),k2_zfmisc_1(esk3_0,esk4_0))|~r1_tarski(X1,esk6_0)), inference(spm,[status(thm)],[c_0_62, c_0_14])).
cnf(c_0_96,negated_conjecture,(r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_91, c_0_36])).
cnf(c_0_97,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk4_0)=k2_zfmisc_1(esk5_0,esk4_0)|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))|~r1_tarski(k3_xboole_0(esk3_0,esk5_0),esk5_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_90]), c_0_31])).
cnf(c_0_98,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk5_0,X1),esk6_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(X1,esk5_0)|~r1_tarski(esk5_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_31]), c_0_31]), c_0_31])).
cnf(c_0_99,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_37, c_0_94])).
cnf(c_0_100,negated_conjecture,(r1_tarski(k2_zfmisc_1(k1_xboole_0,X1),k2_zfmisc_1(esk3_0,esk4_0))|~r1_tarski(X1,esk6_0)), inference(rw,[status(thm)],[c_0_95, c_0_77])).
cnf(c_0_101,negated_conjecture,(r1_tarski(esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_96]), c_0_18])).
cnf(c_0_102,negated_conjecture,(r1_tarski(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),k3_xboole_0(X1,esk4_0)),k2_zfmisc_1(esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_46]), c_0_39])).
cnf(c_0_103,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk4_0)=k2_zfmisc_1(esk5_0,esk4_0)|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_46])])).
cnf(c_0_104,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk4_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_32]), c_0_61])])).
cnf(c_0_105,negated_conjecture,(r1_tarski(esk6_0,k3_xboole_0(esk4_0,esk6_0))|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_78, c_0_13])).
cnf(c_0_106,negated_conjecture,(r2_hidden(X1,esk6_0)|~r1_tarski(esk5_0,esk3_0)|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_21, c_0_53])).
cnf(c_0_107,negated_conjecture,(r1_tarski(k2_zfmisc_1(k1_xboole_0,esk4_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_108,negated_conjecture,(r1_tarski(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk4_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_102, c_0_31])).
cnf(c_0_109,negated_conjecture,(k2_zfmisc_1(esk5_0,esk4_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(esk5_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_110,plain,(r2_hidden(X4,X2)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_111,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk6_0|~r1_tarski(k3_xboole_0(esk4_0,esk6_0),esk6_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_47, c_0_105])).
cnf(c_0_112,negated_conjecture,(X1=k3_xboole_0(X2,esk6_0)|~r1_tarski(esk5_0,esk3_0)|~r2_hidden(esk1_3(X2,esk6_0,X1),esk4_0)|~r2_hidden(esk1_3(X2,esk6_0,X1),X2)|~r2_hidden(esk1_3(X2,esk6_0,X1),X1)), inference(spm,[status(thm)],[c_0_33, c_0_106])).
cnf(c_0_113,negated_conjecture,(r1_tarski(k1_xboole_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_107]), c_0_72])).
cnf(c_0_114,negated_conjecture,(r1_tarski(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),esk4_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_77]), c_0_32])).
cnf(c_0_115,negated_conjecture,(k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(k1_xboole_0,esk4_0)|~r1_tarski(k1_xboole_0,k3_xboole_0(k1_xboole_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_77]), c_0_77]), c_0_77]), c_0_32])).
cnf(c_0_116,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_110])).
cnf(c_0_117,negated_conjecture,(r2_hidden(X1,k3_xboole_0(esk4_0,esk6_0))|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_21, c_0_105])).
cnf(c_0_118,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk6_0|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_46])])).
cnf(c_0_119,negated_conjecture,(X1=k3_xboole_0(X2,esk6_0)|~r2_hidden(esk1_3(X2,esk6_0,X1),esk4_0)|~r2_hidden(esk1_3(X2,esk6_0,X1),X2)|~r2_hidden(esk1_3(X2,esk6_0,X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_77]), c_0_113])])).
cnf(c_0_120,negated_conjecture,(r1_tarski(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),esk4_0),k2_zfmisc_1(k1_xboole_0,esk4_0))|~r1_tarski(k1_xboole_0,k3_xboole_0(k1_xboole_0,esk3_0))), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_121,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),esk4_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(k1_xboole_0,k3_xboole_0(k1_xboole_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_77]), c_0_32]), c_0_77]), c_0_77]), c_0_32])).
cnf(c_0_122,negated_conjecture,(r2_hidden(X1,esk4_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_123,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk6_0)=k2_zfmisc_1(esk3_0,esk6_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_43, c_0_118])).
cnf(c_0_124,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_118]), c_0_31])).
cnf(c_0_125,negated_conjecture,(k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,esk6_0))=k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_77]), c_0_32])).
cnf(c_0_126,negated_conjecture,(k3_xboole_0(esk4_0,esk6_0)=esk4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_28]), c_0_28])).
cnf(c_0_127,negated_conjecture,(r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(k1_xboole_0,esk4_0))|~r1_tarski(k1_xboole_0,k3_xboole_0(k1_xboole_0,esk3_0))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_128,negated_conjecture,(~r1_tarski(esk3_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_75, c_0_77])).
cnf(c_0_129,negated_conjecture,(r1_tarski(esk6_0,X1)|r2_hidden(esk2_2(esk6_0,X1),esk4_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_122, c_0_36])).
cnf(c_0_130,negated_conjecture,(k2_zfmisc_1(esk3_0,esk6_0)=k2_zfmisc_1(esk3_0,esk4_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_31])).
cnf(c_0_131,negated_conjecture,(esk6_0=k1_xboole_0|r1_tarski(k1_xboole_0,X1)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk6_0))), inference(rw,[status(thm)],[c_0_65, c_0_77])).
cnf(c_0_132,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),esk6_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_133,negated_conjecture,(~r1_tarski(k1_xboole_0,k3_xboole_0(k1_xboole_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_127]), c_0_72]), c_0_128])).
cnf(c_0_134,negated_conjecture,(r1_tarski(esk6_0,esk4_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_40, c_0_129])).
cnf(c_0_135,negated_conjecture,(r1_tarski(X1,esk6_0)|~r1_tarski(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(esk3_0,esk4_0))|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_130]), c_0_18])).
cnf(c_0_136,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk5_0,X1),k3_xboole_0(X2,esk6_0))=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),k3_xboole_0(esk4_0,X2))), inference(spm,[status(thm)],[c_0_73, c_0_32])).
cnf(c_0_137,negated_conjecture,(esk6_0=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_61])]), c_0_133])).
cnf(c_0_138,negated_conjecture,(esk6_0=k1_xboole_0|r1_tarski(X1,esk5_0)|~r1_tarski(k2_zfmisc_1(X1,esk6_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_57, c_0_14])).
cnf(c_0_139,negated_conjecture,(esk6_0=esk4_0|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))|~r1_tarski(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_47, c_0_134])).
cnf(c_0_140,negated_conjecture,(r1_tarski(esk4_0,esk6_0)|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_135, c_0_61])).
cnf(c_0_141,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),k3_xboole_0(X2,k1_xboole_0))=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),k3_xboole_0(esk4_0,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_136, c_0_77]), c_0_137])).
cnf(c_0_142,negated_conjecture,(k3_xboole_0(k1_xboole_0,esk4_0)=esk4_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_137]), c_0_32])).
cnf(c_0_143,negated_conjecture,(esk6_0=k1_xboole_0|r1_tarski(esk3_0,esk5_0)|~r1_tarski(esk6_0,esk4_0)), inference(spm,[status(thm)],[c_0_138, c_0_62])).
cnf(c_0_144,negated_conjecture,(esk6_0=esk4_0|~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_145,negated_conjecture,(k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_80, c_0_137])).
cnf(c_0_146,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)=k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_31]), c_0_32]), c_0_142])).
cnf(c_0_147,negated_conjecture,(~r1_tarski(esk3_0,k3_xboole_0(esk3_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_61])]), c_0_72]), c_0_75])).
cnf(c_0_148,negated_conjecture,(r1_tarski(esk3_0,X1)|~r1_tarski(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k2_zfmisc_1(X1,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_145]), c_0_72])).
cnf(c_0_149,negated_conjecture,(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,esk3_0),esk4_0)=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_31]), c_0_145])).
cnf(c_0_150,negated_conjecture,(~r1_tarski(esk3_0,k3_xboole_0(k1_xboole_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_147, c_0_77]), c_0_32])).
cnf(c_0_151,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_61])]), c_0_150]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 152
# Proof object clause steps            : 133
# Proof object formula steps           : 19
# Proof object conjectures             : 112
# Proof object clause conjectures      : 109
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 94
# Proof object simplifying inferences  : 96
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 2025
# ...of these trivial                  : 70
# ...subsumed                          : 1276
# ...remaining for further processing  : 679
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 100
# Backward-rewritten                   : 433
# Generated clauses                    : 12579
# ...of the previous two non-trivial   : 11600
# Contextual simplify-reflections      : 657
# Paramodulations                      : 12546
# Factorizations                       : 20
# Equation resolutions                 : 13
# Current number of processed clauses  : 122
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 88
# Current number of unprocessed clauses: 1173
# ...number of literals in the above   : 3328
# Current number of archived formulas  : 0
# Current number of archived clauses   : 555
# Clause-clause subsumption calls (NU) : 35988
# Rec. Clause-clause subsumption calls : 28996
# Non-unit clause-clause subsumptions  : 1503
# Unit Clause-clause subsumption calls : 685
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 103
# BW rewrite match successes           : 63
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 200982

# -------------------------------------------------
# User time                : 0.188 s
# System time              : 0.000 s
# Total time               : 0.188 s
# Maximum resident set size: 42972 pages
