# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d13_xtuple_0, axiom, (![X1]:![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1)))), file('out/t134_relat_1__3_premises__rerun.E_input', d13_xtuple_0)).
fof(d12_xtuple_0, axiom, (![X1]:![X2]:(X2=k9_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t134_relat_1__3_premises__rerun.E_input', d12_xtuple_0)).
fof(d14_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:![X3]:(X3=k8_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2)))))), file('out/t134_relat_1__3_premises__rerun.E_input', d14_relat_1)).
fof(t134_relat_1, conjecture, (![X1]:(v1_relat_1(X1)=>k8_relat_1(X1,k10_xtuple_0(X1))=k9_xtuple_0(X1))), file('out/t134_relat_1__3_premises__rerun.E_input', t134_relat_1)).
fof(c_0_4, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(esk1_3(X5,X6,X7),X7),X5))|X6!=k10_xtuple_0(X5))&((~r2_hidden(k4_tarski(X9,X7),X5)|r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5)))&(((~r2_hidden(esk2_2(X5,X6),X6)|~r2_hidden(k4_tarski(X11,esk2_2(X5,X6)),X5))|X6=k10_xtuple_0(X5))&((r2_hidden(esk2_2(X5,X6),X6)|r2_hidden(k4_tarski(esk3_2(X5,X6),esk2_2(X5,X6)),X5))|X6=k10_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_xtuple_0])])])])])])])).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk4_3(X5,X6,X7)),X5))|X6!=k9_xtuple_0(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k9_xtuple_0(X5)))&(((~r2_hidden(esk5_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk5_2(X5,X6),X11),X5))|X6=k9_xtuple_0(X5))&((r2_hidden(esk5_2(X5,X6),X6)|r2_hidden(k4_tarski(esk5_2(X5,X6),esk6_2(X5,X6)),X5))|X6=k9_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_xtuple_0])])])])])])])).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:![X7]:![X8]:![X13]:((((((r2_hidden(k4_tarski(X9,esk7_4(X6,X7,X8,X9)),X6)|~r2_hidden(X9,X8))|X8!=k8_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk7_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|X8!=k8_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(k4_tarski(X9,X11),X6)|~r2_hidden(X11,X7))|r2_hidden(X9,X8))|X8!=k8_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(esk8_3(X6,X7,X8),X8)|(~r2_hidden(k4_tarski(esk8_3(X6,X7,X8),X13),X6)|~r2_hidden(X13,X7)))|X8=k8_relat_1(X6,X7))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk8_3(X6,X7,X8),esk9_3(X6,X7,X8)),X6)|r2_hidden(esk8_3(X6,X7,X8),X8))|X8=k8_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk9_3(X6,X7,X8),X7)|r2_hidden(esk8_3(X6,X7,X8),X8))|X8=k8_relat_1(X6,X7))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])])).
cnf(c_0_7,plain,(r2_hidden(X3,X1)|X1!=k10_xtuple_0(X2)|~r2_hidden(k4_tarski(X4,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(X1=k9_xtuple_0(X2)|r2_hidden(k4_tarski(esk5_2(X2,X1),esk6_2(X2,X1)),X2)|r2_hidden(esk5_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(r2_hidden(X4,X2)|~v1_relat_1(X1)|X2!=k8_relat_1(X1,X3)|~r2_hidden(X5,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(X1=k9_xtuple_0(X2)|r2_hidden(esk5_2(X2,X1),X1)|r2_hidden(esk6_2(X2,X1),X3)|X3!=k10_xtuple_0(X2)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_11,plain,(X1=k9_xtuple_0(X2)|r2_hidden(esk5_2(X2,X1),X1)|r2_hidden(esk5_2(X2,X1),X3)|X3!=k8_relat_1(X2,X4)|~v1_relat_1(X2)|~r2_hidden(esk6_2(X2,X1),X4)), inference(spm,[status(thm)],[c_0_9, c_0_8])).
cnf(c_0_12,plain,(X1=k9_xtuple_0(X2)|r2_hidden(esk6_2(X2,X1),k10_xtuple_0(X2))|r2_hidden(esk5_2(X2,X1),X1)), inference(er,[status(thm)],[c_0_10])).
fof(c_0_13, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>k8_relat_1(X1,k10_xtuple_0(X1))=k9_xtuple_0(X1)))), inference(assume_negation,[status(cth)],[t134_relat_1])).
cnf(c_0_14,plain,(X1=k9_xtuple_0(X2)|r2_hidden(esk5_2(X2,X1),X3)|r2_hidden(esk5_2(X2,X1),X1)|X3!=k8_relat_1(X2,k10_xtuple_0(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_15, negated_conjecture, ((v1_relat_1(esk10_0)&k8_relat_1(esk10_0,k10_xtuple_0(esk10_0))!=k9_xtuple_0(esk10_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_16,plain,(r2_hidden(k4_tarski(X4,esk7_4(X1,X3,X2,X4)),X1)|~v1_relat_1(X1)|X2!=k8_relat_1(X1,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17,plain,(X1=k9_xtuple_0(X2)|r2_hidden(esk5_2(X2,X1),k8_relat_1(X2,k10_xtuple_0(X2)))|r2_hidden(esk5_2(X2,X1),X1)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_18,negated_conjecture,(v1_relat_1(esk10_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19,plain,(X1=k9_xtuple_0(X2)|~r2_hidden(k4_tarski(esk5_2(X2,X1),X3),X2)|~r2_hidden(esk5_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_20,plain,(r2_hidden(k4_tarski(X1,esk7_4(X2,X3,k8_relat_1(X2,X3),X1)),X2)|~v1_relat_1(X2)|~r2_hidden(X1,k8_relat_1(X2,X3))), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_21,negated_conjecture,(X1=k9_xtuple_0(esk10_0)|r2_hidden(esk5_2(esk10_0,X1),k8_relat_1(esk10_0,k10_xtuple_0(esk10_0)))|r2_hidden(esk5_2(esk10_0,X1),X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,negated_conjecture,(k8_relat_1(esk10_0,k10_xtuple_0(esk10_0))!=k9_xtuple_0(esk10_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(X1=k9_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(esk5_2(X2,X1),k8_relat_1(X2,X3))|~r2_hidden(esk5_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_24,negated_conjecture,(r2_hidden(esk5_2(esk10_0,k8_relat_1(esk10_0,k10_xtuple_0(esk10_0))),k8_relat_1(esk10_0,k10_xtuple_0(esk10_0)))), inference(sr,[status(thm)],[inference(ef,[status(thm)],[c_0_21]), c_0_22])).
cnf(c_0_25,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_18]), c_0_24])]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 26
# Proof object clause steps            : 17
# Proof object formula steps           : 9
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 213
# ...of these trivial                  : 3
# ...subsumed                          : 28
# ...remaining for further processing  : 182
# Other redundant clauses eliminated   : 309
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 40
# Backward-rewritten                   : 0
# Generated clauses                    : 2332
# ...of the previous two non-trivial   : 1972
# Contextual simplify-reflections      : 37
# Paramodulations                      : 1944
# Factorizations                       : 14
# Equation resolutions                 : 374
# Current number of processed clauses  : 126
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 122
# Current number of unprocessed clauses: 1452
# ...number of literals in the above   : 9721
# Current number of archived formulas  : 0
# Current number of archived clauses   : 56
# Clause-clause subsumption calls (NU) : 3456
# Rec. Clause-clause subsumption calls : 601
# Non-unit clause-clause subsumptions  : 105
# Unit Clause-clause subsumption calls : 51
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 61330

# -------------------------------------------------
# User time                : 0.072 s
# System time              : 0.000 s
# Total time               : 0.072 s
# Maximum resident set size: 45276 pages
