# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t136_orders_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:((v6_orders_2(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r3_orders_1(u1_orders_2(X1),X2)))), file('out/t136_orders_2__10_premises__rerun.E_input', t136_orders_2)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t136_orders_2__10_premises__rerun.E_input', t3_subset)).
fof(t94_orders_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>((r4_relat_2(X3,X1)&r1_tarski(X2,X1))=>r4_relat_2(X3,X2)))), file('out/t136_orders_2__10_premises__rerun.E_input', t94_orders_1)).
fof(t95_orders_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>((r8_relat_2(X3,X1)&r1_tarski(X2,X1))=>r8_relat_2(X3,X2)))), file('out/t136_orders_2__10_premises__rerun.E_input', t95_orders_1)).
fof(d8_orders_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(r3_orders_1(X1,X2)<=>(r1_relat_2(X1,X2)&(r8_relat_2(X1,X2)&(r4_relat_2(X1,X2)&r6_relat_2(X1,X2))))))), file('out/t136_orders_2__10_premises__rerun.E_input', d8_orders_1)).
fof(t92_orders_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(r7_relat_2(X2,X1)<=>(r1_relat_2(X2,X1)&r6_relat_2(X2,X1))))), file('out/t136_orders_2__10_premises__rerun.E_input', t92_orders_1)).
fof(d11_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v6_orders_2(X2,X1)<=>r7_relat_2(u1_orders_2(X1),X2))))), file('out/t136_orders_2__10_premises__rerun.E_input', d11_orders_2)).
fof(d6_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v5_orders_2(X1)<=>r4_relat_2(u1_orders_2(X1),u1_struct_0(X1))))), file('out/t136_orders_2__10_premises__rerun.E_input', d6_orders_2)).
fof(d5_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v4_orders_2(X1)<=>r8_relat_2(u1_orders_2(X1),u1_struct_0(X1))))), file('out/t136_orders_2__10_premises__rerun.E_input', d5_orders_2)).
fof(cc1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('out/t136_orders_2__10_premises__rerun.E_input', cc1_relset_1)).
fof(dt_u1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))), file('out/t136_orders_2__10_premises__rerun.E_input', dt_u1_orders_2)).
fof(c_0_11, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:((v6_orders_2(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r3_orders_1(u1_orders_2(X1),X2))))), inference(assume_negation,[status(cth)],[t136_orders_2])).
fof(c_0_12, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_13, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&((v6_orders_2(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&~r3_orders_1(u1_orders_2(esk1_0),esk2_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|((~r4_relat_2(X6,X4)|~r1_tarski(X5,X4))|r4_relat_2(X6,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_orders_1])])).
cnf(c_0_15,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|((~r8_relat_2(X6,X4)|~r1_tarski(X5,X4))|r8_relat_2(X6,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_orders_1])])).
fof(c_0_18, plain, (![X3]:![X4]:![X4]:((((r1_relat_2(X3,X4)|~r3_orders_1(X3,X4))|~v1_relat_1(X3))&(((r8_relat_2(X3,X4)|~r3_orders_1(X3,X4))|~v1_relat_1(X3))&(((r4_relat_2(X3,X4)|~r3_orders_1(X3,X4))|~v1_relat_1(X3))&((r6_relat_2(X3,X4)|~r3_orders_1(X3,X4))|~v1_relat_1(X3)))))&(((~r1_relat_2(X3,X4)|(~r8_relat_2(X3,X4)|(~r4_relat_2(X3,X4)|~r6_relat_2(X3,X4))))|r3_orders_1(X3,X4))|~v1_relat_1(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_orders_1])])])])])])).
fof(c_0_19, plain, (![X3]:![X4]:((((r1_relat_2(X4,X3)|~r7_relat_2(X4,X3))|~v1_relat_1(X4))&((r6_relat_2(X4,X3)|~r7_relat_2(X4,X3))|~v1_relat_1(X4)))&(((~r1_relat_2(X4,X3)|~r6_relat_2(X4,X3))|r7_relat_2(X4,X3))|~v1_relat_1(X4)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t92_orders_1])])])).
fof(c_0_20, plain, (![X3]:![X4]:((((~v6_orders_2(X4,X3)|r7_relat_2(u1_orders_2(X3),X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_orders_2(X3))&(((~r7_relat_2(u1_orders_2(X3),X4)|v6_orders_2(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_orders_2])])])])])])).
cnf(c_0_21,plain,(r4_relat_2(X1,X2)|~r1_tarski(X2,X3)|~r4_relat_2(X1,X3)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,negated_conjecture,(r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_23, plain, (![X2]:(((~v5_orders_2(X2)|r4_relat_2(u1_orders_2(X2),u1_struct_0(X2)))|~l1_orders_2(X2))&((~r4_relat_2(u1_orders_2(X2),u1_struct_0(X2))|v5_orders_2(X2))|~l1_orders_2(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_orders_2])])])).
cnf(c_0_24,plain,(r8_relat_2(X1,X2)|~r1_tarski(X2,X3)|~r8_relat_2(X1,X3)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, (![X2]:(((~v4_orders_2(X2)|r8_relat_2(u1_orders_2(X2),u1_struct_0(X2)))|~l1_orders_2(X2))&((~r8_relat_2(u1_orders_2(X2),u1_struct_0(X2))|v4_orders_2(X2))|~l1_orders_2(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_orders_2])])])).
cnf(c_0_26,plain,(r3_orders_1(X1,X2)|~v1_relat_1(X1)|~r6_relat_2(X1,X2)|~r4_relat_2(X1,X2)|~r8_relat_2(X1,X2)|~r1_relat_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(r6_relat_2(X1,X2)|~v1_relat_1(X1)|~r7_relat_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(r1_relat_2(X1,X2)|~v1_relat_1(X1)|~r7_relat_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29,plain,(r7_relat_2(u1_orders_2(X1),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v6_orders_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30,negated_conjecture,(v6_orders_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32,negated_conjecture,(r4_relat_2(X1,esk2_0)|~r4_relat_2(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_33,plain,(r4_relat_2(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35,negated_conjecture,(r8_relat_2(X1,esk2_0)|~r8_relat_2(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_24, c_0_22])).
cnf(c_0_36,plain,(r8_relat_2(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_39, plain, (![X2]:(~l1_orders_2(X2)|m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_40,plain,(r3_orders_1(X1,X2)|~r7_relat_2(X1,X2)|~r8_relat_2(X1,X2)|~r4_relat_2(X1,X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_41,negated_conjecture,(r7_relat_2(u1_orders_2(esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_16]), c_0_30]), c_0_31])])).
cnf(c_0_42,negated_conjecture,(~r3_orders_1(u1_orders_2(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_43,negated_conjecture,(r4_relat_2(u1_orders_2(esk1_0),esk2_0)|~v1_relat_1(u1_orders_2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_31]), c_0_34])])).
cnf(c_0_44,negated_conjecture,(r8_relat_2(u1_orders_2(esk1_0),esk2_0)|~v1_relat_1(u1_orders_2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_31]), c_0_37])])).
cnf(c_0_45,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46,plain,(m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47,negated_conjecture,(~v1_relat_1(u1_orders_2(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43]), c_0_44])).
cnf(c_0_48,plain,(v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_49,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_31])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 27
# Proof object formula steps           : 23
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 74
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 29
# ...of the previous two non-trivial   : 19
# Contextual simplify-reflections      : 4
# Paramodulations                      : 29
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 45
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 4
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 263
# Rec. Clause-clause subsumption calls : 152
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2428

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 47696 pages
