# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t137_enumset1, conjecture, (![X1]:![X2]:![X3]:k2_xboole_0(k2_tarski(X2,X1),k2_tarski(X3,X1))=k1_enumset1(X1,X2,X3)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t137_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t98_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t71_enumset1)).
fof(t100_enumset1, axiom, (![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X3,X1)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t100_enumset1)).
fof(t69_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t69_enumset1)).
fof(t98_enumset1, axiom, (![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t98_enumset1)).
fof(t102_enumset1, axiom, (![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t102_enumset1)).
fof(l53_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', l53_enumset1)).
fof(t44_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t44_enumset1)).
fof(t42_enumset1, axiom, (![X1]:k2_enumset1(X1,X1,X1,X1)=k1_tarski(X1)), file('out/t137_enumset1__11_premises__rerun__rerun.E_input', t42_enumset1)).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:![X3]:k2_xboole_0(k2_tarski(X2,X1),k2_tarski(X3,X1))=k1_enumset1(X1,X2,X3))), inference(assume_negation,[status(cth)],[t137_enumset1])).
fof(c_0_12, negated_conjecture, (k2_xboole_0(k2_tarski(esk2_0,esk1_0),k2_tarski(esk3_0,esk1_0))!=k1_enumset1(esk1_0,esk2_0,esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_13, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_14, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:k1_enumset1(X4,X5,X6)=k1_enumset1(X5,X6,X4)), inference(variable_rename,[status(thm)],[t100_enumset1])).
cnf(c_0_17,negated_conjecture,(k2_xboole_0(k2_tarski(esk2_0,esk1_0),k2_tarski(esk3_0,esk1_0))!=k1_enumset1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X6,X8,X5,X7)), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_22,plain,(k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:k1_enumset1(X4,X5,X6)=k1_enumset1(X4,X6,X5)), inference(variable_rename,[status(thm)],[t98_enumset1])).
fof(c_0_24, plain, (![X4]:![X5]:![X6]:k1_enumset1(X4,X5,X6)=k1_enumset1(X6,X5,X4)), inference(variable_rename,[status(thm)],[t102_enumset1])).
cnf(c_0_25,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k4_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0)))!=k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_18]), c_0_19]), c_0_20]), c_0_20]), c_0_20]), c_0_20])).
cnf(c_0_26,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X2,X2,X3,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_20]), c_0_20])).
cnf(c_0_28,plain,(k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29,plain,(k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X5,X6),k2_tarski(X7,X8))), inference(variable_rename,[status(thm)],[l53_enumset1])).
fof(c_0_31, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X5),k1_enumset1(X6,X7,X8))), inference(variable_rename,[status(thm)],[t44_enumset1])).
fof(c_0_32, plain, (![X2]:k2_enumset1(X2,X2,X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t42_enumset1])).
cnf(c_0_33,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k4_xboole_0(k2_enumset1(esk1_0,esk3_0,esk3_0,esk3_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0)))!=k2_enumset1(esk2_0,esk1_0,esk3_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26]), c_0_26]), c_0_26])).
cnf(c_0_34,plain,(k2_enumset1(X1,X2,X3,X2)=k2_enumset1(X1,X1,X3,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X1,X1,X3,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_20]), c_0_20])).
cnf(c_0_36,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X3,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_20]), c_0_20])).
cnf(c_0_37,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(spm,[status(thm)],[c_0_26, c_0_26])).
cnf(c_0_38,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40,plain,(k2_enumset1(X1,X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k4_xboole_0(k2_enumset1(esk1_0,esk1_0,esk3_0,esk3_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0)))!=k2_enumset1(esk2_0,esk2_0,esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_34]), c_0_35])).
cnf(c_0_42,plain,(k2_enumset1(X1,X2,X3,X3)=k2_enumset1(X1,X1,X2,X3)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,plain,(k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k2_enumset1(X1,X1,X1,X2),k4_xboole_0(k2_enumset1(X3,X3,X3,X4),k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_18]), c_0_18]), c_0_19]), c_0_20]), c_0_20]), c_0_20])).
cnf(c_0_44,plain,(k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k2_enumset1(X1,X1,X1,X1),k4_xboole_0(k2_enumset1(X2,X2,X3,X4),k2_enumset1(X1,X1,X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40]), c_0_19]), c_0_20])).
cnf(c_0_45,negated_conjecture,(k2_enumset1(esk2_0,esk1_0,esk1_0,esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42]), c_0_43])).
cnf(c_0_46,plain,(k2_enumset1(X1,X2,X2,X3)=k2_enumset1(X1,X1,X2,X3)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 25
# Proof object formula steps           : 23
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 7
# Processed clauses                    : 121
# ...of these trivial                  : 0
# ...subsumed                          : 82
# ...remaining for further processing  : 39
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 1582
# ...of the previous two non-trivial   : 1373
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1582
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 27
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 1266
# ...number of literals in the above   : 1266
# Current number of archived formulas  : 0
# Current number of archived clauses   : 14
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 397
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 970
# BW rewrite match successes           : 323
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5316

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.008 s
# Total time               : 0.020 s
# Maximum resident set size: 42988 pages
