# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t137_enumset1, conjecture, (![X1]:![X2]:![X3]:k2_xboole_0(k2_tarski(X2,X1),k2_tarski(X3,X1))=k1_enumset1(X1,X2,X3)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t137_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t71_enumset1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t125_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t125_enumset1)).
fof(t113_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X3,X1)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t113_enumset1)).
fof(l57_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_tarski(X4,X5))), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', l57_enumset1)).
fof(l53_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', l53_enumset1)).
fof(t107_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t107_enumset1)).
fof(t105_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X3,X4,X2)), file('out/t137_enumset1__17_premises__rerun__rerun.E_input', t105_enumset1)).
fof(c_0_12, negated_conjecture, (~(![X1]:![X2]:![X3]:k2_xboole_0(k2_tarski(X2,X1),k2_tarski(X3,X1))=k1_enumset1(X1,X2,X3))), inference(assume_negation,[status(cth)],[t137_enumset1])).
fof(c_0_13, negated_conjecture, (k2_xboole_0(k2_tarski(esk2_0,esk1_0),k2_tarski(esk3_0,esk1_0))!=k1_enumset1(esk1_0,esk2_0,esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_14, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_15, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_16, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
cnf(c_0_18,negated_conjecture,(k2_xboole_0(k2_tarski(esk2_0,esk1_0),k2_tarski(esk3_0,esk1_0))!=k1_enumset1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_24,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k5_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0),k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0))))!=k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_22]), c_0_22]), c_0_22]), c_0_22])).
cnf(c_0_25,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_26, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X8,X7,X6,X5)), inference(variable_rename,[status(thm)],[t125_enumset1])).
cnf(c_0_27,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k5_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0),k3_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0))))!=k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_29, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X6,X8,X7,X5)), inference(variable_rename,[status(thm)],[t113_enumset1])).
fof(c_0_30, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k3_enumset1(X6,X7,X8,X9,X10)=k2_xboole_0(k1_enumset1(X6,X7,X8),k2_tarski(X9,X10))), inference(variable_rename,[status(thm)],[l57_enumset1])).
fof(c_0_31, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X5,X6),k2_tarski(X7,X8))), inference(variable_rename,[status(thm)],[l53_enumset1])).
cnf(c_0_32,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k5_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0),k3_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0))))!=k2_enumset1(esk3_0,esk2_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_34, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X5,X8,X7,X6)), inference(variable_rename,[status(thm)],[t107_enumset1])).
cnf(c_0_35,plain,(k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_tarski(X4,X5))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k5_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0),k3_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0))))!=k2_enumset1(esk2_0,esk1_0,esk1_0,esk3_0)), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39,plain,(k3_enumset1(X1,X2,X3,X4,X5)=k5_xboole_0(k2_enumset1(X1,X1,X2,X3),k5_xboole_0(k2_enumset1(X4,X4,X4,X5),k3_xboole_0(k2_enumset1(X4,X4,X4,X5),k2_enumset1(X1,X1,X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_22]), c_0_22]), c_0_22])).
cnf(c_0_40,plain,(k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k2_enumset1(X1,X1,X1,X2),k5_xboole_0(k2_enumset1(X3,X3,X3,X4),k3_xboole_0(k2_enumset1(X3,X3,X3,X4),k2_enumset1(X1,X1,X1,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_19]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_22]), c_0_22]), c_0_22])).
fof(c_0_41, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X5,X7,X8,X6)), inference(variable_rename,[status(thm)],[t105_enumset1])).
cnf(c_0_42,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k5_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0),k3_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk1_0))))!=k2_enumset1(esk2_0,esk3_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43,plain,(k5_xboole_0(k2_enumset1(X1,X1,X2,X3),k5_xboole_0(k2_enumset1(X4,X4,X4,X5),k3_xboole_0(k2_enumset1(X1,X1,X2,X3),k2_enumset1(X4,X4,X4,X5))))=k3_enumset1(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_39, c_0_25])).
cnf(c_0_44,plain,(k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(rw,[status(thm)],[c_0_40, c_0_39])).
cnf(c_0_45,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X3,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 22
# Proof object formula steps           : 25
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 1
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 8
# Processed clauses                    : 127
# ...of these trivial                  : 3
# ...subsumed                          : 85
# ...remaining for further processing  : 39
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 1300
# ...of the previous two non-trivial   : 996
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1300
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 17
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 885
# ...number of literals in the above   : 885
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 157
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 531
# BW rewrite match successes           : 496
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7886

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43012 pages
