# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))))), file('out/t13_yellow_1__11_premises__rerun.E_input', t3_yellow_1)).
fof(t1_yellow_1, axiom, (![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1))), file('out/t13_yellow_1__11_premises__rerun.E_input', t1_yellow_1)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t13_yellow_1__11_premises__rerun.E_input', redefinition_r3_orders_2)).
fof(dt_k2_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1)))), file('out/t13_yellow_1__11_premises__rerun.E_input', dt_k2_yellow_1)).
fof(fc5_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1)))))), file('out/t13_yellow_1__11_premises__rerun.E_input', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1))))), file('out/t13_yellow_1__11_premises__rerun.E_input', fc6_yellow_1)).
fof(t30_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))))))), file('out/t13_yellow_1__11_premises__rerun.E_input', t30_yellow_0)).
fof(t13_yellow_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k1_xboole_0,X1)=>k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0))), file('out/t13_yellow_1__11_premises__rerun.E_input', t13_yellow_1)).
fof(d11_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0))), file('out/t13_yellow_1__11_premises__rerun.E_input', d11_yellow_0)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t13_yellow_1__11_premises__rerun.E_input', t2_xboole_1)).
fof(t1_subset, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2))), file('out/t13_yellow_1__11_premises__rerun.E_input', t1_subset)).
fof(t6_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2))))), file('out/t13_yellow_1__11_premises__rerun.E_input', t6_yellow_0)).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:(((((~r3_orders_2(k2_yellow_1(X4),X5,X6)|r1_tarski(X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4))&((((~r1_tarski(X5,X6)|r3_orders_2(k2_yellow_1(X4),X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])])])).
fof(c_0_13, plain, (![X2]:![X2]:(u1_struct_0(k2_yellow_1(X2))=X2&u1_orders_2(k2_yellow_1(X2))=k1_yellow_1(X2))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t1_yellow_1])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_15,plain,(v1_xboole_0(X1)|r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, (![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&l1_orders_2(k2_yellow_1(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k2_yellow_1])])])).
fof(c_0_18, plain, (![X2]:![X2]:![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&(v3_orders_2(k2_yellow_1(X2))&(v4_orders_2(k2_yellow_1(X2))&v5_orders_2(k2_yellow_1(X2)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc5_yellow_1])])])).
fof(c_0_19, plain, (![X2]:((~v2_struct_0(k2_yellow_1(X2))|v1_xboole_0(X2))&(v1_orders_2(k2_yellow_1(X2))|v1_xboole_0(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r2_lattice3(X5,X7,X6)|(X6!=k1_yellow_0(X5,X7)|~r1_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X7,X8)|r1_orders_2(X5,X6,X8)))|(X6!=k1_yellow_0(X5,X7)|~r1_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k1_yellow_0(X5,X7)|(m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_yellow_0(X5,X7)|(m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k1_yellow_0(X5,X7)|(r2_lattice3(X5,X7,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_yellow_0(X5,X7)|(r2_lattice3(X5,X7,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&((((X6=k1_yellow_0(X5,X7)|(~r1_orders_2(X5,X6,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_yellow_0(X5,X7)|(~r1_orders_2(X5,X6,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])])).
fof(c_0_21, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k1_xboole_0,X1)=>k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t13_yellow_1])).
cnf(c_0_22,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,plain,(r3_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_16])).
cnf(c_0_24,plain,(l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|X2=k1_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_29, negated_conjecture, ((~v1_xboole_0(esk2_0)&(r2_hidden(k1_xboole_0,esk2_0)&k3_yellow_0(k2_yellow_1(esk2_0))!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_21])])])])).
fof(c_0_30, plain, (![X2]:(~l1_orders_2(X2)|k3_yellow_0(X2)=k1_yellow_0(X2,k1_xboole_0))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_31,plain,(X2=k1_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk1_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32,plain,(r1_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_16]), c_0_16])]), c_0_26])).
cnf(c_0_33,plain,(X1=k1_yellow_0(k2_yellow_1(X2),X3)|m1_subset_1(esk1_3(k2_yellow_1(X2),X1,X3),X2)|~r2_lattice3(k2_yellow_1(X2),X3,X1)|~m1_subset_1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_16]), c_0_28]), c_0_24])])).
fof(c_0_34, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_35, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|m1_subset_1(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_36,negated_conjecture,(k3_yellow_0(k2_yellow_1(esk2_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,plain,(k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38,plain,(X1=k1_yellow_0(k2_yellow_1(X2),X3)|v1_xboole_0(X2)|~r2_lattice3(k2_yellow_1(X2),X3,X1)|~r1_tarski(X1,esk1_3(k2_yellow_1(X2),X1,X3))|~m1_subset_1(X1,X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_28]), c_0_24]), c_0_16])]), c_0_33])).
cnf(c_0_39,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41,negated_conjecture,(r2_hidden(k1_xboole_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42,negated_conjecture,(k1_yellow_0(k2_yellow_1(esk2_0),k1_xboole_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_24])])).
cnf(c_0_43,plain,(k1_yellow_0(k2_yellow_1(X1),X2)=k1_xboole_0|v1_xboole_0(X1)|~r2_lattice3(k2_yellow_1(X1),X2,k1_xboole_0)|~m1_subset_1(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(k1_xboole_0,esk2_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_46, plain, (![X3]:![X4]:(((r2_lattice3(X3,k1_xboole_0,X4)|~m1_subset_1(X4,u1_struct_0(X3)))|~l1_orders_2(X3))&((r1_lattice3(X3,k1_xboole_0,X4)|~m1_subset_1(X4,u1_struct_0(X3)))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])])])).
cnf(c_0_47,negated_conjecture,(~r2_lattice3(k2_yellow_1(esk2_0),k1_xboole_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_48,plain,(r2_lattice3(X1,k1_xboole_0,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_49,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_24]), c_0_16]), c_0_44])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 25
# Proof object formula steps           : 25
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 70
# ...of these trivial                  : 2
# ...subsumed                          : 0
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 23
# ...of the previous two non-trivial   : 17
# Contextual simplify-reflections      : 4
# Paramodulations                      : 21
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 39
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 26
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 256
# Rec. Clause-clause subsumption calls : 8
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 12
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2995

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 43328 pages
