# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t185_relat_1, axiom, (![X1]:![X2]:![X3]:((v1_relat_1(X3)&(v4_relat_1(X3,X1)&v5_relat_1(X3,X2)))=>r1_tarski(X3,k2_zfmisc_1(X1,X2)))), file('out/t13_relset_1__2_premises__rerun.E_input', t185_relat_1)).
fof(fc28_relat_1, axiom, (![X1]:(v1_relat_1(k4_relat_1(X1))&(v4_relat_1(k4_relat_1(X1),X1)&v5_relat_1(k4_relat_1(X1),X1)))), file('out/t13_relset_1__2_premises__rerun.E_input', fc28_relat_1)).
fof(t13_relset_1, conjecture, (![X1]:r1_tarski(k4_relat_1(X1),k2_zfmisc_1(X1,X1))), file('out/t13_relset_1__2_premises__rerun.E_input', t13_relset_1)).
fof(c_0_3, plain, (![X4]:![X5]:![X6]:((~v1_relat_1(X6)|(~v4_relat_1(X6,X4)|~v5_relat_1(X6,X5)))|r1_tarski(X6,k2_zfmisc_1(X4,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t185_relat_1])])).
fof(c_0_4, plain, (![X2]:![X2]:![X2]:(v1_relat_1(k4_relat_1(X2))&(v4_relat_1(k4_relat_1(X2),X2)&v5_relat_1(k4_relat_1(X2),X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc28_relat_1])])])).
fof(c_0_5, negated_conjecture, (~(![X1]:r1_tarski(k4_relat_1(X1),k2_zfmisc_1(X1,X1)))), inference(assume_negation,[status(cth)],[t13_relset_1])).
cnf(c_0_6,plain,(r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v5_relat_1(X1,X3)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_7,plain,(v5_relat_1(k4_relat_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(v1_relat_1(k4_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_9, negated_conjecture, (~r1_tarski(k4_relat_1(esk1_0),k2_zfmisc_1(esk1_0,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_10,plain,(r1_tarski(k4_relat_1(X1),k2_zfmisc_1(X2,X1))|~v4_relat_1(k4_relat_1(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_11,plain,(v4_relat_1(k4_relat_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(~r1_tarski(k4_relat_1(esk1_0),k2_zfmisc_1(esk1_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(r1_tarski(k4_relat_1(X1),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 15
# Proof object clause steps            : 8
# Proof object formula steps           : 7
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 12
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 12
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 2
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 6
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 2
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 314

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 43052 pages
