# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l26_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k10_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))))))), file('out/t13_lattice3__3_premises__rerun.E_input', l26_lattice3)).
fof(cc1_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t13_lattice3__3_premises__rerun.E_input', cc1_lattice3)).
fof(d10_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))))))), file('out/t13_lattice3__3_premises__rerun.E_input', d10_lattice3)).
fof(t13_lattice3, conjecture, (![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2))))), file('out/t13_lattice3__3_premises__rerun.E_input', t13_lattice3)).
fof(c_0_4, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X7,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X8,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X7,X10)|~r1_orders_2(X6,X8,X10))|r1_orders_2(X6,X9,X10)))|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,X7,esk1_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,X8,esk1_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,X9,esk1_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l26_lattice3])])])])])])])])).
fof(c_0_5, plain, (![X2]:(~l1_orders_2(X2)|(~v1_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_6,plain,(v2_struct_0(X1)|X4=k10_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X4,esk1_4(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(~v2_struct_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, (![X6]:![X7]:![X8]:![X10]:![X13]:((((((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v1_lattice3(X6))|~l1_orders_2(X6))&(((((r1_orders_2(X6,X7,esk2_3(X6,X7,X8))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v1_lattice3(X6))|~l1_orders_2(X6))&(((((r1_orders_2(X6,X8,esk2_3(X6,X7,X8))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v1_lattice3(X6))|~l1_orders_2(X6))&(((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X7,X10)|~r1_orders_2(X6,X8,X10))|r1_orders_2(X6,esk2_3(X6,X7,X8),X10)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v1_lattice3(X6))|~l1_orders_2(X6)))))&(((m1_subset_1(esk3_1(X6),u1_struct_0(X6))|v1_lattice3(X6))|~l1_orders_2(X6))&(((m1_subset_1(esk4_1(X6),u1_struct_0(X6))|v1_lattice3(X6))|~l1_orders_2(X6))&((((((m1_subset_1(esk5_2(X6,X13),u1_struct_0(X6))|~r1_orders_2(X6,esk4_1(X6),X13))|~r1_orders_2(X6,esk3_1(X6),X13))|~m1_subset_1(X13,u1_struct_0(X6)))|v1_lattice3(X6))|~l1_orders_2(X6))&(((((((r1_orders_2(X6,esk3_1(X6),esk5_2(X6,X13))|~r1_orders_2(X6,esk4_1(X6),X13))|~r1_orders_2(X6,esk3_1(X6),X13))|~m1_subset_1(X13,u1_struct_0(X6)))|v1_lattice3(X6))|~l1_orders_2(X6))&(((((r1_orders_2(X6,esk4_1(X6),esk5_2(X6,X13))|~r1_orders_2(X6,esk4_1(X6),X13))|~r1_orders_2(X6,esk3_1(X6),X13))|~m1_subset_1(X13,u1_struct_0(X6)))|v1_lattice3(X6))|~l1_orders_2(X6)))&(((((~r1_orders_2(X6,X13,esk5_2(X6,X13))|~r1_orders_2(X6,esk4_1(X6),X13))|~r1_orders_2(X6,esk3_1(X6),X13))|~m1_subset_1(X13,u1_struct_0(X6)))|v1_lattice3(X6))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_lattice3])])])])])])])).
cnf(c_0_9,plain,(v2_struct_0(X1)|X4=k10_lattice3(X1,X2,X3)|m1_subset_1(esk1_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,plain,(X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk1_4(X2,X3,X4,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11,plain,(r1_orders_2(X1,esk2_3(X1,X2,X3),X4)|~l1_orders_2(X1)|~v1_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(X1=k10_lattice3(X2,X3,X4)|m1_subset_1(esk1_4(X2,X3,X4,X1),u1_struct_0(X2))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_13,plain,(m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(v2_struct_0(X1)|X4=k10_lattice3(X1,X2,X3)|r1_orders_2(X1,X2,esk1_4(X1,X2,X3,X4))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15,plain,(v2_struct_0(X1)|X4=k10_lattice3(X1,X2,X3)|r1_orders_2(X1,X3,esk1_4(X1,X2,X3,X4))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,plain,(esk2_3(X1,X2,X3)=k10_lattice3(X1,X4,X5)|~r1_orders_2(X1,X3,esk1_4(X1,X4,X5,esk2_3(X1,X2,X3)))|~r1_orders_2(X1,X2,esk1_4(X1,X4,X5,esk2_3(X1,X2,X3)))|~r1_orders_2(X1,X5,esk2_3(X1,X2,X3))|~r1_orders_2(X1,X4,esk2_3(X1,X2,X3))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])).
cnf(c_0_17,plain,(X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X3,esk1_4(X2,X3,X4,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_14, c_0_7])).
cnf(c_0_18,plain,(r1_orders_2(X1,X3,esk2_3(X1,X2,X3))|~l1_orders_2(X1)|~v1_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19,plain,(X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk1_4(X2,X3,X4,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_15, c_0_7])).
fof(c_0_20, negated_conjecture, (~(![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)))))), inference(assume_negation,[status(cth)],[t13_lattice3])).
cnf(c_0_21,plain,(esk2_3(X1,X2,X3)=k10_lattice3(X1,X3,X4)|~r1_orders_2(X1,X2,esk1_4(X1,X3,X4,esk2_3(X1,X2,X3)))|~r1_orders_2(X1,X4,esk2_3(X1,X2,X3))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_13]), c_0_18])).
cnf(c_0_22,plain,(r1_orders_2(X1,X2,esk2_3(X1,X2,X3))|~l1_orders_2(X1)|~v1_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_23,plain,(esk2_3(X1,X2,X3)=k10_lattice3(X1,X4,X3)|~r1_orders_2(X1,X2,esk1_4(X1,X4,X3,esk2_3(X1,X2,X3)))|~r1_orders_2(X1,X4,esk2_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_19]), c_0_13]), c_0_18])).
fof(c_0_24, negated_conjecture, (((v5_orders_2(esk6_0)&(v1_lattice3(esk6_0)&l1_orders_2(esk6_0)))&(m1_subset_1(esk7_0,u1_struct_0(esk6_0))&(m1_subset_1(esk8_0,u1_struct_0(esk6_0))&k10_lattice3(esk6_0,esk7_0,esk8_0)!=k10_lattice3(esk6_0,esk8_0,esk7_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])])).
cnf(c_0_25,plain,(esk2_3(X1,X2,X3)=k10_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_19]), c_0_13]), c_0_18]), c_0_22])).
cnf(c_0_26,plain,(esk2_3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_17]), c_0_13]), c_0_18]), c_0_22])).
cnf(c_0_27,negated_conjecture,(k10_lattice3(esk6_0,esk7_0,esk8_0)!=k10_lattice3(esk6_0,esk8_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28,plain,(k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29,negated_conjecture,(m1_subset_1(esk7_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk8_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,negated_conjecture,(l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,negated_conjecture,(v1_lattice3(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,negated_conjecture,(v5_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 26
# Proof object formula steps           : 9
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 122
# ...of these trivial                  : 0
# ...subsumed                          : 28
# ...remaining for further processing  : 94
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 0
# Generated clauses                    : 436
# ...of the previous two non-trivial   : 309
# Contextual simplify-reflections      : 84
# Paramodulations                      : 433
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 191
# ...number of literals in the above   : 2151
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 3675
# Rec. Clause-clause subsumption calls : 194
# Non-unit clause-clause subsumptions  : 122
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20512

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 45284 pages
