# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', t44_enumset1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', t71_enumset1)).
fof(d1_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3))))))), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', d1_enumset1)).
fof(t69_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', t69_enumset1)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', d1_tarski)).
fof(t13_zfmisc_1, conjecture, (![X1]:![X2]:(k2_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2)), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', t13_zfmisc_1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t13_zfmisc_1__10_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X5),k1_enumset1(X6,X7,X8))), inference(variable_rename,[status(thm)],[t44_enumset1])).
fof(c_0_10, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_11, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_13, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X10]:![X6]:![X7]:![X8]:![X9]:((((~r2_hidden(X10,X9)|(X10=X6|(X10=X7|X10=X8)))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X6|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X7|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&((X10!=X8|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8)))))&((((esk1_4(X6,X7,X8,X9)!=X6|~r2_hidden(esk1_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&(((esk1_4(X6,X7,X8,X9)!=X7|~r2_hidden(esk1_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&((esk1_4(X6,X7,X8,X9)!=X8|~r2_hidden(esk1_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))))&((r2_hidden(esk1_4(X6,X7,X8,X9),X9)|(esk1_4(X6,X7,X8,X9)=X6|(esk1_4(X6,X7,X8,X9)=X7|esk1_4(X6,X7,X8,X9)=X8)))|X9=k1_enumset1(X6,X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])])).
cnf(c_0_14,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X6,X8,X5,X7)), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_19,plain,(X5=X4|X5=X3|X5=X2|X1!=k1_enumset1(X2,X3,X4)|~r2_hidden(X5,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k2_enumset1(X2,X2,X3,X4),k3_xboole_0(k2_enumset1(X2,X2,X3,X4),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17]), c_0_17])).
cnf(c_0_21,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22,plain,(X4=X5|X3=X5|X2=X5|X1!=k2_enumset1(X2,X2,X3,X4)|~r2_hidden(X5,X1)), inference(rw,[status(thm)],[c_0_19, c_0_17])).
cnf(c_0_23,plain,(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k2_enumset1(X2,X3,X4,X3),k3_xboole_0(k2_enumset1(X2,X3,X4,X3),k1_tarski(X1))))=k2_enumset1(X1,X3,X2,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24,plain,(X1=X2|X3=X2|X4=X2|~r2_hidden(X2,k2_enumset1(X4,X4,X3,X1))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_25,plain,(k2_enumset1(X1,X2,X2,X3)=k2_enumset1(X1,X2,X3,X2)), inference(spm,[status(thm)],[c_0_20, c_0_23])).
cnf(c_0_26,plain,(X1=X2|X3=X2|~r2_hidden(X2,k2_enumset1(X3,X3,X3,X1))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_27, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk2_2(X4,X5),X5)|esk2_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk2_2(X4,X5),X5)|esk2_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_28,plain,(X1=X2|~r2_hidden(X2,k2_enumset1(X1,X1,X1,X1))), inference(spm,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_29,plain,(X1=k1_tarski(X2)|esk2_2(X2,X1)=X2|r2_hidden(esk2_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X2), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31,plain,(esk2_2(X1,k2_enumset1(X2,X2,X2,X2))=X1|esk2_2(X1,k2_enumset1(X2,X2,X2,X2))=X2|k2_enumset1(X2,X2,X2,X2)=k1_tarski(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32,plain,(r2_hidden(X5,X1)|X2!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_30, c_0_17])).
fof(c_0_33, negated_conjecture, (~(![X1]:![X2]:(k2_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2))), inference(assume_negation,[status(cth)],[t13_zfmisc_1])).
cnf(c_0_34,plain,(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k2_enumset1(X2,X3,X2,X4),k3_xboole_0(k2_enumset1(X2,X3,X2,X4),k1_tarski(X1))))=k2_enumset1(X1,X2,X4,X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_35,plain,(X1=k1_tarski(X2)|esk2_2(X2,X1)!=X2|~r2_hidden(esk2_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36,plain,(esk2_2(X1,k2_enumset1(X1,X1,X1,X1))=X1|k2_enumset1(X1,X1,X1,X1)=k1_tarski(X1)), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_31])])).
cnf(c_0_37,plain,(r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])])).
fof(c_0_38, negated_conjecture, ((k2_xboole_0(k1_tarski(esk3_0),k1_tarski(esk4_0))=k1_tarski(esk3_0)&esk3_0!=esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).
cnf(c_0_39,plain,(k2_enumset1(X1,X2,X3,X3)=k2_enumset1(X1,X3,X2,X2)), inference(spm,[status(thm)],[c_0_23, c_0_34])).
cnf(c_0_40,plain,(k2_enumset1(X1,X1,X1,X1)=k1_tarski(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
cnf(c_0_41,negated_conjecture,(k2_xboole_0(k1_tarski(esk3_0),k1_tarski(esk4_0))=k1_tarski(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_42, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_43,plain,(k2_enumset1(X1,X2,X3,X3)=k2_enumset1(X3,X2,X1,X2)), inference(spm,[status(thm)],[c_0_21, c_0_39])).
cnf(c_0_44,plain,(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1))))=k2_enumset1(X1,X2,X2,X2)), inference(spm,[status(thm)],[c_0_20, c_0_40])).
cnf(c_0_45,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X4), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_46,negated_conjecture,(k5_xboole_0(k1_tarski(esk3_0),k5_xboole_0(k1_tarski(esk4_0),k3_xboole_0(k1_tarski(esk4_0),k1_tarski(esk3_0))))=k1_tarski(esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_15]), c_0_16])).
cnf(c_0_47,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48,plain,(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1))))=k2_enumset1(X2,X2,X1,X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49,plain,(k2_enumset1(X1,X2,X2,X3)=k2_enumset1(X2,X2,X1,X3)), inference(spm,[status(thm)],[c_0_21, c_0_25])).
cnf(c_0_50,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(spm,[status(thm)],[c_0_21, c_0_21])).
cnf(c_0_51,plain,(r2_hidden(X5,X1)|X4!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_45, c_0_17])).
cnf(c_0_52,negated_conjecture,(k5_xboole_0(k1_tarski(esk3_0),k5_xboole_0(k1_tarski(esk4_0),k3_xboole_0(k1_tarski(esk3_0),k1_tarski(esk4_0))))=k1_tarski(esk3_0)), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53,plain,(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X1),k1_tarski(X2))))=k2_enumset1(X2,X2,X1,X1)), inference(spm,[status(thm)],[c_0_48, c_0_47])).
cnf(c_0_54,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X1,X3,X2,X1)), inference(spm,[status(thm)],[c_0_21, c_0_49])).
cnf(c_0_55,plain,(k2_enumset1(X1,X2,X3,X3)=k2_enumset1(X2,X2,X3,X1)), inference(spm,[status(thm)],[c_0_50, c_0_39])).
cnf(c_0_56,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X3,X1,X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_49])).
cnf(c_0_57,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_58,plain,(r2_hidden(X1,k2_enumset1(X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_51])])).
cnf(c_0_59,negated_conjecture,(k2_enumset1(esk3_0,esk3_0,esk3_0,esk4_0)=k1_tarski(esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_49]), c_0_55]), c_0_56])).
cnf(c_0_60,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_61,negated_conjecture,(r2_hidden(esk4_0,k1_tarski(esk3_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62,negated_conjecture,(esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_63,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 45
# Proof object formula steps           : 19
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 17
# Processed clauses                    : 2059
# ...of these trivial                  : 38
# ...subsumed                          : 1789
# ...remaining for further processing  : 232
# Other redundant clauses eliminated   : 379
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 6
# Generated clauses                    : 49040
# ...of the previous two non-trivial   : 44879
# Contextual simplify-reflections      : 0
# Paramodulations                      : 48434
# Factorizations                       : 231
# Equation resolutions                 : 379
# Current number of processed clauses  : 193
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 59
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 99
# Current number of unprocessed clauses: 38836
# ...number of literals in the above   : 233868
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 4798
# Rec. Clause-clause subsumption calls : 1709
# Non-unit clause-clause subsumptions  : 560
# Unit Clause-clause subsumption calls : 4260
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8181
# BW rewrite match successes           : 1502
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 746154

# -------------------------------------------------
# User time                : 1.084 s
# System time              : 0.016 s
# Total time               : 1.100 s
# Maximum resident set size: 43136 pages
