# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t141_zfmisc_1, conjecture, (![X1]:![X2]:(~(r2_hidden(X1,X2))=>k4_xboole_0(k2_xboole_0(X2,k1_tarski(X1)),k1_tarski(X1))=X2)), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t141_zfmisc_1)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t95_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t71_enumset1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t91_xboole_1)).
fof(t40_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t40_xboole_1)).
fof(t75_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1)), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t75_enumset1)).
fof(t92_xboole_1, axiom, (![X1]:k5_xboole_0(X1,X1)=k1_xboole_0), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t92_xboole_1)).
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', d5_xboole_0)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', t5_boole)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t141_zfmisc_1__17_premises__rerun__rerun.E_input', d1_tarski)).
fof(c_0_13, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_14, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_15, negated_conjecture, (~(![X1]:![X2]:(~(r2_hidden(X1,X2))=>k4_xboole_0(k2_xboole_0(X2,k1_tarski(X1)),k1_tarski(X1))=X2))), inference(assume_negation,[status(cth)],[t141_zfmisc_1])).
fof(c_0_16, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_17,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_20, negated_conjecture, ((~r2_hidden(esk1_0,esk2_0)&k4_xboole_0(k2_xboole_0(esk2_0,k1_tarski(esk1_0)),k1_tarski(esk1_0))!=esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
fof(c_0_21, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_22,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,negated_conjecture,(k4_xboole_0(k2_xboole_0(esk2_0,k1_tarski(esk1_0)),k1_tarski(esk1_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
fof(c_0_28, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
fof(c_0_29, plain, (![X3]:![X4]:k4_xboole_0(k2_xboole_0(X3,X4),X4)=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_30,negated_conjecture,(k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k5_xboole_0(k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k1_tarski(esk1_0)),k3_tarski(k2_enumset1(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k1_tarski(esk1_0)))))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_23]), c_0_26]), c_0_24]), c_0_24]), c_0_27])).
cnf(c_0_31,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_32, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X8,X6,X7,X5)), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_33,plain,(k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34,negated_conjecture,(k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k1_tarski(esk1_0))))))!=esk2_0), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_36,plain,(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k5_xboole_0(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),X2),k3_tarski(k2_enumset1(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,X2)),X2))))=k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_23]), c_0_26]), c_0_26]), c_0_24]), c_0_24]), c_0_27]), c_0_27])).
fof(c_0_37, plain, (![X2]:k5_xboole_0(X2,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_38, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk3_3(X5,X6,X7),X7)|(~r2_hidden(esk3_3(X5,X6,X7),X5)|r2_hidden(esk3_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk3_3(X5,X6,X7),X5)|r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk3_3(X5,X6,X7),X6)|r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
cnf(c_0_39,negated_conjecture,(k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(k1_tarski(esk1_0),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))))))))!=esk2_0), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40,plain,(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k5_xboole_0(X2,k3_tarski(k2_enumset1(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,X2)))))))=k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_35]), c_0_31]), c_0_31])).
cnf(c_0_41,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_42, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_43,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X1)|r2_hidden(esk3_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X3)|~r2_hidden(esk3_3(X2,X3,X1),X2)|~r2_hidden(esk3_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45,negated_conjecture,(k5_xboole_0(esk2_0,k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))))))!=esk2_0), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=k5_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_31, c_0_41])).
cnf(c_0_47,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48,plain,(X1=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k3_tarski(k2_enumset1(X2,X2,X2,X3))))|r2_hidden(esk3_3(X2,X3,X1),X2)|r2_hidden(esk3_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_26]), c_0_27])).
cnf(c_0_49,plain,(X1=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k3_tarski(k2_enumset1(X2,X2,X2,X3))))|r2_hidden(esk3_3(X2,X3,X1),X3)|~r2_hidden(esk3_3(X2,X3,X1),X2)|~r2_hidden(esk3_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_26]), c_0_27])).
cnf(c_0_50,negated_conjecture,(k5_xboole_0(k1_xboole_0,k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0)))))!=esk2_0), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_41]), c_0_47])).
cnf(c_0_52,plain,(X1=k5_xboole_0(X2,k5_xboole_0(X2,k5_xboole_0(X3,k3_tarski(k2_enumset1(X2,X2,X2,X3)))))|r2_hidden(esk3_3(X2,X3,X1),X1)|r2_hidden(esk3_3(X2,X3,X1),X2)), inference(rw,[status(thm)],[c_0_48, c_0_31])).
fof(c_0_53, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_54,plain,(X1=k5_xboole_0(X2,k5_xboole_0(X2,k5_xboole_0(X3,k3_tarski(k2_enumset1(X2,X2,X2,X3)))))|r2_hidden(esk3_3(X2,X3,X1),X3)|~r2_hidden(esk3_3(X2,X3,X1),X2)|~r2_hidden(esk3_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[c_0_49, c_0_31])).
cnf(c_0_55,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))))!=esk2_0), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56,plain,(X1=k5_xboole_0(X2,k3_tarski(k2_enumset1(X3,X3,X3,X2)))|r2_hidden(esk3_3(X3,X2,X1),X3)|r2_hidden(esk3_3(X3,X2,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_46]), c_0_51])).
cnf(c_0_57,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_58,plain,(X1=k5_xboole_0(X2,k3_tarski(k2_enumset1(X3,X3,X3,X2)))|r2_hidden(esk3_3(X3,X2,X1),X2)|~r2_hidden(esk3_3(X3,X2,X1),X3)|~r2_hidden(esk3_3(X3,X2,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_46]), c_0_51])).
cnf(c_0_59,negated_conjecture,(r2_hidden(esk3_3(esk2_0,k1_tarski(esk1_0),esk2_0),esk2_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56])])).
cnf(c_0_60,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_61,negated_conjecture,(r2_hidden(esk3_3(esk2_0,k1_tarski(esk1_0),esk2_0),k1_tarski(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_59])]), c_0_55])).
cnf(c_0_62,negated_conjecture,(esk3_3(esk2_0,k1_tarski(esk1_0),esk2_0)=esk1_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_63,negated_conjecture,(~r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_64,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_62]), c_0_63]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 38
# Proof object formula steps           : 27
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 17
# Processed clauses                    : 89
# ...of these trivial                  : 0
# ...subsumed                          : 29
# ...remaining for further processing  : 60
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 12
# Generated clauses                    : 270
# ...of the previous two non-trivial   : 211
# Contextual simplify-reflections      : 0
# Paramodulations                      : 261
# Factorizations                       : 2
# Equation resolutions                 : 8
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 138
# ...number of literals in the above   : 380
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 103
# Rec. Clause-clause subsumption calls : 85
# Non-unit clause-clause subsumptions  : 22
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 71
# BW rewrite match successes           : 52
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4733

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45288 pages
