# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t146_funct_2, conjecture, (![X1]:![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>(r1_partfun1(X2,X3)<=>![X4]:(r2_hidden(X4,k1_relset_1(X1,X1,X2))=>k1_funct_1(X2,X4)=k1_funct_1(X3,X4)))))), file('out/t146_funct_2__1_premises__rerun.E_input', t146_funct_2)).
fof(t145_funct_2, axiom, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((X2=k1_xboole_0=>X1=k1_xboole_0)=>(r1_partfun1(X3,X4)<=>![X5]:(r2_hidden(X5,k1_relset_1(X1,X2,X3))=>k1_funct_1(X3,X5)=k1_funct_1(X4,X5))))))), file('out/t146_funct_2__1_premises__rerun.E_input', t145_funct_2)).
fof(c_0_2, negated_conjecture, (~(![X1]:![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>(r1_partfun1(X2,X3)<=>![X4]:(r2_hidden(X4,k1_relset_1(X1,X1,X2))=>k1_funct_1(X2,X4)=k1_funct_1(X3,X4))))))), inference(assume_negation,[status(cth)],[t146_funct_2])).
fof(c_0_3, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((((((~r1_partfun1(X8,X9)|(~r2_hidden(X10,k1_relset_1(X6,X7,X8))|k1_funct_1(X8,X10)=k1_funct_1(X9,X10)))|X7=k1_xboole_0)|(~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|(~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))&(((((r2_hidden(esk5_4(X6,X7,X8,X9),k1_relset_1(X6,X7,X8))|r1_partfun1(X8,X9))|X7=k1_xboole_0)|(~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|(~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))&((((k1_funct_1(X8,esk5_4(X6,X7,X8,X9))!=k1_funct_1(X9,esk5_4(X6,X7,X8,X9))|r1_partfun1(X8,X9))|X7=k1_xboole_0)|(~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|(~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))))&(((((~r1_partfun1(X8,X9)|(~r2_hidden(X10,k1_relset_1(X6,X7,X8))|k1_funct_1(X8,X10)=k1_funct_1(X9,X10)))|X6!=k1_xboole_0)|(~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|(~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))&(((((r2_hidden(esk5_4(X6,X7,X8,X9),k1_relset_1(X6,X7,X8))|r1_partfun1(X8,X9))|X6!=k1_xboole_0)|(~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|(~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))&((((k1_funct_1(X8,esk5_4(X6,X7,X8,X9))!=k1_funct_1(X9,esk5_4(X6,X7,X8,X9))|r1_partfun1(X8,X9))|X6!=k1_xboole_0)|(~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|(~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t145_funct_2])])])])])])])).
fof(c_0_4, negated_conjecture, (![X9]:((v1_funct_1(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0))))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk1_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&(((r2_hidden(esk4_0,k1_relset_1(esk1_0,esk1_0,esk2_0))|~r1_partfun1(esk2_0,esk3_0))&(k1_funct_1(esk2_0,esk4_0)!=k1_funct_1(esk3_0,esk4_0)|~r1_partfun1(esk2_0,esk3_0)))&(r1_partfun1(esk2_0,esk3_0)|(~r2_hidden(X9,k1_relset_1(esk1_0,esk1_0,esk2_0))|k1_funct_1(esk2_0,X9)=k1_funct_1(esk3_0,X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])])).
cnf(c_0_5,plain,(X3=k1_xboole_0|k1_funct_1(X1,X5)=k1_funct_1(X4,X5)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4)|~r2_hidden(X5,k1_relset_1(X2,X3,X1))|~r1_partfun1(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(r2_hidden(esk4_0,k1_relset_1(esk1_0,esk1_0,esk2_0))|~r1_partfun1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,negated_conjecture,(v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(k1_funct_1(esk2_0,esk4_0)=k1_funct_1(X1,esk4_0)|k1_xboole_0=esk1_0|~r1_partfun1(esk2_0,esk3_0)|~r1_partfun1(esk2_0,X1)|~v1_funct_2(X1,esk1_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7]), c_0_8])])).
cnf(c_0_10,negated_conjecture,(v1_funct_2(esk3_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13,negated_conjecture,(~r1_partfun1(esk2_0,esk3_0)|k1_funct_1(esk2_0,esk4_0)!=k1_funct_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14,negated_conjecture,(k1_funct_1(esk2_0,X1)=k1_funct_1(esk3_0,X1)|r1_partfun1(esk2_0,esk3_0)|~r2_hidden(X1,k1_relset_1(esk1_0,esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15,plain,(X3=k1_xboole_0|r1_partfun1(X1,X4)|r2_hidden(esk5_4(X2,X3,X1,X4),k1_relset_1(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_16,negated_conjecture,(k1_xboole_0=esk1_0|~r1_partfun1(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])]), c_0_13])).
cnf(c_0_17,plain,(k1_funct_1(X1,X5)=k1_funct_1(X4,X5)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4)|X2!=k1_xboole_0|~r2_hidden(X5,k1_relset_1(X2,X3,X1))|~r1_partfun1(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_18,plain,(X3=k1_xboole_0|r1_partfun1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4)|k1_funct_1(X1,esk5_4(X2,X3,X1,X4))!=k1_funct_1(X4,esk5_4(X2,X3,X1,X4))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_19,negated_conjecture,(k1_funct_1(esk3_0,esk5_4(esk1_0,esk1_0,esk2_0,X1))=k1_funct_1(esk2_0,esk5_4(esk1_0,esk1_0,esk2_0,X1))|k1_xboole_0=esk1_0|r1_partfun1(esk2_0,X1)|~v1_funct_2(X1,esk1_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_7]), c_0_8])]), c_0_16])).
cnf(c_0_20,plain,(k1_funct_1(X1,X2)=k1_funct_1(X3,X2)|~r2_hidden(X2,k1_relset_1(k1_xboole_0,X4,X1))|~r1_partfun1(X1,X3)|~v1_funct_2(X3,k1_xboole_0,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X4)))|~v1_funct_1(X3)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21,negated_conjecture,(k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_10]), c_0_11]), c_0_7]), c_0_12]), c_0_8])]), c_0_16])).
cnf(c_0_22,plain,(r1_partfun1(X1,X4)|r2_hidden(esk5_4(X2,X3,X1,X4),k1_relset_1(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4)|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_23,plain,(k1_funct_1(X1,X2)=k1_funct_1(X3,X2)|~r2_hidden(X2,k1_relset_1(esk1_0,X4,X1))|~r1_partfun1(X1,X3)|~v1_funct_2(X3,esk1_0,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X4)))|~v1_funct_1(X3)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_24,plain,(r1_partfun1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4)|X2!=k1_xboole_0|k1_funct_1(X1,esk5_4(X2,X3,X1,X4))!=k1_funct_1(X4,esk5_4(X2,X3,X1,X4))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_25,plain,(r2_hidden(esk5_4(k1_xboole_0,X1,X2,X3),k1_relset_1(k1_xboole_0,X1,X2))|r1_partfun1(X2,X3)|~v1_funct_2(X3,k1_xboole_0,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X1)))|~v1_funct_1(X3)|~v1_funct_1(X2)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_26,negated_conjecture,(k1_funct_1(esk2_0,esk4_0)=k1_funct_1(X1,esk4_0)|~r1_partfun1(esk2_0,esk3_0)|~r1_partfun1(esk2_0,X1)|~v1_funct_2(X1,esk1_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_6]), c_0_7]), c_0_8])])).
cnf(c_0_27,plain,(r1_partfun1(X1,X2)|k1_funct_1(X1,esk5_4(k1_xboole_0,X3,X1,X2))!=k1_funct_1(X2,esk5_4(k1_xboole_0,X3,X1,X2))|~v1_funct_2(X2,k1_xboole_0,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X3)))|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_28,plain,(r2_hidden(esk5_4(esk1_0,X1,X2,X3),k1_relset_1(esk1_0,X1,X2))|r1_partfun1(X2,X3)|~v1_funct_2(X3,esk1_0,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|~v1_funct_1(X3)|~v1_funct_1(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_21]), c_0_21]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_29,negated_conjecture,(~r1_partfun1(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_10]), c_0_11]), c_0_12])]), c_0_13])).
cnf(c_0_30,plain,(r1_partfun1(X1,X2)|k1_funct_1(X1,esk5_4(esk1_0,X3,X1,X2))!=k1_funct_1(X2,esk5_4(esk1_0,X3,X1,X2))|~v1_funct_2(X2,esk1_0,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X3)))|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_21]), c_0_21]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_31,negated_conjecture,(k1_funct_1(esk3_0,esk5_4(esk1_0,esk1_0,esk2_0,X1))=k1_funct_1(esk2_0,esk5_4(esk1_0,esk1_0,esk2_0,X1))|r1_partfun1(esk2_0,X1)|~v1_funct_2(X1,esk1_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_28]), c_0_7]), c_0_8])]), c_0_29])).
cnf(c_0_32,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_10]), c_0_11]), c_0_7]), c_0_12]), c_0_8])]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 28
# Proof object formula steps           : 5
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 14
# Processed clauses                    : 47
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 47
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 11
# Generated clauses                    : 22
# ...of the previous two non-trivial   : 28
# Contextual simplify-reflections      : 4
# Paramodulations                      : 16
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 18
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 40
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 120
# Rec. Clause-clause subsumption calls : 7
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2539

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45300 pages
