# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t4_xboole_0)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t95_xboole_1)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t28_xboole_1)).
fof(d7_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0)), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', d7_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', d10_xboole_0)).
fof(t14_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&(r1_tarski(X3,X2)&![X4]:((r1_tarski(X1,X4)&r1_tarski(X3,X4))=>r1_tarski(X2,X4))))=>X2=k2_xboole_0(X1,X3))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t14_xboole_1)).
fof(t39_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t39_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t5_boole)).
fof(d6_xboole_0, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', d6_xboole_0)).
fof(t67_zfmisc_1, axiom, (![X1]:![X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)<=>~(r2_hidden(X1,X2)))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t67_zfmisc_1)).
fof(t29_xboole_1, axiom, (![X1]:![X2]:![X3]:r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X1,X3))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t29_xboole_1)).
fof(t16_zfmisc_1, axiom, (![X1]:![X2]:~((r1_xboole_0(k1_tarski(X1),k1_tarski(X2))&X1=X2))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t16_zfmisc_1)).
fof(t44_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t44_enumset1)).
fof(t89_enumset1, axiom, (![X1]:![X2]:![X3]:k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t89_enumset1)).
fof(t48_zfmisc_1, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k2_xboole_0(k2_tarski(X1,X3),X2)=X2)), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t48_zfmisc_1)).
fof(t77_enumset1, axiom, (![X1]:![X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t77_enumset1)).
fof(t18_zfmisc_1, axiom, (![X1]:![X2]:(k3_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2)), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t18_zfmisc_1)).
fof(t147_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:(~(X1=X2)=>k4_xboole_0(k2_xboole_0(X3,k1_tarski(X1)),k1_tarski(X2))=k2_xboole_0(k4_xboole_0(X3,k1_tarski(X2)),k1_tarski(X1)))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t147_zfmisc_1)).
fof(t87_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_xboole_0(X1,X2)=>k2_xboole_0(k4_xboole_0(X3,X1),X2)=k4_xboole_0(k2_xboole_0(X3,X2),X1))), file('out/t147_zfmisc_1__21_premises__rerun__rerun.E_input', t87_xboole_1)).
fof(c_0_21, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk1_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
fof(c_0_22, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
fof(c_0_23, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_24,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_27, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k3_xboole_0(X3,X4)=k1_xboole_0)&(k3_xboole_0(X3,X4)!=k1_xboole_0|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])])])).
cnf(c_0_28,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,plain,(k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_26, c_0_25])).
fof(c_0_30, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk2_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk2_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_31,plain,(r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32,plain,(~r1_tarski(X1,X2)|~r2_hidden(X3,X1)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33,plain,(r1_tarski(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_34,plain,(r1_xboole_0(X1,X2)|k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_31, c_0_25])).
fof(c_0_35, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_36,plain,(r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37,plain,(r1_xboole_0(k1_xboole_0,X1)|~r1_tarski(k1_xboole_0,X1)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_29])])).
cnf(c_0_38,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_39, plain, (![X5]:![X6]:![X7]:(((((r1_tarski(X5,esk3_3(X5,X6,X7))|~r1_tarski(X7,X6))|~r1_tarski(X5,X6))|X6=k2_xboole_0(X5,X7))&(((r1_tarski(X7,esk3_3(X5,X6,X7))|~r1_tarski(X7,X6))|~r1_tarski(X5,X6))|X6=k2_xboole_0(X5,X7)))&(((~r1_tarski(X6,esk3_3(X5,X6,X7))|~r1_tarski(X7,X6))|~r1_tarski(X5,X6))|X6=k2_xboole_0(X5,X7)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_xboole_1])])])])])])).
cnf(c_0_40,plain,(r1_tarski(k1_xboole_0,X1)|~r1_tarski(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_38])).
fof(c_0_42, plain, (![X3]:![X4]:k2_xboole_0(X3,k4_xboole_0(X4,X3))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_43, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_44,plain,(X1=k2_xboole_0(X2,X3)|~r1_tarski(X2,X1)|~r1_tarski(X3,X1)|~r1_tarski(X1,esk3_3(X2,X1,X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45,plain,(X1=k2_xboole_0(X2,X3)|r1_tarski(X3,esk3_3(X2,X1,X3))|~r1_tarski(X2,X1)|~r1_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_47,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48,plain,(k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_41])])).
fof(c_0_51, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_52,plain,(X1=k2_xboole_0(X2,X3)|r1_tarski(X2,esk3_3(X2,X1,X3))|~r1_tarski(X2,X1)|~r1_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_53, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k2_xboole_0(k4_xboole_0(X3,X4),k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[d6_xboole_0])).
cnf(c_0_54,plain,(k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_46, c_0_25])).
cnf(c_0_55,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_47])])).
cnf(c_0_56,plain,(k2_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X1),k2_xboole_0(X2,X1))))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49]), c_0_25])).
cnf(c_0_57,plain,(k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_50, c_0_47])).
cnf(c_0_58,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59,plain,(k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_52]), c_0_41])])).
fof(c_0_60, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(k1_tarski(X3),X4)!=k1_tarski(X3)|~r2_hidden(X3,X4))&(r2_hidden(X3,X4)|k4_xboole_0(k1_tarski(X3),X4)=k1_tarski(X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t67_zfmisc_1])])])])])).
cnf(c_0_61,plain,(k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_62,plain,(k5_xboole_0(k5_xboole_0(k1_xboole_0,X1),k2_xboole_0(k1_xboole_0,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_63,plain,(k5_xboole_0(X1,k5_xboole_0(X1,k2_xboole_0(X1,k1_xboole_0)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_57])).
cnf(c_0_64,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_59, c_0_47])).
cnf(c_0_65,plain,(~r2_hidden(X1,X2)|k4_xboole_0(k1_tarski(X1),X2)!=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66,plain,(k5_xboole_0(X1,X2)=k2_xboole_0(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))),k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X1),k2_xboole_0(X2,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_49]), c_0_49]), c_0_25]), c_0_25])).
cnf(c_0_67,plain,(k5_xboole_0(k5_xboole_0(k1_xboole_0,X1),X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_62, c_0_57])).
cnf(c_0_68,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[c_0_63, c_0_64])).
fof(c_0_69, plain, (![X4]:![X5]:![X6]:r1_tarski(k3_xboole_0(X4,X5),k2_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[t29_xboole_1])).
cnf(c_0_70,plain,(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_xboole_0(k1_tarski(X1),X2),k2_xboole_0(k1_tarski(X1),X2)))!=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_49]), c_0_25])).
cnf(c_0_71,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_64]), c_0_57]), c_0_67]), c_0_58]), c_0_58]), c_0_57]), c_0_68])).
cnf(c_0_72,plain,(r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_73,plain,(r1_xboole_0(X1,k1_xboole_0)|k5_xboole_0(X1,k2_xboole_0(X1,k1_xboole_0))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_58])).
cnf(c_0_74,plain,(k5_xboole_0(k1_tarski(X1),k1_tarski(X1))!=k1_tarski(X1)|~r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_70, c_0_29])).
cnf(c_0_75,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_67, c_0_71])).
cnf(c_0_76,plain,(r1_tarski(k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)),k2_xboole_0(X1,X3))), inference(rw,[status(thm)],[c_0_72, c_0_25])).
cnf(c_0_77,plain,(k2_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_50, c_0_41])).
cnf(c_0_78,plain,(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)|r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_79,plain,(r1_xboole_0(X1,k1_xboole_0)|k5_xboole_0(X1,X1)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_73, c_0_64])).
cnf(c_0_80,plain,(k1_tarski(X1)!=k1_xboole_0|~r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_81,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_75]), c_0_77]), c_0_71])).
cnf(c_0_82,plain,(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_xboole_0(k1_tarski(X1),X2),k2_xboole_0(k1_tarski(X1),X2)))=k1_tarski(X1)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_49]), c_0_25])).
cnf(c_0_83,plain,(~r2_hidden(X1,k5_xboole_0(k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(k5_xboole_0(X3,X2),k2_xboole_0(X3,X2)))),k2_xboole_0(X2,X3)))|~r1_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(k5_xboole_0(X3,X2),k2_xboole_0(X3,X2))))), inference(spm,[status(thm)],[c_0_28, c_0_56])).
cnf(c_0_84,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_75])])).
cnf(c_0_85,plain,(k1_tarski(X1)!=k1_xboole_0|~r2_hidden(X1,k2_xboole_0(k1_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_86,plain,(k2_xboole_0(k1_tarski(X1),X2)=k5_xboole_0(k1_tarski(X1),X2)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_82]), c_0_56])).
cnf(c_0_87,plain,(k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_81])).
fof(c_0_88, plain, (![X3]:![X4]:(~r1_xboole_0(k1_tarski(X3),k1_tarski(X4))|X3!=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_zfmisc_1])])).
cnf(c_0_89,plain,(~r1_tarski(X1,X2)|~r2_hidden(X3,k5_xboole_0(X2,k2_xboole_0(X2,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_29]), c_0_75]), c_0_58]), c_0_75]), c_0_84])])).
cnf(c_0_90,plain,(k5_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X1),X2))=k2_xboole_0(k1_tarski(X1),X2)|k1_tarski(X1)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87])).
cnf(c_0_91,plain,(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)))=k5_xboole_0(X1,X2)|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_29]), c_0_75]), c_0_64])).
cnf(c_0_92,plain,(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)))=k5_xboole_0(X2,X1)|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_29]), c_0_75]), c_0_57])).
fof(c_0_93, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X5),k1_enumset1(X6,X7,X8))), inference(variable_rename,[status(thm)],[t44_enumset1])).
fof(c_0_94, plain, (![X4]:![X5]:![X6]:k5_enumset1(X4,X4,X4,X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t89_enumset1])).
cnf(c_0_95,plain,(X1!=X2|~r1_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_96,plain,(r1_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))|k5_xboole_0(X1,k2_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)))!=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_34, c_0_29])).
cnf(c_0_97,plain,(k1_tarski(X1)!=k1_xboole_0|~r1_tarski(X2,k1_tarski(X1))|~r2_hidden(X3,k2_xboole_0(k1_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_98,plain,(k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
fof(c_0_99, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~r2_hidden(X6,X5))|k2_xboole_0(k2_tarski(X4,X6),X5)=X5)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_zfmisc_1])])).
fof(c_0_100, plain, (![X3]:![X4]:k2_enumset1(X3,X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t77_enumset1])).
cnf(c_0_101,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_102,plain,(k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_103,plain,(~r1_xboole_0(k1_tarski(X1),k1_tarski(X1))), inference(er,[status(thm)],[c_0_95])).
cnf(c_0_104,plain,(r1_xboole_0(X1,X1)|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_64]), c_0_58]), c_0_58]), c_0_77]), c_0_75])])).
cnf(c_0_105,plain,(k1_tarski(X1)!=k1_xboole_0|~r2_hidden(X2,k1_tarski(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_64]), c_0_47])])).
fof(c_0_106, plain, (![X3]:![X4]:(k3_xboole_0(k1_tarski(X3),k1_tarski(X4))!=k1_tarski(X3)|X3=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_zfmisc_1])])).
cnf(c_0_107,plain,(k5_xboole_0(k2_xboole_0(X1,X2),X1)=k5_xboole_0(X1,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_98, c_0_81])).
cnf(c_0_108,plain,(k2_xboole_0(k2_xboole_0(X1,X2),X1)=k2_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_59, c_0_81])).
cnf(c_0_109,plain,(k5_xboole_0(k5_xboole_0(X1,k2_xboole_0(X1,X2)),k2_xboole_0(X1,X2))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_87]), c_0_81])])).
cnf(c_0_110,plain,(k2_xboole_0(k2_tarski(X1,X2),X3)=X3|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_111,plain,(k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_112,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X2,X2,X2,X2,X3,X4))), inference(rw,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_113,plain,(~r1_tarski(k1_tarski(X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_114,plain,(r1_tarski(k1_tarski(X1),X2)|k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_105, c_0_33])).
cnf(c_0_115,plain,(X1=X2|k3_xboole_0(k1_tarski(X1),k1_tarski(X2))!=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_116,plain,(r1_tarski(X1,k2_xboole_0(k2_xboole_0(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_107]), c_0_108]), c_0_109])).
cnf(c_0_117,plain,(k2_xboole_0(k2_xboole_0(k1_tarski(X1),k5_enumset1(X1,X1,X1,X1,X1,X1,X2)),X3)=X3|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_111]), c_0_112])).
cnf(c_0_118,plain,(k1_tarski(X1)=k1_xboole_0|r2_hidden(X1,k1_tarski(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_75]), c_0_77]), c_0_71]), c_0_75])).
cnf(c_0_119,plain,(k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_113, c_0_114])).
fof(c_0_120, negated_conjecture, (~(![X1]:![X2]:![X3]:(~(X1=X2)=>k4_xboole_0(k2_xboole_0(X3,k1_tarski(X1)),k1_tarski(X2))=k2_xboole_0(k4_xboole_0(X3,k1_tarski(X2)),k1_tarski(X1))))), inference(assume_negation,[status(cth)],[t147_zfmisc_1])).
cnf(c_0_121,plain,(X1=X2|k5_xboole_0(k5_xboole_0(k1_tarski(X1),k1_tarski(X2)),k2_xboole_0(k1_tarski(X1),k1_tarski(X2)))!=k1_tarski(X1)), inference(rw,[status(thm)],[c_0_115, c_0_25])).
cnf(c_0_122,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X3,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_123,plain,(r2_hidden(X1,k1_tarski(X1))), inference(sr,[status(thm)],[c_0_118, c_0_119])).
fof(c_0_124, negated_conjecture, ((esk4_0!=esk5_0&k4_xboole_0(k2_xboole_0(esk6_0,k1_tarski(esk4_0)),k1_tarski(esk5_0))!=k2_xboole_0(k4_xboole_0(esk6_0,k1_tarski(esk5_0)),k1_tarski(esk4_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_120])])])])])).
fof(c_0_125, plain, (![X4]:![X5]:![X6]:(~r1_xboole_0(X4,X5)|k2_xboole_0(k4_xboole_0(X6,X4),X5)=k4_xboole_0(k2_xboole_0(X6,X5),X4))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_xboole_1])])])])).
cnf(c_0_126,plain,(X1=X2|~r1_tarski(k1_tarski(X1),k1_tarski(X2))), inference(spm,[status(thm)],[c_0_121, c_0_29])).
cnf(c_0_127,plain,(r1_tarski(k1_tarski(X1),k1_tarski(X2))|~r2_hidden(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_128,negated_conjecture,(k4_xboole_0(k2_xboole_0(esk6_0,k1_tarski(esk4_0)),k1_tarski(esk5_0))!=k2_xboole_0(k4_xboole_0(esk6_0,k1_tarski(esk5_0)),k1_tarski(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_129,plain,(k2_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k2_xboole_0(X1,X3),X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_130,plain,(X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_131,negated_conjecture,(k5_xboole_0(k2_xboole_0(esk6_0,k1_tarski(esk4_0)),k5_xboole_0(k5_xboole_0(k2_xboole_0(esk6_0,k1_tarski(esk4_0)),k1_tarski(esk5_0)),k2_xboole_0(k2_xboole_0(esk6_0,k1_tarski(esk4_0)),k1_tarski(esk5_0))))!=k2_xboole_0(k5_xboole_0(esk6_0,k5_xboole_0(k5_xboole_0(esk6_0,k1_tarski(esk5_0)),k2_xboole_0(esk6_0,k1_tarski(esk5_0)))),k1_tarski(esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_49]), c_0_49]), c_0_25]), c_0_25])).
cnf(c_0_132,plain,(k5_xboole_0(k2_xboole_0(X1,X3),k5_xboole_0(k5_xboole_0(k2_xboole_0(X1,X3),X2),k2_xboole_0(k2_xboole_0(X1,X3),X2)))=k2_xboole_0(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))),X3)|~r1_xboole_0(X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_49]), c_0_49]), c_0_25]), c_0_25])).
cnf(c_0_133,plain,(k2_xboole_0(k1_tarski(X1),k1_tarski(X2))=k5_xboole_0(k1_tarski(X1),k1_tarski(X2))|X1=X2), inference(spm,[status(thm)],[c_0_130, c_0_86])).
cnf(c_0_134,negated_conjecture,(~r1_xboole_0(k1_tarski(esk5_0),k1_tarski(esk4_0))), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_135,plain,(X1=X2|r1_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_133]), c_0_75])])).
cnf(c_0_136,negated_conjecture,(esk4_0!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_137,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_136]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 138
# Proof object clause steps            : 95
# Proof object formula steps           : 43
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 82
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 26
# Processed clauses                    : 711
# ...of these trivial                  : 33
# ...subsumed                          : 431
# ...remaining for further processing  : 247
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 32
# Generated clauses                    : 5391
# ...of the previous two non-trivial   : 4504
# Contextual simplify-reflections      : 25
# Paramodulations                      : 5379
# Factorizations                       : 0
# Equation resolutions                 : 11
# Current number of processed clauses  : 178
#    Positive orientable unit clauses  : 35
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 133
# Current number of unprocessed clauses: 3636
# ...number of literals in the above   : 9173
# Current number of archived formulas  : 0
# Current number of archived clauses   : 70
# Clause-clause subsumption calls (NU) : 3532
# Rec. Clause-clause subsumption calls : 3011
# Non-unit clause-clause subsumptions  : 243
# Unit Clause-clause subsumption calls : 211
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 110
# BW rewrite match successes           : 20
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 233123

# -------------------------------------------------
# User time                : 0.092 s
# System time              : 0.008 s
# Total time               : 0.100 s
# Maximum resident set size: 47144 pages
