# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t148_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:![X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4))), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', t148_zfmisc_1)).
fof(t21_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', t21_xboole_1)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', t48_xboole_1)).
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', d4_xboole_0)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', t12_xboole_1)).
fof(l1_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', l1_zfmisc_1)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', d1_tarski)).
fof(t16_xboole_1, axiom, (![X1]:![X2]:![X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), file('out/t148_zfmisc_1__15_premises__rerun__rerun.E_input', t16_xboole_1)).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4)))), inference(assume_negation,[status(cth)],[t148_zfmisc_1])).
fof(c_0_9, plain, (![X3]:![X4]:k3_xboole_0(X3,k2_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_10, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk5_3(X5,X6,X7),X7)|(~r2_hidden(esk5_3(X5,X6,X7),X5)|~r2_hidden(esk5_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk5_3(X5,X6,X7),X5)|r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk5_3(X5,X6,X7),X6)|r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
fof(c_0_12, negated_conjecture, (((r1_tarski(esk1_0,esk2_0)&(k3_xboole_0(esk2_0,esk3_0)=k1_tarski(esk4_0)&r2_hidden(esk4_0,esk1_0)))&k3_xboole_0(esk1_0,esk3_0)!=k1_tarski(esk4_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_13,plain,(k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_16,plain,(r2_hidden(X4,X2)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(k3_xboole_0(esk1_0,esk3_0)!=k1_tarski(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk5_3(X2,X3,X1),X1)|r2_hidden(esk5_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k2_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])])])).
cnf(c_0_22,plain,(r2_hidden(X4,X2)|X1!=k4_xboole_0(X2,k4_xboole_0(X2,X3))|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_23,negated_conjecture,(k3_xboole_0(esk2_0,esk3_0)=k1_tarski(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_24, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk6_2(X4,X5),X5)|esk6_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk6_2(X4,X5),X5)|esk6_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:k3_xboole_0(k3_xboole_0(X4,X5),X6)=k3_xboole_0(X4,k3_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_26,negated_conjecture,(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))!=k1_tarski(esk4_0)), inference(rw,[status(thm)],[c_0_17, c_0_14])).
cnf(c_0_27,plain,(X1=k4_xboole_0(X2,k4_xboole_0(X2,X3))|r2_hidden(esk5_3(X2,X3,X1),X3)|r2_hidden(esk5_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[c_0_18, c_0_14])).
cnf(c_0_28,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_29,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_31,negated_conjecture,(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))=k1_tarski(esk4_0)), inference(rw,[status(thm)],[c_0_23, c_0_14])).
cnf(c_0_32,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,plain,(r2_hidden(X4,X3)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_34,plain,(k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35,negated_conjecture,(r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_36,negated_conjecture,(r2_hidden(esk5_3(X1,X2,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)))|r2_hidden(esk5_3(X1,X2,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),X2)|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k1_tarski(esk4_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_37,plain,(k4_xboole_0(k1_tarski(X1),k4_xboole_0(k1_tarski(X1),X2))=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38,negated_conjecture,(r2_hidden(X1,esk2_0)|~r2_hidden(X1,k1_tarski(esk4_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])])).
cnf(c_0_40,plain,(r2_hidden(X4,X3)|X1!=k4_xboole_0(X2,k4_xboole_0(X2,X3))|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[c_0_33, c_0_14])).
cnf(c_0_41,plain,(k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3))=k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_14]), c_0_14]), c_0_14]), c_0_14])).
cnf(c_0_42,negated_conjecture,(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))=esk1_0), inference(spm,[status(thm)],[c_0_28, c_0_35])).
cnf(c_0_43,negated_conjecture,(r2_hidden(esk5_3(k1_tarski(X1),X2,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)))|r2_hidden(esk5_3(k1_tarski(X1),X2,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),X2)|k1_tarski(X1)!=k1_tarski(esk4_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44,negated_conjecture,(r2_hidden(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X3,k4_xboole_0(X3,X2)))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_46,negated_conjecture,(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X1))))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,X1))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47,plain,(X1=k3_xboole_0(X2,X3)|~r2_hidden(esk5_3(X2,X3,X1),X3)|~r2_hidden(esk5_3(X2,X3,X1),X2)|~r2_hidden(esk5_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_48,negated_conjecture,(r2_hidden(esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)))|r2_hidden(esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49,negated_conjecture,(r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_45, c_0_42])).
cnf(c_0_50,negated_conjecture,(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,k1_tarski(esk4_0)))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_46, c_0_31])).
cnf(c_0_51,plain,(r2_hidden(X4,X1)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_52,plain,(X1=k4_xboole_0(X2,k4_xboole_0(X2,X3))|~r2_hidden(esk5_3(X2,X3,X1),X3)|~r2_hidden(esk5_3(X2,X3,X1),X2)|~r2_hidden(esk5_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[c_0_47, c_0_14])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_48]), c_0_49])).
cnf(c_0_54,negated_conjecture,(r2_hidden(X1,k1_tarski(esk4_0))|~r2_hidden(X1,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)))), inference(spm,[status(thm)],[c_0_45, c_0_50])).
cnf(c_0_55,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk5_3(X2,X3,X1),X1)|r2_hidden(esk5_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_56,plain,(r2_hidden(X4,X1)|X1!=k4_xboole_0(X2,k4_xboole_0(X2,X3))|~r2_hidden(X4,X3)|~r2_hidden(X4,X2)), inference(rw,[status(thm)],[c_0_51, c_0_14])).
cnf(c_0_57,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_58,negated_conjecture,(k4_xboole_0(k1_tarski(esk4_0),k4_xboole_0(k1_tarski(esk4_0),esk2_0))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))|~r2_hidden(esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_59,plain,(X1=k4_xboole_0(X2,k4_xboole_0(X2,X3))|r2_hidden(esk5_3(X2,X3,X1),X2)|r2_hidden(esk5_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[c_0_55, c_0_14])).
cnf(c_0_60,plain,(r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_61,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_62,negated_conjecture,(k4_xboole_0(k1_tarski(esk4_0),k4_xboole_0(k1_tarski(esk4_0),esk2_0))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))|r2_hidden(esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),k1_tarski(esk4_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63,negated_conjecture,(r2_hidden(X1,esk3_0)|~r2_hidden(X1,k1_tarski(esk4_0))), inference(spm,[status(thm)],[c_0_45, c_0_31])).
cnf(c_0_64,negated_conjecture,(k4_xboole_0(k1_tarski(esk4_0),k4_xboole_0(k1_tarski(esk4_0),esk2_0))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))|~r2_hidden(esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),esk3_0)|~r2_hidden(esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))),esk1_0)), inference(spm,[status(thm)],[c_0_58, c_0_60])).
cnf(c_0_65,negated_conjecture,(k4_xboole_0(k1_tarski(esk4_0),k4_xboole_0(k1_tarski(esk4_0),esk2_0))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))|esk5_3(k1_tarski(esk4_0),esk2_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)))=esk4_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_66,negated_conjecture,(r2_hidden(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_63, c_0_39])).
cnf(c_0_67,negated_conjecture,(r2_hidden(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_68,negated_conjecture,(k4_xboole_0(k1_tarski(esk4_0),k4_xboole_0(k1_tarski(esk4_0),esk2_0))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66]), c_0_67])])).
cnf(c_0_69,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_68]), c_0_44])]), c_0_26]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 70
# Proof object clause steps            : 53
# Proof object formula steps           : 17
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 19
# Processed clauses                    : 90
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 5
# Generated clauses                    : 397
# ...of the previous two non-trivial   : 359
# Contextual simplify-reflections      : 6
# Paramodulations                      : 384
# Factorizations                       : 4
# Equation resolutions                 : 10
# Current number of processed clauses  : 56
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 37
# Current number of unprocessed clauses: 282
# ...number of literals in the above   : 890
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 529
# Rec. Clause-clause subsumption calls : 481
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9202

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43084 pages
