# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t148_zfmisc_1__10_premises__rerun__rerun.E_input', d1_tarski)).
fof(t148_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:![X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4))), file('out/t148_zfmisc_1__10_premises__rerun__rerun.E_input', t148_zfmisc_1)).
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t148_zfmisc_1__10_premises__rerun__rerun.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t148_zfmisc_1__10_premises__rerun__rerun.E_input', d3_tarski)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t148_zfmisc_1__10_premises__rerun__rerun.E_input', t28_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t148_zfmisc_1__10_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t148_zfmisc_1__10_premises__rerun__rerun.E_input', d10_xboole_0)).
fof(c_0_7, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4)))), inference(assume_negation,[status(cth)],[t148_zfmisc_1])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk6_3(X5,X6,X7),X7)|(~r2_hidden(esk6_3(X5,X6,X7),X5)|~r2_hidden(esk6_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk6_3(X5,X6,X7),X5)|r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk6_3(X5,X6,X7),X6)|r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
cnf(c_0_10,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, negated_conjecture, (((r1_tarski(esk2_0,esk3_0)&(k3_xboole_0(esk3_0,esk4_0)=k1_tarski(esk5_0)&r2_hidden(esk5_0,esk2_0)))&k3_xboole_0(esk2_0,esk4_0)!=k1_tarski(esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk7_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk7_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_13,plain,(r2_hidden(X4,X1)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_15,negated_conjecture,(k3_xboole_0(esk3_0,esk4_0)=k1_tarski(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,plain,(r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(r2_hidden(X1,k3_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(esk5_0=X1|~r2_hidden(X1,k3_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20,plain,(r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_21, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_22,plain,(r2_hidden(X4,X2)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_16])])).
cnf(c_0_24,plain,(r1_tarski(X1,k3_xboole_0(X2,X3))|~r2_hidden(esk7_2(X1,k3_xboole_0(X2,X3)),X3)|~r2_hidden(esk7_2(X1,k3_xboole_0(X2,X3)),X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25,negated_conjecture,(esk7_2(k3_xboole_0(esk3_0,esk4_0),X1)=esk5_0|r1_tarski(k3_xboole_0(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,negated_conjecture,(r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29,negated_conjecture,(r2_hidden(esk5_0,k3_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_23, c_0_15])).
cnf(c_0_30,negated_conjecture,(r1_tarski(k3_xboole_0(esk3_0,esk4_0),k3_xboole_0(X1,X2))|~r2_hidden(esk5_0,X2)|~r2_hidden(esk5_0,X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,negated_conjecture,(k3_xboole_0(esk2_0,esk3_0)=esk2_0), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,negated_conjecture,(r2_hidden(esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33,negated_conjecture,(r2_hidden(esk5_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_34, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_35,negated_conjecture,(r1_tarski(k3_xboole_0(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33])])).
cnf(c_0_36,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_37,plain,(r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk7_2(k3_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_28, c_0_20])).
cnf(c_0_38,negated_conjecture,(k3_xboole_0(esk2_0,k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_35]), c_0_36])).
cnf(c_0_39,plain,(r2_hidden(X4,X3)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40,negated_conjecture,(r1_tarski(k3_xboole_0(esk3_0,esk4_0),X1)|r2_hidden(esk7_2(k3_xboole_0(esk3_0,esk4_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_41,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_42,negated_conjecture,(r1_tarski(k3_xboole_0(esk3_0,esk4_0),k3_xboole_0(X1,esk2_0))|~r2_hidden(esk7_2(k3_xboole_0(esk3_0,esk4_0),k3_xboole_0(X1,esk2_0)),X1)), inference(spm,[status(thm)],[c_0_24, c_0_40])).
cnf(c_0_43,plain,(r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk7_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_41, c_0_20])).
fof(c_0_44, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_45,negated_conjecture,(k3_xboole_0(esk2_0,esk4_0)!=k1_tarski(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_46,plain,(r1_tarski(X1,k3_xboole_0(X2,X1))|~r2_hidden(esk7_2(X1,k3_xboole_0(X2,X1)),X2)), inference(spm,[status(thm)],[c_0_24, c_0_20])).
cnf(c_0_47,negated_conjecture,(r1_tarski(k3_xboole_0(esk3_0,esk4_0),k3_xboole_0(esk2_0,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_36])).
cnf(c_0_48,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49,negated_conjecture,(k3_xboole_0(esk3_0,esk4_0)!=k3_xboole_0(esk2_0,esk4_0)), inference(rw,[status(thm)],[c_0_45, c_0_15])).
cnf(c_0_50,negated_conjecture,(r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_41, c_0_31])).
cnf(c_0_51,plain,(r1_tarski(X1,k3_xboole_0(X1,X2))|~r2_hidden(esk7_2(X1,k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_46, c_0_36])).
cnf(c_0_52,negated_conjecture,(k3_xboole_0(k3_xboole_0(esk2_0,esk4_0),k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_47]), c_0_36])).
cnf(c_0_53,negated_conjecture,(~r1_tarski(k3_xboole_0(esk2_0,esk4_0),k3_xboole_0(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_47]), c_0_49])).
cnf(c_0_54,negated_conjecture,(r1_tarski(X1,esk3_0)|~r2_hidden(esk7_2(X1,esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_50])).
cnf(c_0_55,negated_conjecture,(~r2_hidden(esk7_2(k3_xboole_0(esk2_0,esk4_0),k3_xboole_0(esk3_0,esk4_0)),k3_xboole_0(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_56,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_57,negated_conjecture,(r1_tarski(k3_xboole_0(esk2_0,X1),esk3_0)), inference(spm,[status(thm)],[c_0_54, c_0_37])).
cnf(c_0_58,negated_conjecture,(~r2_hidden(esk7_2(k3_xboole_0(esk2_0,esk4_0),k3_xboole_0(esk3_0,esk4_0)),esk4_0)|~r2_hidden(esk7_2(k3_xboole_0(esk2_0,esk4_0),k3_xboole_0(esk3_0,esk4_0)),esk3_0)), inference(spm,[status(thm)],[c_0_55, c_0_18])).
cnf(c_0_59,negated_conjecture,(r2_hidden(X1,esk3_0)|~r2_hidden(X1,k3_xboole_0(esk2_0,X2))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_60,negated_conjecture,(~r2_hidden(esk7_2(k3_xboole_0(esk2_0,esk4_0),k3_xboole_0(esk3_0,esk4_0)),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_43]), c_0_53])).
cnf(c_0_61,negated_conjecture,(r1_tarski(k3_xboole_0(esk2_0,X1),X2)|r2_hidden(esk7_2(k3_xboole_0(esk2_0,X1),X2),esk3_0)), inference(spm,[status(thm)],[c_0_59, c_0_20])).
cnf(c_0_62,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_53]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 48
# Proof object formula steps           : 15
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 426
# ...of these trivial                  : 9
# ...subsumed                          : 248
# ...remaining for further processing  : 169
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 3
# Generated clauses                    : 1491
# ...of the previous two non-trivial   : 1212
# Contextual simplify-reflections      : 38
# Paramodulations                      : 1455
# Factorizations                       : 20
# Equation resolutions                 : 17
# Current number of processed clauses  : 137
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 109
# Current number of unprocessed clauses: 824
# ...number of literals in the above   : 2831
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 2077
# Rec. Clause-clause subsumption calls : 1890
# Non-unit clause-clause subsumptions  : 286
# Unit Clause-clause subsumption calls : 39
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 39
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20669

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.004 s
# Total time               : 0.036 s
# Maximum resident set size: 43080 pages
