# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t149_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:![X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2))), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t149_zfmisc_1)).
fof(d7_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0)), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', d7_xboole_0)).
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t4_xboole_0)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t3_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t28_xboole_1)).
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t16_xboole_1, axiom, (![X1]:![X2]:![X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t16_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t71_enumset1)).
fof(l27_zfmisc_1, axiom, (![X1]:![X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2))), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', l27_zfmisc_1)).
fof(t78_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_xboole_0(X1,X2)=>k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3))), file('out/t149_zfmisc_1__15_premises__rerun__rerun.E_input', t78_xboole_1)).
fof(c_0_12, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2)))), inference(assume_negation,[status(cth)],[t149_zfmisc_1])).
fof(c_0_13, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k3_xboole_0(X3,X4)=k1_xboole_0)&(k3_xboole_0(X3,X4)!=k1_xboole_0|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])])])).
fof(c_0_14, negated_conjecture, (((r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))&(r2_hidden(esk4_0,esk3_0)&r1_xboole_0(esk3_0,esk2_0)))&~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
fof(c_0_15, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk5_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
cnf(c_0_16,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(r1_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19,negated_conjecture,(k3_xboole_0(esk3_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_20, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk6_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk6_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_21, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_22, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_23, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_24, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:k3_xboole_0(k3_xboole_0(X4,X5),X6)=k3_xboole_0(X4,k3_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_26,negated_conjecture,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_17])])).
cnf(c_0_27,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,negated_conjecture,(r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_31,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_33, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
cnf(c_0_34,plain,(k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35,negated_conjecture,(r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_18, c_0_28])).
cnf(c_0_37,negated_conjecture,(k3_xboole_0(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))=k3_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38,negated_conjecture,(~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_39,plain,(k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41,negated_conjecture,(k3_xboole_0(esk3_0,k3_xboole_0(esk2_0,X1))=k3_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_34, c_0_19])).
cnf(c_0_42,negated_conjecture,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_16, c_0_35])).
cnf(c_0_43,plain,(~r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r2_hidden(X4,k3_xboole_0(X2,k3_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_36, c_0_34])).
cnf(c_0_44,negated_conjecture,(k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,k1_tarski(esk4_0)))=k3_xboole_0(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_37, c_0_34])).
fof(c_0_45, plain, (![X3]:![X4]:(r2_hidden(X3,X4)|r1_xboole_0(k1_tarski(X3),X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_46,negated_conjecture,(~r1_xboole_0(k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_47,plain,(r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_48, plain, (![X4]:![X5]:![X6]:(~r1_xboole_0(X4,X5)|k3_xboole_0(X4,k2_xboole_0(X5,X6))=k3_xboole_0(X4,X6))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_xboole_1])])])])).
cnf(c_0_49,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_50,negated_conjecture,(k3_xboole_0(esk3_0,k3_xboole_0(X1,esk2_0))=k3_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_41, c_0_28])).
cnf(c_0_51,negated_conjecture,(k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_42])).
cnf(c_0_52,negated_conjecture,(~r1_xboole_0(k1_tarski(esk4_0),k3_xboole_0(esk1_0,esk2_0))|~r2_hidden(X1,k3_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_53,plain,(r1_xboole_0(k1_tarski(X1),X2)|r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54,negated_conjecture,(k3_xboole_0(k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)),esk2_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55,plain,(k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56,plain,(k3_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_49, c_0_47])).
cnf(c_0_57,negated_conjecture,(k3_xboole_0(esk3_0,k3_xboole_0(X1,esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58,negated_conjecture,(r2_hidden(esk4_0,k3_xboole_0(esk1_0,esk2_0))|~r2_hidden(X1,k3_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_59,plain,(r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_60,negated_conjecture,(k3_xboole_0(esk2_0,k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_54, c_0_28])).
cnf(c_0_61,plain,(k3_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X3)))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_39]), c_0_40])).
cnf(c_0_62,negated_conjecture,(~r2_hidden(X1,k3_xboole_0(X2,esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63,negated_conjecture,(r1_xboole_0(esk1_0,esk2_0)|r2_hidden(esk4_0,k3_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_64,negated_conjecture,(r2_hidden(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_65,negated_conjecture,(~r1_xboole_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_28]), c_0_19])])).
cnf(c_0_66,negated_conjecture,(r1_xboole_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])])).
cnf(c_0_67,negated_conjecture,(k3_xboole_0(esk1_0,esk2_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_47]), c_0_28])).
cnf(c_0_68,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_66]), c_0_67]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 44
# Proof object formula steps           : 25
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 16
# Processed clauses                    : 225
# ...of these trivial                  : 13
# ...subsumed                          : 105
# ...remaining for further processing  : 107
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 10
# Generated clauses                    : 807
# ...of the previous two non-trivial   : 603
# Contextual simplify-reflections      : 1
# Paramodulations                      : 807
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 81
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 380
# ...number of literals in the above   : 660
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 526
# Rec. Clause-clause subsumption calls : 486
# Non-unit clause-clause subsumptions  : 35
# Unit Clause-clause subsumption calls : 34
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 86
# BW rewrite match successes           : 81
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8632

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 45372 pages
