# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t3_yellow_1)).
fof(t1_yellow_1, axiom, (![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t1_yellow_1)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', redefinition_r3_orders_2)).
fof(dt_k2_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1)))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', dt_k2_yellow_1)).
fof(fc5_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1)))))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1))))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', fc6_yellow_1)).
fof(t31_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))))))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t31_yellow_0)).
fof(t74_zfmisc_1, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2)))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t74_zfmisc_1)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t2_subset)).
fof(t14_yellow_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t14_yellow_1)).
fof(t6_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2))))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t6_yellow_0)).
fof(t1_subset, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', t1_subset)).
fof(d12_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>k4_yellow_0(X1)=k2_yellow_0(X1,k1_xboole_0))), file('out/t14_yellow_1__14_premises__rerun__rerun.E_input', d12_yellow_0)).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:(((((~r3_orders_2(k2_yellow_1(X4),X5,X6)|r1_tarski(X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4))&((((~r1_tarski(X5,X6)|r3_orders_2(k2_yellow_1(X4),X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])])])).
fof(c_0_14, plain, (![X2]:![X2]:(u1_struct_0(k2_yellow_1(X2))=X2&u1_orders_2(k2_yellow_1(X2))=k1_yellow_1(X2))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t1_yellow_1])])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_16,plain,(v1_xboole_0(X1)|r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, (![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&l1_orders_2(k2_yellow_1(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k2_yellow_1])])])).
fof(c_0_19, plain, (![X2]:![X2]:![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&(v3_orders_2(k2_yellow_1(X2))&(v4_orders_2(k2_yellow_1(X2))&v5_orders_2(k2_yellow_1(X2)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc5_yellow_1])])])).
fof(c_0_20, plain, (![X2]:((~v2_struct_0(k2_yellow_1(X2))|v1_xboole_0(X2))&(v1_orders_2(k2_yellow_1(X2))|v1_xboole_0(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
fof(c_0_21, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r1_lattice3(X5,X7,X6)|(X6!=k2_yellow_0(X5,X7)|~r2_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X7,X8)|r1_orders_2(X5,X8,X6)))|(X6!=k2_yellow_0(X5,X7)|~r2_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k2_yellow_0(X5,X7)|(m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k2_yellow_0(X5,X7)|(r1_lattice3(X5,X7,esk1_3(X5,X6,X7))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(r1_lattice3(X5,X7,esk1_3(X5,X6,X7))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&((((X6=k2_yellow_0(X5,X7)|(~r1_orders_2(X5,esk1_3(X5,X6,X7),X6)|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(~r1_orders_2(X5,esk1_3(X5,X6,X7),X6)|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])])).
cnf(c_0_22,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(r3_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_17])).
cnf(c_0_24,plain,(l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,plain,(v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,plain,(m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|X2=k2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28,plain,(v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29,plain,(X2=k2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,esk1_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,plain,(r1_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_17]), c_0_17])]), c_0_26])).
cnf(c_0_31,plain,(X1=k2_yellow_0(k2_yellow_1(X2),X3)|m1_subset_1(esk1_3(k2_yellow_1(X2),X1,X3),X2)|~r1_lattice3(k2_yellow_1(X2),X3,X1)|~m1_subset_1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_17]), c_0_28]), c_0_24])])).
fof(c_0_32, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|r1_tarski(X3,k3_tarski(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t74_zfmisc_1])])).
cnf(c_0_33,plain,(X1=k2_yellow_0(k2_yellow_1(X2),X3)|v1_xboole_0(X2)|~r1_lattice3(k2_yellow_1(X2),X3,X1)|~r1_tarski(esk1_3(k2_yellow_1(X2),X1,X3),X1)|~m1_subset_1(X1,X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_28]), c_0_24]), c_0_17])]), c_0_31])).
cnf(c_0_34,plain,(r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_35, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_36,plain,(k3_tarski(X1)=k2_yellow_0(k2_yellow_1(X2),X3)|v1_xboole_0(X2)|~r2_hidden(esk1_3(k2_yellow_1(X2),k3_tarski(X1),X3),X1)|~r1_lattice3(k2_yellow_1(X2),X3,k3_tarski(X1))|~m1_subset_1(k3_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_38, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))))), inference(assume_negation,[status(cth)],[t14_yellow_1])).
cnf(c_0_39,plain,(k3_tarski(X1)=k2_yellow_0(k2_yellow_1(X2),X3)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_lattice3(k2_yellow_1(X2),X3,k3_tarski(X1))|~m1_subset_1(esk1_3(k2_yellow_1(X2),k3_tarski(X1),X3),X1)|~m1_subset_1(k3_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_40, plain, (![X3]:![X4]:(((r2_lattice3(X3,k1_xboole_0,X4)|~m1_subset_1(X4,u1_struct_0(X3)))|~l1_orders_2(X3))&((r1_lattice3(X3,k1_xboole_0,X4)|~m1_subset_1(X4,u1_struct_0(X3)))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])])])).
fof(c_0_41, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|m1_subset_1(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_42, negated_conjecture, ((~v1_xboole_0(esk2_0)&(r2_hidden(k3_tarski(esk2_0),esk2_0)&k4_yellow_0(k2_yellow_1(esk2_0))!=k3_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_38])])])])).
fof(c_0_43, plain, (![X2]:(~l1_orders_2(X2)|k4_yellow_0(X2)=k2_yellow_0(X2,k1_xboole_0))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_yellow_0])])).
cnf(c_0_44,plain,(k2_yellow_0(k2_yellow_1(X1),X2)=k3_tarski(X1)|v1_xboole_0(X1)|~r1_lattice3(k2_yellow_1(X1),X2,k3_tarski(X1))|~m1_subset_1(k3_tarski(X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_27]), c_0_17]), c_0_17]), c_0_17]), c_0_17]), c_0_28]), c_0_24]), c_0_17]), c_0_17])])).
cnf(c_0_45,plain,(r1_lattice3(X1,k1_xboole_0,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47,negated_conjecture,(r2_hidden(k3_tarski(esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48,negated_conjecture,(k4_yellow_0(k2_yellow_1(esk2_0))!=k3_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49,plain,(k4_yellow_0(X1)=k2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50,plain,(k2_yellow_0(k2_yellow_1(X1),k1_xboole_0)=k3_tarski(X1)|v1_xboole_0(X1)|~m1_subset_1(k3_tarski(X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_24]), c_0_17])])).
cnf(c_0_51,negated_conjecture,(m1_subset_1(k3_tarski(esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52,negated_conjecture,(k2_yellow_0(k2_yellow_1(esk2_0),k1_xboole_0)!=k3_tarski(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_24])])).
cnf(c_0_53,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_53]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 28
# Proof object formula steps           : 27
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 30
# Processed clauses                    : 81
# ...of these trivial                  : 2
# ...subsumed                          : 3
# ...remaining for further processing  : 76
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 33
# ...of the previous two non-trivial   : 28
# Contextual simplify-reflections      : 6
# Paramodulations                      : 31
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 46
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 30
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 526
# Rec. Clause-clause subsumption calls : 36
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3595

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 43328 pages
