# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d15_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2)))))))))), file('out/t14_waybel_0__5_premises__rerun.E_input', d15_waybel_0)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t14_waybel_0__5_premises__rerun.E_input', t4_subset)).
fof(fraenkel_a_2_0_waybel_0, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_0_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(X5,X3)))))))), file('out/t14_waybel_0__5_premises__rerun.E_input', fraenkel_a_2_0_waybel_0)).
fof(t14_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2)))), file('out/t14_waybel_0__5_premises__rerun.E_input', t14_waybel_0)).
fof(dt_k3_waybel_0, axiom, (![X1]:![X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t14_waybel_0__5_premises__rerun.E_input', dt_k3_waybel_0)).
fof(t1_tarski, axiom, (![X1]:![X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2)), file('out/t14_waybel_0__5_premises__rerun.E_input', t1_tarski)).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:![X13]:(((((((((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((((r1_orders_2(X6,X9,esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((r2_hidden(esk1_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))))&(((((((~m1_subset_1(X11,u1_struct_0(X6))|(~r1_orders_2(X6,X9,X11)|~r2_hidden(X11,X7)))|r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))&(((((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X6))|X8=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((~r2_hidden(esk2_3(X6,X7,X8),X8)|(~m1_subset_1(X13,u1_struct_0(X6))|(~r1_orders_2(X6,esk2_3(X6,X7,X8),X13)|~r2_hidden(X13,X7))))|X8=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((m1_subset_1(esk3_3(X6,X7,X8),u1_struct_0(X6))|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((r1_orders_2(X6,esk2_3(X6,X7,X8),esk3_3(X6,X7,X8))|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((r2_hidden(esk3_3(X6,X7,X8),X7)|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k3_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])])])).
fof(c_0_7, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_8, plain, (![X6]:![X7]:![X8]:![X11]:![X12]:((((m1_subset_1(esk4_3(X6,X7,X8),u1_struct_0(X7))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((X6=esk4_3(X6,X7,X8)|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((m1_subset_1(esk5_3(X6,X7,X8),u1_struct_0(X7))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((r1_orders_2(X7,esk4_3(X6,X7,X8),esk5_3(X6,X7,X8))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&((r2_hidden(esk5_3(X6,X7,X8),X8)|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))))))&(((~m1_subset_1(X11,u1_struct_0(X7))|(X6!=X11|(~m1_subset_1(X12,u1_struct_0(X7))|(~r1_orders_2(X7,X11,X12)|~r2_hidden(X12,X8)))))|r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_waybel_0])])])])])])])])).
fof(c_0_9, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2))))), inference(assume_negation,[status(cth)],[t14_waybel_0])).
fof(c_0_10, plain, (![X3]:![X4]:((~l1_orders_2(X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|m1_subset_1(k3_waybel_0(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_waybel_0])])).
cnf(c_0_11,plain,(r1_orders_2(X1,X4,esk1_4(X1,X2,X3,X4))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(r2_hidden(esk1_4(X1,X2,X3,X4),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,plain,(v2_struct_0(X2)|r2_hidden(X3,a_2_0_waybel_0(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X4,X1)|~r1_orders_2(X2,X5,X4)|~m1_subset_1(X4,u1_struct_0(X2))|X3!=X5|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_15, negated_conjecture, (((~v2_struct_0(esk6_0)&l1_orders_2(esk6_0))&(m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(esk6_0)))&k3_waybel_0(esk6_0,esk7_0)!=a_2_0_waybel_0(esk6_0,esk7_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_16,plain,(m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,plain,(r1_orders_2(X1,X2,esk1_4(X1,X3,X4,X2))|X4!=k3_waybel_0(X1,X3)|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18,plain,(r2_hidden(esk1_4(X1,X2,X3,X4),X2)|X3!=k3_waybel_0(X1,X2)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_13, c_0_12])).
cnf(c_0_19,plain,(v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X4)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_14, c_0_12])])).
cnf(c_0_20,negated_conjecture,(m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,negated_conjecture,(l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,negated_conjecture,(~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k3_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_12, c_0_16])).
cnf(c_0_24,plain,(r1_orders_2(X1,X2,esk1_4(X1,X3,k3_waybel_0(X1,X4),X2))|k3_waybel_0(X1,X4)!=k3_waybel_0(X1,X3)|~r2_hidden(X2,k3_waybel_0(X1,X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_17, c_0_16])).
cnf(c_0_25,plain,(r2_hidden(esk1_4(X1,X2,k3_waybel_0(X1,X3),X4),X2)|k3_waybel_0(X1,X3)!=k3_waybel_0(X1,X2)|~r2_hidden(X4,k3_waybel_0(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_16])).
cnf(c_0_26,negated_conjecture,(r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,X1,X2)|~r2_hidden(X2,esk7_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_27,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_20]), c_0_21])])).
cnf(c_0_28,negated_conjecture,(r1_orders_2(esk6_0,X1,esk1_4(esk6_0,esk7_0,k3_waybel_0(esk6_0,X2),X1))|k3_waybel_0(esk6_0,X2)!=k3_waybel_0(esk6_0,esk7_0)|~r2_hidden(X1,k3_waybel_0(esk6_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_20]), c_0_21])])).
cnf(c_0_29,negated_conjecture,(r2_hidden(esk1_4(esk6_0,esk7_0,k3_waybel_0(esk6_0,X1),X2),esk7_0)|k3_waybel_0(esk6_0,X1)!=k3_waybel_0(esk6_0,esk7_0)|~r2_hidden(X2,k3_waybel_0(esk6_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_20]), c_0_21])])).
cnf(c_0_30,plain,(v2_struct_0(X2)|X3=esk4_3(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_31,negated_conjecture,(r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,X1,X2)|~r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))|~r2_hidden(X2,esk7_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,negated_conjecture,(r1_orders_2(esk6_0,X1,esk1_4(esk6_0,esk7_0,k3_waybel_0(esk6_0,esk7_0),X1))|~r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_28, c_0_20])).
cnf(c_0_33,negated_conjecture,(r2_hidden(esk1_4(esk6_0,esk7_0,k3_waybel_0(esk6_0,esk7_0),X1),esk7_0)|~r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_29, c_0_20])).
fof(c_0_34, plain, (![X4]:![X5]:(((~r2_hidden(esk8_2(X4,X5),X4)|~r2_hidden(esk8_2(X4,X5),X5))|X4=X5)&((r2_hidden(esk8_2(X4,X5),X4)|r2_hidden(esk8_2(X4,X5),X5))|X4=X5))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tarski])])])])])])).
cnf(c_0_35,plain,(r2_hidden(X4,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X5,X2)|~r1_orders_2(X1,X4,X5)|~m1_subset_1(X5,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_36,negated_conjecture,(esk4_3(X1,esk6_0,esk7_0)=X1|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_37,negated_conjecture,(r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_38,plain,(X1=X2|r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(esk8_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39,plain,(r2_hidden(X1,X2)|X2!=k3_waybel_0(X3,X4)|~r1_orders_2(X3,X1,X5)|~r2_hidden(X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_35, c_0_12])).
cnf(c_0_40,plain,(v2_struct_0(X2)|m1_subset_1(esk4_3(X3,X2,X1),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_41,negated_conjecture,(esk4_3(X1,esk6_0,esk7_0)=X1|~r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42,negated_conjecture,(esk4_3(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0)=esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1)|a_2_0_waybel_0(esk6_0,esk7_0)=X1|r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),X1)), inference(spm,[status(thm)],[c_0_36, c_0_38])).
cnf(c_0_43,negated_conjecture,(k3_waybel_0(esk6_0,esk7_0)!=a_2_0_waybel_0(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_44,plain,(r2_hidden(X1,k3_waybel_0(X2,X3))|k3_waybel_0(X2,X3)!=k3_waybel_0(X2,X4)|~r1_orders_2(X2,X1,X5)|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_39, c_0_16])).
cnf(c_0_45,negated_conjecture,(m1_subset_1(esk4_3(X1,esk6_0,esk7_0),u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_46,negated_conjecture,(esk4_3(esk8_2(X1,k3_waybel_0(esk6_0,esk7_0)),esk6_0,esk7_0)=esk8_2(X1,k3_waybel_0(esk6_0,esk7_0))|X1=k3_waybel_0(esk6_0,esk7_0)|r2_hidden(esk8_2(X1,k3_waybel_0(esk6_0,esk7_0)),X1)), inference(spm,[status(thm)],[c_0_41, c_0_38])).
cnf(c_0_47,negated_conjecture,(esk4_3(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),esk6_0,esk7_0)=esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])).
cnf(c_0_48,plain,(v2_struct_0(X2)|r1_orders_2(X2,esk4_3(X3,X2,X1),esk5_3(X3,X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_49,plain,(v2_struct_0(X2)|r2_hidden(esk5_3(X3,X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_50,negated_conjecture,(r2_hidden(X1,k3_waybel_0(esk6_0,X2))|k3_waybel_0(esk6_0,X2)!=k3_waybel_0(esk6_0,esk7_0)|~r1_orders_2(esk6_0,X1,X3)|~r2_hidden(X3,esk7_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_20]), c_0_21])])).
cnf(c_0_51,negated_conjecture,(X1=k3_waybel_0(esk6_0,esk7_0)|r2_hidden(esk8_2(X1,k3_waybel_0(esk6_0,esk7_0)),X1)|m1_subset_1(esk8_2(X1,k3_waybel_0(esk6_0,esk7_0)),u1_struct_0(esk6_0))|~r2_hidden(esk8_2(X1,k3_waybel_0(esk6_0,esk7_0)),a_2_0_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52,negated_conjecture,(m1_subset_1(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),u1_struct_0(esk6_0))|~r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),a_2_0_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_45, c_0_47])).
cnf(c_0_53,negated_conjecture,(r1_orders_2(esk6_0,esk4_3(X1,esk6_0,esk7_0),esk5_3(X1,esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_54,negated_conjecture,(r2_hidden(esk5_3(X1,esk6_0,esk7_0),esk7_0)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_55,negated_conjecture,(r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,X1,X2)|~r2_hidden(X2,esk7_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(spm,[status(thm)],[c_0_50, c_0_20])).
cnf(c_0_56,negated_conjecture,(m1_subset_1(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_38]), c_0_43]), c_0_37]), c_0_52])).
cnf(c_0_57,negated_conjecture,(a_2_0_waybel_0(esk6_0,esk7_0)=X1|r1_orders_2(esk6_0,esk4_3(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0),esk5_3(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0))|r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),X1)), inference(spm,[status(thm)],[c_0_53, c_0_38])).
cnf(c_0_58,negated_conjecture,(r2_hidden(esk5_3(X1,esk6_0,esk7_0),esk7_0)|~r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_54, c_0_37])).
cnf(c_0_59,negated_conjecture,(a_2_0_waybel_0(esk6_0,esk7_0)=X1|r2_hidden(esk5_3(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0),esk7_0)|r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),X1)), inference(spm,[status(thm)],[c_0_54, c_0_38])).
cnf(c_0_60,negated_conjecture,(r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),k3_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),X1)|~r2_hidden(X1,esk7_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_61,negated_conjecture,(a_2_0_waybel_0(esk6_0,esk7_0)=X1|r1_orders_2(esk6_0,esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),esk5_3(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0))|r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),X1),X1)), inference(spm,[status(thm)],[c_0_57, c_0_42])).
cnf(c_0_62,negated_conjecture,(r2_hidden(esk5_3(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),esk6_0,esk7_0),esk7_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_43])).
cnf(c_0_63,plain,(X1=X2|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(esk8_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_64,negated_conjecture,(r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),k3_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])]), c_0_43])).
cnf(c_0_65,negated_conjecture,(~r2_hidden(esk8_2(a_2_0_waybel_0(esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0)),a_2_0_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_43])).
cnf(c_0_66,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_37]), c_0_64])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 67
# Proof object clause steps            : 54
# Proof object formula steps           : 13
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 40
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 497
# ...of these trivial                  : 7
# ...subsumed                          : 145
# ...remaining for further processing  : 345
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 1
# Generated clauses                    : 2225
# ...of the previous two non-trivial   : 2011
# Contextual simplify-reflections      : 214
# Paramodulations                      : 2222
# Factorizations                       : 2
# Equation resolutions                 : 1
# Current number of processed clauses  : 306
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 294
# Current number of unprocessed clauses: 1480
# ...number of literals in the above   : 7851
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 16682
# Rec. Clause-clause subsumption calls : 6156
# Non-unit clause-clause subsumptions  : 369
# Unit Clause-clause subsumption calls : 74
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 85291

# -------------------------------------------------
# User time                : 0.080 s
# System time              : 0.000 s
# Total time               : 0.080 s
# Maximum resident set size: 45284 pages
