# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t148_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t148_relat_1)).
fof(t100_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k7_relat_1(k7_relat_1(X3,X1),X2)=k7_relat_1(X3,k3_xboole_0(X1,X2)))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t100_relat_1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t145_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k9_relat_1(X2,X1)=k9_relat_1(X2,k3_xboole_0(k1_relat_1(X2),X1)))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t145_relat_1)).
fof(t90_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k1_relat_1(k7_relat_1(X2,X1))=k3_xboole_0(k1_relat_1(X2),X1))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t90_relat_1)).
fof(t101_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k7_relat_1(k7_relat_1(X2,X1),X1)=k7_relat_1(X2,X1))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t101_relat_1)).
fof(t150_funct_1, conjecture, (![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>k9_relat_1(X3,k3_xboole_0(X1,k10_relat_1(X3,X2)))=k3_xboole_0(k9_relat_1(X3,X1),X2))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t150_funct_1)).
fof(fc8_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k7_relat_1(X1,X2))&v1_funct_1(k7_relat_1(X1,X2))))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', fc8_funct_1)).
fof(t139_funct_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k10_relat_1(k7_relat_1(X3,X1),X2)=k3_xboole_0(X1,k10_relat_1(X3,X2)))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t139_funct_1)).
fof(t148_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>k9_relat_1(X2,k10_relat_1(X2,X1))=k3_xboole_0(X1,k9_relat_1(X2,k1_relat_1(X2))))), file('out/t150_funct_1__11_premises__rerun__rerun.E_input', t148_funct_1)).
fof(c_0_10, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k2_relat_1(k7_relat_1(X4,X3))=k9_relat_1(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k7_relat_1(k7_relat_1(X6,X4),X5)=k7_relat_1(X6,k3_xboole_0(X4,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t100_relat_1])])).
fof(c_0_12, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_13,plain,(k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k9_relat_1(X4,X3)=k9_relat_1(X4,k3_xboole_0(k1_relat_1(X4),X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t145_relat_1])])).
fof(c_0_17, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k1_relat_1(k7_relat_1(X4,X3))=k3_xboole_0(k1_relat_1(X4),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t90_relat_1])])).
fof(c_0_18, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k7_relat_1(k7_relat_1(X4,X3),X3)=k7_relat_1(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t101_relat_1])])).
cnf(c_0_19,plain,(k2_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3))=k9_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20,plain,(k7_relat_1(X1,k3_xboole_0(X2,X3))=k7_relat_1(k7_relat_1(X1,X3),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21,plain,(k9_relat_1(X1,X2)=k9_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(k7_relat_1(k7_relat_1(X1,X2),X2)=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(k2_relat_1(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3),X4))=k9_relat_1(X1,k3_xboole_0(k3_xboole_0(X3,X2),X4))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25,plain,(k2_relat_1(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3),X4))=k9_relat_1(X1,k3_xboole_0(X2,k3_xboole_0(X3,X4)))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_14])).
cnf(c_0_26,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27,plain,(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3),k3_xboole_0(X2,X3))=k7_relat_1(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_14])).
cnf(c_0_28,plain,(k9_relat_1(X1,k3_xboole_0(X2,X3))=k9_relat_1(k7_relat_1(X1,X2),X3)|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_13, c_0_19])).
cnf(c_0_29,plain,(k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,k3_xboole_0(X2,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_23])).
cnf(c_0_30,plain,(k9_relat_1(X1,k3_xboole_0(X2,k1_relat_1(X1)))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_21, c_0_15])).
cnf(c_0_31,plain,(k9_relat_1(X1,k3_xboole_0(k3_xboole_0(X2,X3),X4))=k9_relat_1(X1,k3_xboole_0(X3,k3_xboole_0(X2,X4)))|~v1_relat_1(k7_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32,plain,(k2_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3))=k9_relat_1(X1,k3_xboole_0(X3,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_13, c_0_20])).
cnf(c_0_33,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3)))=k9_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_14])).
cnf(c_0_34,plain,(k7_relat_1(k7_relat_1(X1,X2),k3_xboole_0(X2,X2))=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_23])).
cnf(c_0_35,plain,(k9_relat_1(k7_relat_1(X1,k1_relat_1(X1)),X2)=k9_relat_1(X1,X2)|~v1_relat_1(k7_relat_1(X1,k1_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_21, c_0_28])).
cnf(c_0_36,plain,(k9_relat_1(k7_relat_1(X1,X2),k3_xboole_0(X3,X3))=k9_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_29])).
cnf(c_0_37,plain,(k9_relat_1(X1,k3_xboole_0(X2,k3_xboole_0(X3,k1_relat_1(X1))))=k9_relat_1(X1,k3_xboole_0(X3,X2))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38,plain,(k2_relat_1(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3),X4))=k9_relat_1(X1,k3_xboole_0(X4,k3_xboole_0(X2,X3)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_14])).
cnf(c_0_39,plain,(k2_relat_1(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3),X4))=k9_relat_1(X1,k3_xboole_0(k3_xboole_0(X3,X4),X2))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_14])).
fof(c_0_40, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>k9_relat_1(X3,k3_xboole_0(X1,k10_relat_1(X3,X2)))=k3_xboole_0(k9_relat_1(X3,X1),X2)))), inference(assume_negation,[status(cth)],[t150_funct_1])).
cnf(c_0_41,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k9_relat_1(X1,k3_xboole_0(X2,k3_xboole_0(X2,X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_42,plain,(k9_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))=k9_relat_1(X1,k3_xboole_0(X2,X2))|~v1_relat_1(k7_relat_1(X1,k1_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43,plain,(k9_relat_1(X1,k3_xboole_0(k3_xboole_0(X2,k1_relat_1(X1)),X3))=k9_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(k7_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_37, c_0_15])).
cnf(c_0_44,plain,(k9_relat_1(X1,k3_xboole_0(k3_xboole_0(X2,X3),X4))=k9_relat_1(X1,k3_xboole_0(X3,k3_xboole_0(X4,X2)))|~v1_relat_1(k7_relat_1(X1,X4))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_45, plain, (![X3]:![X4]:![X4]:((v1_relat_1(k7_relat_1(X3,X4))|(~v1_relat_1(X3)|~v1_funct_1(X3)))&(v1_funct_1(k7_relat_1(X3,X4))|(~v1_relat_1(X3)|~v1_funct_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_funct_1])])])])])).
fof(c_0_46, negated_conjecture, (((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&k9_relat_1(esk3_0,k3_xboole_0(esk1_0,k10_relat_1(esk3_0,esk2_0)))!=k3_xboole_0(k9_relat_1(esk3_0,esk1_0),esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])).
cnf(c_0_47,plain,(k9_relat_1(X1,k3_xboole_0(X2,X3))=k9_relat_1(k7_relat_1(X1,X3),X2)|~v1_relat_1(k7_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_15])).
cnf(c_0_48,plain,(k9_relat_1(X1,k3_xboole_0(X2,k3_xboole_0(X2,X2)))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_41])).
cnf(c_0_49,plain,(k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(k7_relat_1(X1,X2),X2)|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_13, c_0_23])).
cnf(c_0_50,plain,(k9_relat_1(X1,k3_xboole_0(X2,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_13, c_0_29])).
cnf(c_0_51,plain,(k9_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))=k2_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(k7_relat_1(X1,k1_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_29, c_0_42])).
cnf(c_0_52,plain,(k9_relat_1(X1,k3_xboole_0(k1_relat_1(X1),k3_xboole_0(X2,X3)))=k9_relat_1(X1,k3_xboole_0(X3,X2))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_53,plain,(v1_relat_1(k7_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56,plain,(k9_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)),X2)=k9_relat_1(X1,X2)|~v1_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_57,plain,(k2_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)))=k9_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)),X2)|~v1_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58,plain,(k2_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X3)))=k9_relat_1(X1,k3_xboole_0(X3,X2))|~v1_relat_1(k7_relat_1(X1,k1_relat_1(X1)))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59,negated_conjecture,(v1_relat_1(k7_relat_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
cnf(c_0_60,plain,(k7_relat_1(k7_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X3)),X2),X3)=k7_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X3)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_14])).
cnf(c_0_61,plain,(v1_funct_1(k7_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_62,plain,(k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(k7_relat_1(X1,X3),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_14, c_0_20])).
cnf(c_0_63,plain,(k2_relat_1(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3),X4))=k9_relat_1(X1,k3_xboole_0(k3_xboole_0(X2,X3),X4))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_14])).
cnf(c_0_64,plain,(k2_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)))=k9_relat_1(X1,X2)|~v1_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65,negated_conjecture,(k2_relat_1(k7_relat_1(esk3_0,k3_xboole_0(X1,X2)))=k9_relat_1(esk3_0,k3_xboole_0(X2,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_59]), c_0_55])])).
cnf(c_0_66,plain,(k7_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X3)),X3)=k7_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(k7_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X3)),X2))|~v1_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X3)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_60])).
cnf(c_0_67,negated_conjecture,(v1_relat_1(k7_relat_1(k7_relat_1(esk3_0,X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_14]), c_0_55])])).
cnf(c_0_68,plain,(v1_funct_1(k7_relat_1(k7_relat_1(X1,X2),X3))|~v1_funct_1(k7_relat_1(X1,X3))|~v1_relat_1(k7_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_69,plain,(k7_relat_1(k7_relat_1(X1,X2),k3_xboole_0(X2,k3_xboole_0(X2,X2)))=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_34])).
cnf(c_0_70,plain,(k9_relat_1(X1,k3_xboole_0(k3_xboole_0(X2,X3),k3_xboole_0(X2,X3)))=k2_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_63, c_0_27])).
cnf(c_0_71,negated_conjecture,(k9_relat_1(esk3_0,k3_xboole_0(X1,X1))=k9_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_59]), c_0_55])])).
cnf(c_0_72,negated_conjecture,(k7_relat_1(k7_relat_1(esk3_0,k3_xboole_0(X1,X2)),X2)=k7_relat_1(esk3_0,k3_xboole_0(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_59]), c_0_55])])).
cnf(c_0_73,plain,(v1_funct_1(k7_relat_1(X1,X2))|~v1_funct_1(k7_relat_1(X1,k3_xboole_0(X2,k3_xboole_0(X2,X2))))|~v1_relat_1(k7_relat_1(X1,k3_xboole_0(X2,k3_xboole_0(X2,X2))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_74,negated_conjecture,(k2_relat_1(k7_relat_1(k7_relat_1(esk3_0,X1),X2))=k9_relat_1(esk3_0,k3_xboole_0(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_55])])).
cnf(c_0_75,negated_conjecture,(k7_relat_1(k7_relat_1(esk3_0,X1),k3_xboole_0(X2,X1))=k7_relat_1(esk3_0,k3_xboole_0(X2,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_72]), c_0_55])])).
cnf(c_0_76,plain,(v1_funct_1(k7_relat_1(X1,X2))|~v1_funct_1(k7_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)),X2))|~v1_relat_1(k7_relat_1(k7_relat_1(X1,k3_xboole_0(X2,X2)),X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_73, c_0_20])).
fof(c_0_77, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k10_relat_1(k7_relat_1(X6,X4),X5)=k3_xboole_0(X4,k10_relat_1(X6,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_funct_1])])).
fof(c_0_78, plain, (![X3]:![X4]:((~v1_relat_1(X4)|~v1_funct_1(X4))|k9_relat_1(X4,k10_relat_1(X4,X3))=k3_xboole_0(X3,k9_relat_1(X4,k1_relat_1(X4))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_funct_1])])).
cnf(c_0_79,negated_conjecture,(k9_relat_1(esk3_0,k3_xboole_0(X1,X2))=k9_relat_1(k7_relat_1(esk3_0,X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_74]), c_0_59])])).
cnf(c_0_80,plain,(k3_xboole_0(X1,k1_relat_1(X2))=k1_relat_1(k7_relat_1(X2,X1))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_15, c_0_22])).
cnf(c_0_81,negated_conjecture,(k7_relat_1(esk3_0,k3_xboole_0(X1,X1))=k7_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_75]), c_0_55])])).
cnf(c_0_82,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k9_relat_1(X1,k3_xboole_0(X2,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_23])).
cnf(c_0_83,negated_conjecture,(v1_funct_1(k7_relat_1(esk3_0,X1))|~v1_funct_1(k7_relat_1(esk3_0,k3_xboole_0(X1,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_72]), c_0_59]), c_0_55])])).
cnf(c_0_84,negated_conjecture,(k9_relat_1(esk3_0,k3_xboole_0(esk1_0,k10_relat_1(esk3_0,esk2_0)))!=k3_xboole_0(k9_relat_1(esk3_0,esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_85,plain,(k10_relat_1(k7_relat_1(X1,X2),X3)=k3_xboole_0(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_86,plain,(k9_relat_1(X1,k10_relat_1(X1,X2))=k3_xboole_0(X2,k9_relat_1(X1,k1_relat_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_87,negated_conjecture,(k9_relat_1(k7_relat_1(esk3_0,X1),k1_relat_1(X2))=k9_relat_1(esk3_0,k1_relat_1(k7_relat_1(X2,X1)))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_88,negated_conjecture,(k7_relat_1(k7_relat_1(esk3_0,X1),X1)=k7_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_81]), c_0_55])])).
cnf(c_0_89,negated_conjecture,(k9_relat_1(esk3_0,k1_relat_1(k7_relat_1(esk3_0,X1)))=k9_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_71]), c_0_55])])).
cnf(c_0_90,negated_conjecture,(v1_funct_1(k7_relat_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_61]), c_0_54]), c_0_55])])).
cnf(c_0_91,negated_conjecture,(k9_relat_1(esk3_0,k3_xboole_0(esk1_0,k10_relat_1(esk3_0,esk2_0)))!=k3_xboole_0(esk2_0,k9_relat_1(esk3_0,esk1_0))), inference(rw,[status(thm)],[c_0_84, c_0_15])).
cnf(c_0_92,plain,(k9_relat_1(k7_relat_1(X1,X2),k10_relat_1(X3,X4))=k9_relat_1(X1,k10_relat_1(k7_relat_1(X3,X2),X4))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_28, c_0_85])).
cnf(c_0_93,negated_conjecture,(k9_relat_1(k7_relat_1(esk3_0,X1),k10_relat_1(k7_relat_1(esk3_0,X1),X2))=k3_xboole_0(X2,k9_relat_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_59])])).
cnf(c_0_94,negated_conjecture,(k9_relat_1(esk3_0,k10_relat_1(k7_relat_1(esk3_0,esk1_0),esk2_0))!=k3_xboole_0(esk2_0,k9_relat_1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_85]), c_0_55])])).
cnf(c_0_95,negated_conjecture,(k9_relat_1(esk3_0,k10_relat_1(k7_relat_1(esk3_0,X1),X2))=k3_xboole_0(X2,k9_relat_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_88]), c_0_59]), c_0_55])])).
cnf(c_0_96,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 97
# Proof object clause steps            : 76
# Proof object formula steps           : 21
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 61
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 2636
# ...of these trivial                  : 409
# ...subsumed                          : 1737
# ...remaining for further processing  : 490
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 80
# Generated clauses                    : 43897
# ...of the previous two non-trivial   : 38376
# Contextual simplify-reflections      : 589
# Paramodulations                      : 43897
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 391
#    Positive orientable unit clauses  : 79
#    Positive unorientable unit clauses: 9
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 302
# Current number of unprocessed clauses: 31907
# ...number of literals in the above   : 132467
# Current number of archived formulas  : 0
# Current number of archived clauses   : 99
# Clause-clause subsumption calls (NU) : 42420
# Rec. Clause-clause subsumption calls : 32464
# Non-unit clause-clause subsumptions  : 1978
# Unit Clause-clause subsumption calls : 1665
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 812
# BW rewrite match successes           : 194
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1028212

# -------------------------------------------------
# User time                : 0.580 s
# System time              : 0.012 s
# Total time               : 0.592 s
# Maximum resident set size: 44996 pages
