# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t145_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k9_relat_1(X2,X1)=k9_relat_1(X2,k3_xboole_0(k1_relat_1(X2),X1)))), file('out/t150_funct_1__10_premises__rerun.E_input', t145_relat_1)).
fof(t12_setfam_1, axiom, (![X1]:![X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), file('out/t150_funct_1__10_premises__rerun.E_input', t12_setfam_1)).
fof(t90_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k1_relat_1(k7_relat_1(X2,X1))=k3_xboole_0(k1_relat_1(X2),X1))), file('out/t150_funct_1__10_premises__rerun.E_input', t90_relat_1)).
fof(t100_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k7_relat_1(k7_relat_1(X3,X1),X2)=k7_relat_1(X3,k3_xboole_0(X1,X2)))), file('out/t150_funct_1__10_premises__rerun.E_input', t100_relat_1)).
fof(t150_funct_1, conjecture, (![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>k9_relat_1(X3,k3_xboole_0(X1,k10_relat_1(X3,X2)))=k3_xboole_0(k9_relat_1(X3,X1),X2))), file('out/t150_funct_1__10_premises__rerun.E_input', t150_funct_1)).
fof(t101_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k7_relat_1(k7_relat_1(X2,X1),X1)=k7_relat_1(X2,X1))), file('out/t150_funct_1__10_premises__rerun.E_input', t101_relat_1)).
fof(fc8_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k7_relat_1(X1,X2))&v1_funct_1(k7_relat_1(X1,X2))))), file('out/t150_funct_1__10_premises__rerun.E_input', fc8_funct_1)).
fof(t148_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1))), file('out/t150_funct_1__10_premises__rerun.E_input', t148_relat_1)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t150_funct_1__10_premises__rerun.E_input', commutativity_k2_tarski)).
fof(t148_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>k9_relat_1(X2,k10_relat_1(X2,X1))=k3_xboole_0(X1,k9_relat_1(X2,k1_relat_1(X2))))), file('out/t150_funct_1__10_premises__rerun.E_input', t148_funct_1)).
fof(t139_funct_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k10_relat_1(k7_relat_1(X3,X1),X2)=k3_xboole_0(X1,k10_relat_1(X3,X2)))), file('out/t150_funct_1__10_premises__rerun.E_input', t139_funct_1)).
fof(c_0_11, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k9_relat_1(X4,X3)=k9_relat_1(X4,k3_xboole_0(k1_relat_1(X4),X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t145_relat_1])])).
fof(c_0_12, plain, (![X3]:![X4]:k1_setfam_1(k2_tarski(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_13, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k1_relat_1(k7_relat_1(X4,X3))=k3_xboole_0(k1_relat_1(X4),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t90_relat_1])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k7_relat_1(k7_relat_1(X6,X4),X5)=k7_relat_1(X6,k3_xboole_0(X4,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t100_relat_1])])).
fof(c_0_15, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>k9_relat_1(X3,k3_xboole_0(X1,k10_relat_1(X3,X2)))=k3_xboole_0(k9_relat_1(X3,X1),X2)))), inference(assume_negation,[status(cth)],[t150_funct_1])).
cnf(c_0_16,plain,(k9_relat_1(X1,X2)=k9_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k7_relat_1(k7_relat_1(X4,X3),X3)=k7_relat_1(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t101_relat_1])])).
cnf(c_0_20,plain,(k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, (![X3]:![X4]:![X4]:((v1_relat_1(k7_relat_1(X3,X4))|(~v1_relat_1(X3)|~v1_funct_1(X3)))&(v1_funct_1(k7_relat_1(X3,X4))|(~v1_relat_1(X3)|~v1_funct_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_funct_1])])])])])).
fof(c_0_22, negated_conjecture, (((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&k9_relat_1(esk3_0,k3_xboole_0(esk1_0,k10_relat_1(esk3_0,esk2_0)))!=k3_xboole_0(k9_relat_1(esk3_0,esk1_0),esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
cnf(c_0_23,plain,(k9_relat_1(X1,X2)=k9_relat_1(X1,k1_setfam_1(k2_tarski(k1_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24,plain,(k1_relat_1(k7_relat_1(X1,X2))=k1_setfam_1(k2_tarski(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_25,plain,(k7_relat_1(k7_relat_1(X1,X2),X2)=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_20, c_0_17])).
cnf(c_0_27,plain,(v1_relat_1(k7_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31,plain,(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3),k1_setfam_1(k2_tarski(X2,X3)))=k7_relat_1(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_32, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k2_relat_1(k7_relat_1(X4,X3))=k9_relat_1(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
cnf(c_0_33,plain,(k7_relat_1(k7_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3))),X2),X3)=k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3)))|~v1_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34,negated_conjecture,(v1_relat_1(k7_relat_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])])).
cnf(c_0_35,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3)))=k9_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_26])).
cnf(c_0_36,plain,(k7_relat_1(k7_relat_1(X1,X2),k1_setfam_1(k2_tarski(X2,X2)))=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_31, c_0_25])).
cnf(c_0_37,plain,(k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38,plain,(k7_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3))),X3)=k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3)))|~v1_relat_1(k7_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3))),X2))|~v1_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_33])).
cnf(c_0_39,negated_conjecture,(v1_relat_1(k7_relat_1(k7_relat_1(esk3_0,X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_26]), c_0_29])])).
fof(c_0_40, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_41,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k9_relat_1(X1,k1_setfam_1(k2_tarski(X2,k1_setfam_1(k2_tarski(X2,X2)))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42,plain,(k2_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3))=k9_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_37, c_0_26])).
cnf(c_0_43,negated_conjecture,(k7_relat_1(k7_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,X2))),X2)=k7_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_34]), c_0_29])])).
cnf(c_0_44,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45,plain,(k9_relat_1(X1,k1_setfam_1(k2_tarski(X2,k1_setfam_1(k2_tarski(X2,X2)))))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_41])).
cnf(c_0_46,negated_conjecture,(k9_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,k1_setfam_1(k2_tarski(X2,X1)))))=k2_relat_1(k7_relat_1(esk3_0,k1_setfam_1(k2_tarski(X2,X1))))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_29])]), c_0_44])).
cnf(c_0_47,plain,(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X3)))=k7_relat_1(k7_relat_1(X1,X3),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_44])).
cnf(c_0_48,negated_conjecture,(k2_relat_1(k7_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,X1))))=k9_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_29])])).
cnf(c_0_49,plain,(v1_funct_1(k7_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_50,plain,(k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(k7_relat_1(X1,X3),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_47])).
cnf(c_0_51,plain,(k9_relat_1(X1,k1_setfam_1(k2_tarski(X2,X2)))=k2_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_42, c_0_25])).
cnf(c_0_52,negated_conjecture,(k9_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,X1)))=k9_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_48]), c_0_29])])).
cnf(c_0_53,plain,(v1_funct_1(k7_relat_1(k7_relat_1(X1,X2),X3))|~v1_funct_1(k7_relat_1(X1,X3))|~v1_relat_1(k7_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54,plain,(k7_relat_1(k7_relat_1(X1,X2),k1_setfam_1(k2_tarski(X2,k1_setfam_1(k2_tarski(X2,X2)))))=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_31, c_0_36])).
cnf(c_0_55,negated_conjecture,(k2_relat_1(k7_relat_1(esk3_0,X1))=k9_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_29])])).
cnf(c_0_56,plain,(v1_funct_1(k7_relat_1(X1,X2))|~v1_funct_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,k1_setfam_1(k2_tarski(X2,X2))))))|~v1_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,k1_setfam_1(k2_tarski(X2,X2))))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
fof(c_0_57, plain, (![X3]:![X4]:((~v1_relat_1(X4)|~v1_funct_1(X4))|k9_relat_1(X4,k10_relat_1(X4,X3))=k3_xboole_0(X3,k9_relat_1(X4,k1_relat_1(X4))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_funct_1])])).
cnf(c_0_58,negated_conjecture,(k2_relat_1(k7_relat_1(k7_relat_1(esk3_0,X1),X2))=k9_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_26]), c_0_29])])).
cnf(c_0_59,plain,(v1_funct_1(k7_relat_1(X1,X2))|~v1_funct_1(k7_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X2))),X2))|~v1_relat_1(k7_relat_1(k7_relat_1(X1,k1_setfam_1(k2_tarski(X2,X2))),X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_56, c_0_47])).
fof(c_0_60, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k10_relat_1(k7_relat_1(X6,X4),X5)=k3_xboole_0(X4,k10_relat_1(X6,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_funct_1])])).
cnf(c_0_61,negated_conjecture,(k9_relat_1(esk3_0,k3_xboole_0(esk1_0,k10_relat_1(esk3_0,esk2_0)))!=k3_xboole_0(k9_relat_1(esk3_0,esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_62,plain,(k9_relat_1(X1,k10_relat_1(X1,X2))=k3_xboole_0(X2,k9_relat_1(X1,k1_relat_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63,negated_conjecture,(k9_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,X2)))=k9_relat_1(k7_relat_1(esk3_0,X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_58]), c_0_34])])).
cnf(c_0_64,plain,(k1_setfam_1(k2_tarski(X1,k1_relat_1(X2)))=k1_relat_1(k7_relat_1(X2,X1))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_24, c_0_44])).
cnf(c_0_65,plain,(k7_relat_1(k7_relat_1(k7_relat_1(X1,X2),X2),X2)=k7_relat_1(X1,X2)|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_36])).
cnf(c_0_66,negated_conjecture,(k7_relat_1(k7_relat_1(k7_relat_1(esk3_0,X1),X2),X2)=k7_relat_1(k7_relat_1(esk3_0,X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_26]), c_0_29])])).
cnf(c_0_67,plain,(k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k9_relat_1(X1,k1_setfam_1(k2_tarski(X2,X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_25])).
cnf(c_0_68,negated_conjecture,(v1_funct_1(k7_relat_1(esk3_0,X1))|~v1_funct_1(k7_relat_1(esk3_0,k1_setfam_1(k2_tarski(X1,X1))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_43]), c_0_34]), c_0_29])])).
cnf(c_0_69,plain,(k10_relat_1(k7_relat_1(X1,X2),X3)=k3_xboole_0(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70,negated_conjecture,(k9_relat_1(esk3_0,k1_setfam_1(k2_tarski(esk1_0,k10_relat_1(esk3_0,esk2_0))))!=k1_setfam_1(k2_tarski(k9_relat_1(esk3_0,esk1_0),esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_17]), c_0_17])).
cnf(c_0_71,plain,(k9_relat_1(X1,k10_relat_1(X1,X2))=k1_setfam_1(k2_tarski(X2,k9_relat_1(X1,k1_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(rw,[status(thm)],[c_0_62, c_0_17])).
cnf(c_0_72,negated_conjecture,(k9_relat_1(k7_relat_1(esk3_0,X1),k1_relat_1(X2))=k9_relat_1(esk3_0,k1_relat_1(k7_relat_1(X2,X1)))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_73,negated_conjecture,(k7_relat_1(k7_relat_1(esk3_0,X1),X1)=k7_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_34]), c_0_29])])).
cnf(c_0_74,negated_conjecture,(k9_relat_1(esk3_0,k1_relat_1(k7_relat_1(esk3_0,X1)))=k9_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_52]), c_0_29])])).
cnf(c_0_75,negated_conjecture,(v1_funct_1(k7_relat_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_49]), c_0_28]), c_0_29])])).
cnf(c_0_76,plain,(k10_relat_1(k7_relat_1(X1,X2),X3)=k1_setfam_1(k2_tarski(X2,k10_relat_1(X1,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_69, c_0_17])).
cnf(c_0_77,negated_conjecture,(k9_relat_1(esk3_0,k1_setfam_1(k2_tarski(esk1_0,k10_relat_1(esk3_0,esk2_0))))!=k1_setfam_1(k2_tarski(esk2_0,k9_relat_1(esk3_0,esk1_0)))), inference(rw,[status(thm)],[c_0_70, c_0_44])).
cnf(c_0_78,negated_conjecture,(k9_relat_1(k7_relat_1(esk3_0,X1),k10_relat_1(k7_relat_1(esk3_0,X1),X2))=k1_setfam_1(k2_tarski(X2,k9_relat_1(esk3_0,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_74]), c_0_75]), c_0_34])])).
cnf(c_0_79,negated_conjecture,(k9_relat_1(k7_relat_1(esk3_0,X1),k10_relat_1(X2,X3))=k9_relat_1(esk3_0,k10_relat_1(k7_relat_1(X2,X1),X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_63, c_0_76])).
cnf(c_0_80,negated_conjecture,(k9_relat_1(esk3_0,k10_relat_1(k7_relat_1(esk3_0,esk1_0),esk2_0))!=k1_setfam_1(k2_tarski(esk2_0,k9_relat_1(esk3_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_76]), c_0_29])])).
cnf(c_0_81,negated_conjecture,(k9_relat_1(esk3_0,k10_relat_1(k7_relat_1(esk3_0,X1),X2))=k1_setfam_1(k2_tarski(X2,k9_relat_1(esk3_0,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_73]), c_0_34])])).
cnf(c_0_82,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_81])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 60
# Proof object formula steps           : 23
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 13
# Processed clauses                    : 1453
# ...of these trivial                  : 229
# ...subsumed                          : 861
# ...remaining for further processing  : 363
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 82
# Generated clauses                    : 19325
# ...of the previous two non-trivial   : 15136
# Contextual simplify-reflections      : 333
# Paramodulations                      : 19325
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 266
#    Positive orientable unit clauses  : 83
#    Positive unorientable unit clauses: 7
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 175
# Current number of unprocessed clauses: 12297
# ...number of literals in the above   : 48147
# Current number of archived formulas  : 0
# Current number of archived clauses   : 98
# Clause-clause subsumption calls (NU) : 16120
# Rec. Clause-clause subsumption calls : 11344
# Non-unit clause-clause subsumptions  : 1056
# Unit Clause-clause subsumption calls : 856
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 798
# BW rewrite match successes           : 171
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 482337

# -------------------------------------------------
# User time                : 0.288 s
# System time              : 0.004 s
# Total time               : 0.292 s
# Maximum resident set size: 44996 pages
