# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t79_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2)), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t79_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t2_boole, axiom, (![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t2_boole)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(d13_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:![X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X5,X4),X1)&r2_hidden(X5,X2)))))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', d13_relat_1)).
fof(d4_relat_1, axiom, (![X1]:![X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', d4_relat_1)).
fof(t148_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t148_relat_1)).
fof(t95_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(k7_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1)))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t95_relat_1)).
fof(t151_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(k9_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1)))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t151_relat_1)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t5_boole)).
fof(t60_relat_1, axiom, ((k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0)), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t60_relat_1)).
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t4_xboole_0)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t151_relat_1__13_premises__rerun__rerun.E_input', t3_xboole_0)).
fof(c_0_13, plain, (![X3]:![X4]:r1_xboole_0(k4_xboole_0(X3,X4),X4)), inference(variable_rename,[status(thm)],[t79_xboole_1])).
fof(c_0_14, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_15, plain, (![X2]:k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_16, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_17, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:![X7]:![X8]:![X13]:((((((r2_hidden(k4_tarski(esk1_4(X6,X7,X8,X9),X9),X6)|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk1_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(k4_tarski(X11,X9),X6)|~r2_hidden(X11,X7))|r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(esk2_3(X6,X7,X8),X8)|(~r2_hidden(k4_tarski(X13,esk2_3(X6,X7,X8)),X6)|~r2_hidden(X13,X7)))|X8=k9_relat_1(X6,X7))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk3_3(X6,X7,X8),esk2_3(X6,X7,X8)),X6)|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk3_3(X6,X7,X8),X7)|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_relat_1])])])])])])])).
fof(c_0_18, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk4_3(X5,X6,X7)),X5))|X6!=k1_relat_1(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k1_relat_1(X5)))&(((~r2_hidden(esk5_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk5_2(X5,X6),X11),X5))|X6=k1_relat_1(X5))&((r2_hidden(esk5_2(X5,X6),X6)|r2_hidden(k4_tarski(esk5_2(X5,X6),esk6_2(X5,X6)),X5))|X6=k1_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])])).
fof(c_0_19, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k2_relat_1(k7_relat_1(X4,X3))=k9_relat_1(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
fof(c_0_20, plain, (![X3]:![X4]:(((k7_relat_1(X4,X3)!=k1_xboole_0|r1_xboole_0(k1_relat_1(X4),X3))|~v1_relat_1(X4))&((~r1_xboole_0(k1_relat_1(X4),X3)|k7_relat_1(X4,X3)=k1_xboole_0)|~v1_relat_1(X4)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_relat_1])])])).
fof(c_0_21, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(k9_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1))))), inference(assume_negation,[status(cth)],[t151_relat_1])).
cnf(c_0_22,plain,(r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_27,plain,(r2_hidden(X4,X2)|~v1_relat_1(X1)|X2!=k9_relat_1(X1,X3)|~r2_hidden(X5,X3)|~r2_hidden(k4_tarski(X5,X4),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28,plain,(r2_hidden(k4_tarski(X3,esk4_3(X2,X1,X3)),X2)|X1!=k1_relat_1(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29,plain,(k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,plain,(k7_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(k1_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31,plain,(k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
fof(c_0_32, negated_conjecture, ((v1_relat_1(esk8_0)&((k9_relat_1(esk8_0,esk7_0)!=k1_xboole_0|~r1_xboole_0(k1_relat_1(esk8_0),esk7_0))&(k9_relat_1(esk8_0,esk7_0)=k1_xboole_0|r1_xboole_0(k1_relat_1(esk8_0),esk7_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
fof(c_0_33, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk9_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
cnf(c_0_34,plain,(r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_35,plain,(k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_36,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37,plain,(r2_hidden(X1,k9_relat_1(X2,X3))|~r2_hidden(k4_tarski(X4,X1),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_38,plain,(r2_hidden(k4_tarski(X1,esk4_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_39,plain,(k9_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(k1_relat_1(X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_40,negated_conjecture,(r1_xboole_0(k1_relat_1(esk8_0),esk7_0)|k9_relat_1(esk8_0,esk7_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41,negated_conjecture,(v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_44,plain,(r2_hidden(esk4_3(X1,k1_relat_1(X1),X2),k9_relat_1(X1,X3))|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X2,X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45,negated_conjecture,(k9_relat_1(esk8_0,esk7_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
cnf(c_0_46,plain,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_35]), c_0_43])])).
fof(c_0_47, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk10_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk10_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
cnf(c_0_48,plain,(r2_hidden(esk9_2(X1,X2),k3_xboole_0(X1,X2))|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_49,negated_conjecture,(~r2_hidden(X1,k1_relat_1(esk8_0))|~r2_hidden(X1,esk7_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_41])]), c_0_46])).
cnf(c_0_50,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk10_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_51,plain,(r2_hidden(esk9_2(X1,X2),k3_xboole_0(X2,X1))|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_48, c_0_25])).
cnf(c_0_52,negated_conjecture,(r1_xboole_0(X1,k1_relat_1(esk8_0))|~r2_hidden(esk10_2(X1,k1_relat_1(esk8_0)),esk7_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_53,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk10_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54,negated_conjecture,(~r1_xboole_0(k1_relat_1(esk8_0),esk7_0)|k9_relat_1(esk8_0,esk7_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_55,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_42, c_0_51])).
cnf(c_0_56,negated_conjecture,(r1_xboole_0(esk7_0,k1_relat_1(esk8_0))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_57,negated_conjecture,(~r1_xboole_0(k1_relat_1(esk8_0),esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_45])])).
cnf(c_0_58,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 33
# Proof object formula steps           : 26
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 27
# Processed clauses                    : 10173
# ...of these trivial                  : 6
# ...subsumed                          : 9467
# ...remaining for further processing  : 700
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 82954
# ...of the previous two non-trivial   : 78269
# Contextual simplify-reflections      : 3078
# Paramodulations                      : 82947
# Factorizations                       : 2
# Equation resolutions                 : 5
# Current number of processed clauses  : 691
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 671
# Current number of unprocessed clauses: 68120
# ...number of literals in the above   : 270096
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 243351
# Rec. Clause-clause subsumption calls : 161739
# Non-unit clause-clause subsumptions  : 11285
# Unit Clause-clause subsumption calls : 12
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1195902

# -------------------------------------------------
# User time                : 1.208 s
# System time              : 0.028 s
# Total time               : 1.236 s
# Maximum resident set size: 53736 pages
