# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_zfmisc_1, axiom, (![X1]:![X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), file('out/t155_funct_2__11_premises__rerun.E_input', t113_zfmisc_1)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t155_funct_2__11_premises__rerun.E_input', l13_xboole_0)).
fof(rc1_xboole_0, axiom, (?[X1]:v1_xboole_0(X1)), file('out/t155_funct_2__11_premises__rerun.E_input', rc1_xboole_0)).
fof(dt_k6_partfun1, axiom, (![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))), file('out/t155_funct_2__11_premises__rerun.E_input', dt_k6_partfun1)).
fof(d7_partfun1, axiom, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:(X4=k5_partfun1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>?[X6]:((v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(X6=X5&(v1_partfun1(X6,X1)&r1_partfun1(X3,X6)))))))), file('out/t155_funct_2__11_premises__rerun.E_input', d7_partfun1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t155_funct_2__11_premises__rerun.E_input', d3_tarski)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t155_funct_2__11_premises__rerun.E_input', t3_subset)).
fof(t155_funct_2, conjecture, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_partfun1(X3,X4)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|r2_hidden(X4,k5_partfun1(X1,X2,X3))))))), file('out/t155_funct_2__11_premises__rerun.E_input', t155_funct_2)).
fof(t5_subset, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3))))), file('out/t155_funct_2__11_premises__rerun.E_input', t5_subset)).
fof(cc5_funct_2, axiom, (![X1]:![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1)))))), file('out/t155_funct_2__11_premises__rerun.E_input', cc5_funct_2)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t155_funct_2__11_premises__rerun.E_input', d1_xboole_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t155_funct_2__11_premises__rerun.E_input', d10_xboole_0)).
fof(c_0_12, plain, (![X3]:![X4]:![X3]:![X4]:((k2_zfmisc_1(X3,X4)!=k1_xboole_0|(X3=k1_xboole_0|X4=k1_xboole_0))&((X3!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)&(X4!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])])])).
fof(c_0_13, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
fof(c_0_14, plain, (v1_xboole_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_xboole_0])])).
fof(c_0_15, plain, (![X2]:![X2]:(v1_partfun1(k6_partfun1(X2),X2)&m1_subset_1(k6_partfun1(X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k6_partfun1])])])).
cnf(c_0_16,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X11]:![X13]:![X10]:![X15]:(((((((v1_funct_1(esk7_5(X7,X8,X9,X10,X11))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((m1_subset_1(esk7_5(X7,X8,X9,X10,X11),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((esk7_5(X7,X8,X9,X10,X11)=X11|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&((((v1_partfun1(esk7_5(X7,X8,X9,X10,X11),X7)|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((r1_partfun1(X9,esk7_5(X7,X8,X9,X10,X11))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))))&(((((~v1_funct_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))|(X13!=X11|(~v1_partfun1(X13,X7)|~r1_partfun1(X9,X13))))|r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((~r2_hidden(esk8_4(X7,X8,X9,X10),X10)|((~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))|(X15!=esk8_4(X7,X8,X9,X10)|(~v1_partfun1(X15,X7)|~r1_partfun1(X9,X15)))))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((((v1_funct_1(esk9_4(X7,X8,X9,X10))|r2_hidden(esk8_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((m1_subset_1(esk9_4(X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|r2_hidden(esk8_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((esk9_4(X7,X8,X9,X10)=esk8_4(X7,X8,X9,X10)|r2_hidden(esk8_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&((((v1_partfun1(esk9_4(X7,X8,X9,X10),X7)|r2_hidden(esk8_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((r1_partfun1(X9,esk9_4(X7,X8,X9,X10))|r2_hidden(esk8_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_partfun1])])])])])])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk2_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk2_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_21, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_22, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_partfun1(X3,X4)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|r2_hidden(X4,k5_partfun1(X1,X2,X3)))))))), inference(assume_negation,[status(cth)],[t155_funct_2])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:(~r2_hidden(X4,X5)|(~m1_subset_1(X5,k1_zfmisc_1(X6))|~v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])])])).
cnf(c_0_24,plain,(m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_26,plain,(esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_27,plain,(m1_subset_1(esk7_5(X2,X3,X1,X4,X5),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r2_hidden(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_30, negated_conjecture, (((v1_funct_1(esk5_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,esk3_0,esk4_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0)))))&(r1_partfun1(esk5_0,esk6_0)&((esk4_0!=k1_xboole_0|esk3_0=k1_xboole_0)&~r2_hidden(esk6_0,k5_partfun1(esk3_0,esk4_0,esk5_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])])])).
cnf(c_0_31,plain,(r2_hidden(X5,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r1_partfun1(X1,X6)|~v1_partfun1(X6,X2)|X6!=X5|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X6)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_32, plain, (![X4]:![X5]:![X6]:((((v1_funct_1(X6)|(~v1_funct_1(X6)|~v1_funct_2(X6,X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))|v1_xboole_0(X5))&(((v1_partfun1(X6,X4)|(~v1_funct_1(X6)|~v1_funct_2(X6,X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))|v1_xboole_0(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])])])).
cnf(c_0_33,plain,(~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34,plain,(m1_subset_1(k6_partfun1(k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_35,plain,(v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_18, c_0_26])).
fof(c_0_36, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk10_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_37,plain,(m1_subset_1(esk7_5(X1,X2,X3,k5_partfun1(X1,X2,X3),X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~r2_hidden(X4,k5_partfun1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_38,plain,(esk7_5(X2,X3,X1,X4,X5)=X5|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r2_hidden(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_39,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41,plain,(r2_hidden(X1,k5_partfun1(X2,X3,X4))|~r1_partfun1(X4,X1)|~v1_funct_1(X1)|~v1_funct_1(X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_partfun1(X1,X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
cnf(c_0_42,negated_conjecture,(r1_partfun1(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44,negated_conjecture,(v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_45,plain,(v1_xboole_0(X1)|v1_partfun1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X2,X3,X1)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46,negated_conjecture,(v1_funct_2(esk6_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_47,negated_conjecture,(m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_48,plain,(~r2_hidden(X1,k6_partfun1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_49,plain,(v1_xboole_0(X1)|r2_hidden(esk10_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50,plain,(~v1_funct_1(X1)|~r2_hidden(X2,esk7_5(X3,X4,X1,k5_partfun1(X3,X4,X1),X5))|~r2_hidden(X5,k5_partfun1(X3,X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(k2_zfmisc_1(X3,X4))), inference(spm,[status(thm)],[c_0_33, c_0_37])).
cnf(c_0_51,plain,(esk7_5(X1,X2,X3,k5_partfun1(X1,X2,X3),X4)=X4|~v1_funct_1(X3)|~r2_hidden(X4,k5_partfun1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_52,plain,(r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_53,negated_conjecture,(r2_hidden(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(X1,esk5_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_54,negated_conjecture,(r2_hidden(esk6_0,k5_partfun1(X1,X2,esk5_0))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_partfun1(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_55,negated_conjecture,(v1_partfun1(esk6_0,esk3_0)|v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_43]), c_0_47])])).
cnf(c_0_56,plain,(v1_xboole_0(k6_partfun1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57,negated_conjecture,(r2_hidden(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_39, c_0_47])).
cnf(c_0_58,plain,(~v1_funct_1(X1)|~r2_hidden(X2,k5_partfun1(X3,X4,X1))|~r2_hidden(X5,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(k2_zfmisc_1(X3,X4))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_59, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_60,negated_conjecture,(r1_tarski(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(esk2_2(X1,k2_zfmisc_1(esk3_0,esk4_0)),esk5_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61,plain,(r1_tarski(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_62,negated_conjecture,(~r2_hidden(esk6_0,k5_partfun1(esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_63,negated_conjecture,(r2_hidden(esk6_0,k5_partfun1(esk3_0,X1,esk5_0))|v1_xboole_0(esk4_0)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X1)))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X1)))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64,plain,(k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_17, c_0_56])).
cnf(c_0_65,negated_conjecture,(r1_tarski(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(esk2_2(X1,k2_zfmisc_1(esk3_0,esk4_0)),esk6_0)), inference(spm,[status(thm)],[c_0_52, c_0_57])).
cnf(c_0_66,plain,(v1_xboole_0(k5_partfun1(X1,X2,X3))|~v1_funct_1(X3)|~r2_hidden(X4,esk10_1(k5_partfun1(X1,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_58, c_0_49])).
cnf(c_0_67,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68,negated_conjecture,(r1_tarski(esk5_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_69,negated_conjecture,(v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_47]), c_0_40])])).
cnf(c_0_70,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_71,plain,(~r2_hidden(X1,k1_xboole_0)), inference(rw,[status(thm)],[c_0_48, c_0_64])).
cnf(c_0_72,negated_conjecture,(r1_tarski(esk6_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_65, c_0_61])).
cnf(c_0_73,plain,(v1_xboole_0(esk10_1(k5_partfun1(X1,X2,X3)))|v1_xboole_0(k5_partfun1(X1,X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_66, c_0_49])).
cnf(c_0_74,negated_conjecture,(k2_zfmisc_1(esk3_0,esk4_0)=esk5_0|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),esk5_0)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_75,negated_conjecture,(esk4_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_17, c_0_69])).
cnf(c_0_76,plain,(k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_77,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_71, c_0_61])).
cnf(c_0_78,negated_conjecture,(k2_zfmisc_1(esk3_0,esk4_0)=esk6_0|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),esk6_0)), inference(spm,[status(thm)],[c_0_67, c_0_72])).
cnf(c_0_79,plain,(esk10_1(k5_partfun1(X1,X2,X3))=k1_xboole_0|v1_xboole_0(k5_partfun1(X1,X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_17, c_0_73])).
cnf(c_0_80,negated_conjecture,(esk5_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_75]), c_0_76]), c_0_75]), c_0_76]), c_0_77])])).
cnf(c_0_81,negated_conjecture,(esk6_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_75]), c_0_76]), c_0_75]), c_0_76]), c_0_77])])).
cnf(c_0_82,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_83,negated_conjecture,(esk3_0=k1_xboole_0|esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_84,plain,(r2_hidden(k1_xboole_0,k5_partfun1(X1,X2,X3))|v1_xboole_0(k5_partfun1(X1,X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_49, c_0_79])).
cnf(c_0_85,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_86,negated_conjecture,(r1_partfun1(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_80]), c_0_81])).
cnf(c_0_87,negated_conjecture,(v1_funct_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_44, c_0_80])).
cnf(c_0_88,plain,(m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_82, c_0_77])).
cnf(c_0_89,plain,(v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_90,negated_conjecture,(esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_75])])).
cnf(c_0_91,plain,(k5_partfun1(X1,X2,X3)=k1_xboole_0|r2_hidden(k1_xboole_0,k5_partfun1(X1,X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_17, c_0_84])).
cnf(c_0_92,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_85])).
cnf(c_0_93,negated_conjecture,(r2_hidden(k1_xboole_0,k5_partfun1(X1,X2,k1_xboole_0))|~v1_partfun1(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_86]), c_0_87]), c_0_88])])).
cnf(c_0_94,plain,(v1_partfun1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_89, c_0_64])).
cnf(c_0_95,negated_conjecture,(~r2_hidden(k1_xboole_0,k5_partfun1(k1_xboole_0,k1_xboole_0,k1_xboole_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_75]), c_0_81]), c_0_90]), c_0_80])).
cnf(c_0_96,plain,(k5_partfun1(X1,k1_xboole_0,X2)=k1_xboole_0|r2_hidden(k1_xboole_0,k5_partfun1(X1,k1_xboole_0,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_76]), c_0_35])])).
cnf(c_0_97,plain,(m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_82, c_0_92])).
cnf(c_0_98,negated_conjecture,(r2_hidden(k1_xboole_0,k5_partfun1(k1_xboole_0,X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_99,negated_conjecture,(k5_partfun1(k1_xboole_0,k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_87]), c_0_97])])).
cnf(c_0_100,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_71]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 101
# Proof object clause steps            : 76
# Proof object formula steps           : 25
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 39
# Processed clauses                    : 189
# ...of these trivial                  : 8
# ...subsumed                          : 29
# ...remaining for further processing  : 152
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 31
# Generated clauses                    : 186
# ...of the previous two non-trivial   : 182
# Contextual simplify-reflections      : 0
# Paramodulations                      : 175
# Factorizations                       : 0
# Equation resolutions                 : 12
# Current number of processed clauses  : 72
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 59
# ...number of literals in the above   : 273
# Current number of archived formulas  : 0
# Current number of archived clauses   : 69
# Clause-clause subsumption calls (NU) : 2192
# Rec. Clause-clause subsumption calls : 692
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 162
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6453

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.004 s
# Total time               : 0.024 s
# Maximum resident set size: 45316 pages
