# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t155_funct_2, conjecture, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_partfun1(X3,X4)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|r2_hidden(X4,k5_partfun1(X1,X2,X3))))))), file('out/t155_funct_2__9_premises__rerun.E_input', t155_funct_2)).
fof(d7_partfun1, axiom, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:(X4=k5_partfun1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>?[X6]:((v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(X6=X5&(v1_partfun1(X6,X1)&r1_partfun1(X3,X6)))))))), file('out/t155_funct_2__9_premises__rerun.E_input', d7_partfun1)).
fof(t113_zfmisc_1, axiom, (![X1]:![X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), file('out/t155_funct_2__9_premises__rerun.E_input', t113_zfmisc_1)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t155_funct_2__9_premises__rerun.E_input', l13_xboole_0)).
fof(rc1_xboole_0, axiom, (?[X1]:v1_xboole_0(X1)), file('out/t155_funct_2__9_premises__rerun.E_input', rc1_xboole_0)).
fof(cc5_funct_2, axiom, (![X1]:![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1)))))), file('out/t155_funct_2__9_premises__rerun.E_input', cc5_funct_2)).
fof(dt_k6_partfun1, axiom, (![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))), file('out/t155_funct_2__9_premises__rerun.E_input', dt_k6_partfun1)).
fof(t5_subset, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3))))), file('out/t155_funct_2__9_premises__rerun.E_input', t5_subset)).
fof(cc4_relset_1, axiom, (![X1]:![X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3)))), file('out/t155_funct_2__9_premises__rerun.E_input', cc4_relset_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t155_funct_2__9_premises__rerun.E_input', d1_xboole_0)).
fof(c_0_10, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_partfun1(X3,X4)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|r2_hidden(X4,k5_partfun1(X1,X2,X3)))))))), inference(assume_negation,[status(cth)],[t155_funct_2])).
fof(c_0_11, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X11]:![X13]:![X10]:![X15]:(((((((v1_funct_1(esk6_5(X7,X8,X9,X10,X11))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((m1_subset_1(esk6_5(X7,X8,X9,X10,X11),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((esk6_5(X7,X8,X9,X10,X11)=X11|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&((((v1_partfun1(esk6_5(X7,X8,X9,X10,X11),X7)|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((r1_partfun1(X9,esk6_5(X7,X8,X9,X10,X11))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))))&(((((~v1_funct_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))|(X13!=X11|(~v1_partfun1(X13,X7)|~r1_partfun1(X9,X13))))|r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((~r2_hidden(esk7_4(X7,X8,X9,X10),X10)|((~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))|(X15!=esk7_4(X7,X8,X9,X10)|(~v1_partfun1(X15,X7)|~r1_partfun1(X9,X15)))))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((((v1_funct_1(esk8_4(X7,X8,X9,X10))|r2_hidden(esk7_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((m1_subset_1(esk8_4(X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|r2_hidden(esk7_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((esk8_4(X7,X8,X9,X10)=esk7_4(X7,X8,X9,X10)|r2_hidden(esk7_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&((((v1_partfun1(esk8_4(X7,X8,X9,X10),X7)|r2_hidden(esk7_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((r1_partfun1(X9,esk8_4(X7,X8,X9,X10))|r2_hidden(esk7_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_partfun1])])])])])])])).
fof(c_0_12, plain, (![X3]:![X4]:![X3]:![X4]:((k2_zfmisc_1(X3,X4)!=k1_xboole_0|(X3=k1_xboole_0|X4=k1_xboole_0))&((X3!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)&(X4!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])])])).
fof(c_0_13, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
fof(c_0_14, plain, (v1_xboole_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_xboole_0])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:((((v1_funct_1(X6)|(~v1_funct_1(X6)|~v1_funct_2(X6,X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))|v1_xboole_0(X5))&(((v1_partfun1(X6,X4)|(~v1_funct_1(X6)|~v1_funct_2(X6,X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))|v1_xboole_0(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])])])).
fof(c_0_16, negated_conjecture, (((v1_funct_1(esk4_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0))))&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,esk2_0,esk3_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))))&(r1_partfun1(esk4_0,esk5_0)&((esk3_0!=k1_xboole_0|esk2_0=k1_xboole_0)&~r2_hidden(esk5_0,k5_partfun1(esk2_0,esk3_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])).
cnf(c_0_17,plain,(r2_hidden(X5,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r1_partfun1(X1,X6)|~v1_partfun1(X6,X2)|X6!=X5|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X6)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_18, plain, (![X2]:![X2]:(v1_partfun1(k6_partfun1(X2),X2)&m1_subset_1(k6_partfun1(X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k6_partfun1])])])).
cnf(c_0_19,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,plain,(v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,plain,(v1_xboole_0(X1)|v1_partfun1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X2,X3,X1)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,negated_conjecture,(v1_funct_2(esk5_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,negated_conjecture,(v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26,negated_conjecture,(~r2_hidden(esk5_0,k5_partfun1(esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27,plain,(r2_hidden(X1,k5_partfun1(X2,X3,X4))|~r1_partfun1(X4,X1)|~v1_funct_1(X1)|~v1_funct_1(X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_partfun1(X1,X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_17])])).
cnf(c_0_28,negated_conjecture,(r1_partfun1(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_31, plain, (![X4]:![X5]:![X6]:(~r2_hidden(X4,X5)|(~m1_subset_1(X5,k1_zfmisc_1(X6))|~v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])])])).
cnf(c_0_32,plain,(m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_33,plain,(k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_34,plain,(esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_35,negated_conjecture,(v1_partfun1(esk5_0,esk2_0)|v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_36,negated_conjecture,(~v1_partfun1(esk5_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_24]), c_0_29]), c_0_25]), c_0_30])])).
fof(c_0_37, plain, (![X4]:![X5]:![X6]:(~v1_xboole_0(X4)|(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X4)))|v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])])])).
cnf(c_0_38,plain,(~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39,plain,(m1_subset_1(k6_partfun1(k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40,plain,(v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_21, c_0_34])).
fof(c_0_41, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk9_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_42,negated_conjecture,(v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43,plain,(v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44,plain,(~r2_hidden(X1,k6_partfun1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
cnf(c_0_45,plain,(v1_xboole_0(X1)|r2_hidden(esk9_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46,negated_conjecture,(esk2_0=k1_xboole_0|esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_47,negated_conjecture,(esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_42])).
cnf(c_0_48,negated_conjecture,(v1_xboole_0(esk5_0)|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_43, c_0_25])).
cnf(c_0_49,plain,(v1_xboole_0(k6_partfun1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50,negated_conjecture,(esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47])])).
cnf(c_0_51,negated_conjecture,(v1_xboole_0(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_42])])).
cnf(c_0_52,plain,(v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_53,plain,(k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_49])).
cnf(c_0_54,negated_conjecture,(~v1_partfun1(esk5_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_36, c_0_50])).
cnf(c_0_55,negated_conjecture,(esk5_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_51])).
cnf(c_0_56,plain,(v1_partfun1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_57,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55]), c_0_56])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 37
# Proof object formula steps           : 21
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 98
# ...of these trivial                  : 2
# ...subsumed                          : 2
# ...remaining for further processing  : 94
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 24
# Generated clauses                    : 43
# ...of the previous two non-trivial   : 56
# Contextual simplify-reflections      : 0
# Paramodulations                      : 33
# Factorizations                       : 0
# Equation resolutions                 : 10
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 71
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 700
# Rec. Clause-clause subsumption calls : 195
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 106
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3035

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45296 pages
