# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d7_partfun1, axiom, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:(X4=k5_partfun1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>?[X6]:((v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(X6=X5&(v1_partfun1(X6,X1)&r1_partfun1(X3,X6)))))))), file('out/t158_funct_2__2_premises__rerun.E_input', d7_partfun1)).
fof(t158_funct_2, conjecture, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:(r2_hidden(X4,k5_partfun1(X1,X2,X3))=>(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))))), file('out/t158_funct_2__2_premises__rerun.E_input', t158_funct_2)).
fof(cc1_funct_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_partfun1(X3,X1)=>v1_funct_2(X3,X1,X2)))), file('out/t158_funct_2__2_premises__rerun.E_input', cc1_funct_2)).
fof(c_0_3, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X11]:![X13]:![X10]:![X15]:(((((((v1_funct_1(esk1_5(X7,X8,X9,X10,X11))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((m1_subset_1(esk1_5(X7,X8,X9,X10,X11),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((esk1_5(X7,X8,X9,X10,X11)=X11|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&((((v1_partfun1(esk1_5(X7,X8,X9,X10,X11),X7)|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((r1_partfun1(X9,esk1_5(X7,X8,X9,X10,X11))|~r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))))&(((((~v1_funct_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))|(X13!=X11|(~v1_partfun1(X13,X7)|~r1_partfun1(X9,X13))))|r2_hidden(X11,X10))|X10!=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((~r2_hidden(esk2_4(X7,X8,X9,X10),X10)|((~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))|(X15!=esk2_4(X7,X8,X9,X10)|(~v1_partfun1(X15,X7)|~r1_partfun1(X9,X15)))))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((((v1_funct_1(esk3_4(X7,X8,X9,X10))|r2_hidden(esk2_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((m1_subset_1(esk3_4(X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|r2_hidden(esk2_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((((esk3_4(X7,X8,X9,X10)=esk2_4(X7,X8,X9,X10)|r2_hidden(esk2_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&((((v1_partfun1(esk3_4(X7,X8,X9,X10),X7)|r2_hidden(esk2_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(((r1_partfun1(X9,esk3_4(X7,X8,X9,X10))|r2_hidden(esk2_4(X7,X8,X9,X10),X10))|X10=k5_partfun1(X7,X8,X9))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_partfun1])])])])])])])).
fof(c_0_4, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:(r2_hidden(X4,k5_partfun1(X1,X2,X3))=>(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))))), inference(assume_negation,[status(cth)],[t158_funct_2])).
cnf(c_0_5,plain,(v1_funct_1(esk1_5(X2,X3,X1,X4,X5))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r2_hidden(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,plain,(esk1_5(X2,X3,X1,X4,X5)=X5|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r2_hidden(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_7, negated_conjecture, (((v1_funct_1(esk6_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0))))&(r2_hidden(esk7_0,k5_partfun1(esk4_0,esk5_0,esk6_0))&(~v1_funct_1(esk7_0)|(~v1_funct_2(esk7_0,esk4_0,esk5_0)|~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
fof(c_0_8, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|(~v1_partfun1(X6,X4)|v1_funct_2(X6,X4,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])).
cnf(c_0_9,plain,(v1_funct_1(esk1_5(X1,X2,X3,k5_partfun1(X1,X2,X3),X4))|~r2_hidden(X4,k5_partfun1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(er,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(esk1_5(X1,X2,X3,k5_partfun1(X1,X2,X3),X4)=X4|~r2_hidden(X4,k5_partfun1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(v1_partfun1(esk1_5(X2,X3,X1,X4,X5),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r2_hidden(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_12,negated_conjecture,(~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))|~v1_funct_2(esk7_0,esk4_0,esk5_0)|~v1_funct_1(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(v1_funct_1(X1)|~r2_hidden(X1,k5_partfun1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_15,negated_conjecture,(r2_hidden(esk7_0,k5_partfun1(esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_18,plain,(v1_partfun1(esk1_5(X1,X2,X3,k5_partfun1(X1,X2,X3),X4),X1)|~r2_hidden(X4,k5_partfun1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,(m1_subset_1(esk1_5(X2,X3,X1,X4,X5),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|X4!=k5_partfun1(X2,X3,X1)|~r2_hidden(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_20,negated_conjecture,(~v1_partfun1(esk7_0,esk4_0)|~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))|~v1_funct_1(esk7_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_21,negated_conjecture,(v1_funct_1(esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_22,plain,(v1_partfun1(X1,X2)|~r2_hidden(X1,k5_partfun1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_18, c_0_10])).
cnf(c_0_23,plain,(m1_subset_1(esk1_5(X1,X2,X3,k5_partfun1(X1,X2,X3),X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_hidden(X4,k5_partfun1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_24,negated_conjecture,(~v1_partfun1(esk7_0,esk4_0)|~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21])])).
cnf(c_0_25,negated_conjecture,(v1_partfun1(esk7_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_26,plain,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X1,k5_partfun1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_23, c_0_10])).
cnf(c_0_27,negated_conjecture,(~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25])])).
cnf(c_0_28,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_15]), c_0_16]), c_0_17])]), c_0_27]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 22
# Proof object formula steps           : 7
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 54
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 39
# ...of the previous two non-trivial   : 34
# Contextual simplify-reflections      : 0
# Paramodulations                      : 32
# Factorizations                       : 0
# Equation resolutions                 : 8
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 19
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 84
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 220
# Rec. Clause-clause subsumption calls : 36
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 29
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2775

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45332 pages
