# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_zfmisc_1, axiom, (![X1]:![X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), file('out/t159_funct_2__18_premises__rerun.E_input', t113_zfmisc_1)).
fof(t194_relat_1, axiom, (![X1]:![X2]:r1_tarski(k2_relat_1(k2_zfmisc_1(X1,X2)),X2)), file('out/t159_funct_2__18_premises__rerun.E_input', t194_relat_1)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t159_funct_2__18_premises__rerun.E_input', t3_subset)).
fof(t5_subset, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3))))), file('out/t159_funct_2__18_premises__rerun.E_input', t5_subset)).
fof(t159_funct_2, conjecture, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>r1_tarski(k5_partfun1(X1,X2,X3),k1_funct_2(X1,X2)))), file('out/t159_funct_2__18_premises__rerun.E_input', t159_funct_2)).
fof(t4_subset_1, axiom, (![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), file('out/t159_funct_2__18_premises__rerun.E_input', t4_subset_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t159_funct_2__18_premises__rerun.E_input', d1_xboole_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t159_funct_2__18_premises__rerun.E_input', d10_xboole_0)).
fof(cc4_relset_1, axiom, (![X1]:![X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3)))), file('out/t159_funct_2__18_premises__rerun.E_input', cc4_relset_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t159_funct_2__18_premises__rerun.E_input', d3_tarski)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t159_funct_2__18_premises__rerun.E_input', fc1_xboole_0)).
fof(t4_relset_1, axiom, (![X1]:![X2]:![X3]:![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>(r1_tarski(X1,X4)=>m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))), file('out/t159_funct_2__18_premises__rerun.E_input', t4_relset_1)).
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t159_funct_2__18_premises__rerun.E_input', t1_xboole_1)).
fof(dt_k1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k1_relset_1(X1,X2,X3),k1_zfmisc_1(X1)))), file('out/t159_funct_2__18_premises__rerun.E_input', dt_k1_relset_1)).
fof(redefinition_k1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3))), file('out/t159_funct_2__18_premises__rerun.E_input', redefinition_k1_relset_1)).
fof(fc2_partfun1, axiom, (![X1]:![X2]:![X3]:((~(v1_xboole_0(X1))&(v1_xboole_0(X2)&(v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>v1_xboole_0(k5_partfun1(X1,X2,X3)))), file('out/t159_funct_2__18_premises__rerun.E_input', fc2_partfun1)).
fof(t11_funct_2, axiom, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((X2=k1_xboole_0=>X1=k1_xboole_0)=>r2_hidden(X3,k1_funct_2(X1,X2))))), file('out/t159_funct_2__18_premises__rerun.E_input', t11_funct_2)).
fof(t158_funct_2, axiom, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:(r2_hidden(X4,k5_partfun1(X1,X2,X3))=>(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))))), file('out/t159_funct_2__18_premises__rerun.E_input', t158_funct_2)).
fof(t12_funct_2, axiom, (![X1]:![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>r2_hidden(X2,k1_funct_2(X1,X1)))), file('out/t159_funct_2__18_premises__rerun.E_input', t12_funct_2)).
fof(c_0_19, plain, (![X3]:![X4]:![X3]:![X4]:((k2_zfmisc_1(X3,X4)!=k1_xboole_0|(X3=k1_xboole_0|X4=k1_xboole_0))&((X3!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)&(X4!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])])])).
fof(c_0_20, plain, (![X3]:![X4]:r1_tarski(k2_relat_1(k2_zfmisc_1(X3,X4)),X4)), inference(variable_rename,[status(thm)],[t194_relat_1])).
cnf(c_0_21,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_23,plain,(r1_tarski(k2_relat_1(k2_zfmisc_1(X1,X2)),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_21])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:(~r2_hidden(X4,X5)|(~m1_subset_1(X5,k1_zfmisc_1(X6))|~v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])])])).
cnf(c_0_26,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27,plain,(r1_tarski(k2_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_28, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>r1_tarski(k5_partfun1(X1,X2,X3),k1_funct_2(X1,X2))))), inference(assume_negation,[status(cth)],[t159_funct_2])).
fof(c_0_29, plain, (![X2]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_30,plain,(~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31,plain,(m1_subset_1(k2_relat_1(k1_xboole_0),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_32, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk1_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
fof(c_0_33, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:(~v1_xboole_0(X4)|(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X4)))|v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])])])).
fof(c_0_35, negated_conjecture, (((v1_funct_1(esk4_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0))))&~r1_tarski(k5_partfun1(esk2_0,esk3_0,esk4_0),k1_funct_2(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])).
cnf(c_0_36,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_37,plain,(m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_39,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k2_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40,plain,(v1_xboole_0(X1)|r2_hidden(esk1_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42,plain,(v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44,plain,(k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_45,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_46,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_30, c_0_37])).
cnf(c_0_47,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48,plain,(v1_xboole_0(k2_relat_1(k1_xboole_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_50,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_51,negated_conjecture,(v1_xboole_0(esk4_0)|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52,plain,(v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_44]), c_0_45])])).
fof(c_0_53, plain, (![X5]:![X6]:![X7]:![X8]:(~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|(~r1_tarski(X5,X8)|m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_relset_1])])).
cnf(c_0_54,plain,(r1_tarski(k1_xboole_0,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55,plain,(v1_xboole_0(k2_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_48, c_0_45])).
cnf(c_0_56,plain,(m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_26, c_0_49])).
cnf(c_0_57,plain,(r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_50, c_0_47])).
cnf(c_0_58,negated_conjecture,(v1_xboole_0(esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59,plain,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_61,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_62,plain,(v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_42, c_0_56])).
cnf(c_0_63,negated_conjecture,(r1_tarski(esk4_0,X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_64, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_65, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|m1_subset_1(k1_relset_1(X4,X5,X6),k1_zfmisc_1(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
fof(c_0_66, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|k1_relset_1(X4,X5,X6)=k1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_67,negated_conjecture,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~r1_tarski(X1,esk4_0)), inference(spm,[status(thm)],[c_0_59, c_0_43])).
cnf(c_0_68,plain,(X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_69,plain,(r1_tarski(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_57, c_0_62])).
cnf(c_0_70,negated_conjecture,(X1=esk4_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,esk4_0)), inference(spm,[status(thm)],[c_0_60, c_0_63])).
cnf(c_0_71,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_73,plain,(m1_subset_1(k1_relset_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74,plain,(k1_relset_1(X1,X2,X3)=k1_relat_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_75,negated_conjecture,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~r1_tarski(X2,esk4_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_59, c_0_67])).
cnf(c_0_76,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_77, plain, (![X4]:![X5]:![X6]:((v1_xboole_0(X4)|(~v1_xboole_0(X5)|(~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))|v1_xboole_0(k5_partfun1(X4,X5,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_partfun1])])])).
cnf(c_0_78,negated_conjecture,(k2_zfmisc_1(X1,X2)=esk4_0|~v1_xboole_0(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_70, c_0_69])).
cnf(c_0_79,plain,(r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_80,plain,(m1_subset_1(k1_relat_1(X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_81,plain,(X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_60, c_0_72])).
cnf(c_0_82,negated_conjecture,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(X2,k1_zfmisc_1(esk4_0))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_75, c_0_72])).
cnf(c_0_83,negated_conjecture,(~v1_xboole_0(k2_zfmisc_1(esk2_0,esk3_0))|~r2_hidden(X1,X2)|~r1_tarski(X2,esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_67])).
cnf(c_0_84,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_76, c_0_52])).
cnf(c_0_85,plain,(v1_xboole_0(k5_partfun1(X1,X2,X3))|v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_86,negated_conjecture,(k2_zfmisc_1(X1,esk4_0)=esk4_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_78, c_0_58])).
cnf(c_0_87,negated_conjecture,(r1_tarski(X1,k2_zfmisc_1(esk2_0,esk3_0))|~r1_tarski(X1,esk4_0)), inference(spm,[status(thm)],[c_0_79, c_0_43])).
cnf(c_0_88,negated_conjecture,(m1_subset_1(k1_relat_1(esk4_0),k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_80, c_0_43])).
cnf(c_0_89,plain,(k2_zfmisc_1(X1,X2)=X3|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_81, c_0_69])).
cnf(c_0_90,negated_conjecture,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk4_0))), inference(spm,[status(thm)],[c_0_82, c_0_49])).
cnf(c_0_91,negated_conjecture,(~r2_hidden(X1,X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X2,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_45])])).
cnf(c_0_92,plain,(r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_71, c_0_61])).
cnf(c_0_93,negated_conjecture,(~r1_tarski(k5_partfun1(esk2_0,esk3_0,esk4_0),k1_funct_2(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_94,plain,(v1_xboole_0(X1)|r1_tarski(k5_partfun1(X1,X2,X3),X4)|~v1_funct_1(X3)|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_57, c_0_85])).
cnf(c_0_95,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_96,negated_conjecture,(r1_tarski(k2_relat_1(esk4_0),esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_23, c_0_86])).
cnf(c_0_97,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=X1|~r1_tarski(k2_zfmisc_1(esk2_0,esk3_0),X1)|~r1_tarski(X1,esk4_0)), inference(spm,[status(thm)],[c_0_60, c_0_87])).
cnf(c_0_98,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k1_relset_1(X1,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(spm,[status(thm)],[c_0_30, c_0_73])).
cnf(c_0_99,negated_conjecture,(r1_tarski(X1,esk2_0)|~r1_tarski(X1,k1_relat_1(esk4_0))), inference(spm,[status(thm)],[c_0_79, c_0_88])).
cnf(c_0_100,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=X1|~v1_xboole_0(esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(esk4_0))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_101,negated_conjecture,(v1_xboole_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,esk4_0)), inference(spm,[status(thm)],[c_0_91, c_0_40])).
cnf(c_0_102,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_92, c_0_72])).
cnf(c_0_103,plain,(m1_subset_1(k2_zfmisc_1(X1,X2),k1_zfmisc_1(X3))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_26, c_0_69])).
cnf(c_0_104,negated_conjecture,(v1_xboole_0(esk2_0)|~v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), c_0_43])])).
cnf(c_0_105,negated_conjecture,(m1_subset_1(k2_relat_1(esk4_0),k1_zfmisc_1(esk4_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_26, c_0_96])).
cnf(c_0_106,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=X1|~v1_xboole_0(esk3_0)|~r1_tarski(X1,esk4_0)), inference(spm,[status(thm)],[c_0_97, c_0_69])).
cnf(c_0_107,negated_conjecture,(r1_tarski(X1,X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,esk4_0)), inference(spm,[status(thm)],[c_0_71, c_0_63])).
cnf(c_0_108,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k1_relat_1(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X4)))), inference(spm,[status(thm)],[c_0_98, c_0_74])).
cnf(c_0_109,negated_conjecture,(r1_tarski(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_relat_1(esk4_0)))), inference(spm,[status(thm)],[c_0_99, c_0_72])).
cnf(c_0_110,plain,(m1_subset_1(k1_relat_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_80, c_0_56])).
cnf(c_0_111,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=X1|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(esk4_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102])).
cnf(c_0_112,negated_conjecture,(m1_subset_1(k2_zfmisc_1(X1,esk4_0),k1_zfmisc_1(X2))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_103, c_0_58])).
cnf(c_0_113,negated_conjecture,(v1_xboole_0(esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_104, c_0_52])).
cnf(c_0_114,negated_conjecture,(~r2_hidden(X1,k2_relat_1(esk4_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_105]), c_0_58])).
cnf(c_0_115,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=X1|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_101]), c_0_102])).
cnf(c_0_116,negated_conjecture,(r1_tarski(k2_relat_1(esk4_0),X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_107, c_0_96])).
cnf(c_0_117,plain,(X1=X2|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_60, c_0_102])).
cnf(c_0_118,negated_conjecture,(r1_tarski(X1,X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(esk4_0))), inference(spm,[status(thm)],[c_0_107, c_0_72])).
cnf(c_0_119,negated_conjecture,(~v1_xboole_0(esk2_0)|~r2_hidden(X1,k1_relat_1(esk4_0))), inference(spm,[status(thm)],[c_0_108, c_0_43])).
cnf(c_0_120,negated_conjecture,(r1_tarski(k1_relat_1(k2_zfmisc_1(k1_relat_1(esk4_0),X1)),esk2_0)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_121,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=k2_zfmisc_1(X1,esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_122,negated_conjecture,(r1_tarski(esk2_0,X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_57, c_0_113])).
cnf(c_0_123,negated_conjecture,(v1_xboole_0(k2_relat_1(esk4_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_114, c_0_40])).
cnf(c_0_124,negated_conjecture,(k2_relat_1(esk4_0)=k2_zfmisc_1(esk2_0,esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_125,plain,(k2_relat_1(k2_zfmisc_1(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_117, c_0_23])).
cnf(c_0_126,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=esk4_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_115, c_0_63])).
cnf(c_0_127,negated_conjecture,(X1=X2|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X2,k1_zfmisc_1(esk4_0))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_60, c_0_118])).
cnf(c_0_128,negated_conjecture,(r1_tarski(k1_relat_1(esk4_0),X1)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_119, c_0_47])).
cnf(c_0_129,negated_conjecture,(m1_subset_1(X1,k1_zfmisc_1(esk4_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X2,k1_zfmisc_1(esk4_0))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_59, c_0_86])).
cnf(c_0_130,negated_conjecture,(r1_tarski(k1_relat_1(k2_zfmisc_1(esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_131,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(X1))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_26, c_0_122])).
cnf(c_0_132,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=k2_zfmisc_1(X1,X2)|~v1_xboole_0(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_115, c_0_69])).
cnf(c_0_133,negated_conjecture,(v1_xboole_0(k2_zfmisc_1(esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_134,negated_conjecture,(k2_relat_1(esk4_0)=esk3_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_135,negated_conjecture,(k1_relat_1(esk4_0)=X1|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(esk4_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_113])).
cnf(c_0_136,negated_conjecture,(m1_subset_1(k1_relat_1(k2_zfmisc_1(esk2_0,esk3_0)),k1_zfmisc_1(esk4_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_131])).
cnf(c_0_137,negated_conjecture,(k2_zfmisc_1(X1,k2_zfmisc_1(esk2_0,esk3_0))=k2_zfmisc_1(esk2_0,esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_138,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=esk3_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_124, c_0_134])).
cnf(c_0_139,plain,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_37])).
cnf(c_0_140,negated_conjecture,(k1_relat_1(k2_zfmisc_1(esk2_0,esk3_0))=k1_relat_1(esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_141,negated_conjecture,(k2_zfmisc_1(X1,esk3_0)=esk3_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_142,plain,(~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,X4)|~r1_tarski(X4,k1_xboole_0)), inference(spm,[status(thm)],[c_0_30, c_0_139])).
cnf(c_0_143,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k1_relat_1(X3))|~r1_tarski(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_108, c_0_139])).
cnf(c_0_144,negated_conjecture,(k1_relat_1(esk3_0)=k1_relat_1(esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_145,plain,(~r2_hidden(X1,X2)|~r1_tarski(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_24]), c_0_45])])).
cnf(c_0_146,negated_conjecture,(~v1_xboole_0(X1)|~r2_hidden(X2,k1_relat_1(esk4_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_26])).
cnf(c_0_147,plain,(v1_xboole_0(X1)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_145, c_0_40])).
cnf(c_0_148,negated_conjecture,(v1_xboole_0(k1_relat_1(esk4_0))|~v1_xboole_0(X1)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_146, c_0_40])).
cnf(c_0_149,negated_conjecture,(v1_xboole_0(esk4_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_147])).
cnf(c_0_150,negated_conjecture,(v1_xboole_0(k1_relat_1(esk4_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_148, c_0_149])).
cnf(c_0_151,negated_conjecture,(r1_tarski(k1_relat_1(esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_99, c_0_49])).
cnf(c_0_152,negated_conjecture,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~r1_tarski(X1,k2_relat_1(k2_zfmisc_1(X2,esk4_0)))), inference(spm,[status(thm)],[c_0_75, c_0_23])).
cnf(c_0_153,negated_conjecture,(r1_tarski(k1_relat_1(esk4_0),X1)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_57, c_0_150])).
cnf(c_0_154,negated_conjecture,(k1_relat_1(esk4_0)=esk2_0|~r1_tarski(esk2_0,k1_relat_1(esk4_0))), inference(spm,[status(thm)],[c_0_60, c_0_151])).
cnf(c_0_155,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k1_relat_1(k2_zfmisc_1(X1,X3)))), inference(spm,[status(thm)],[c_0_108, c_0_56])).
cnf(c_0_156,negated_conjecture,(k2_zfmisc_1(esk2_0,esk3_0)=esk2_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_115, c_0_122])).
cnf(c_0_157,plain,(r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_71, c_0_102])).
cnf(c_0_158,negated_conjecture,(m1_subset_1(k1_relat_1(esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_152, c_0_153])).
cnf(c_0_159,negated_conjecture,(k1_relat_1(esk4_0)=esk2_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_154, c_0_122])).
cnf(c_0_160,negated_conjecture,(~r2_hidden(X1,k1_relat_1(esk2_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_113])).
cnf(c_0_161,plain,(r1_tarski(k2_relat_1(k2_zfmisc_1(X1,X2)),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_157, c_0_23])).
cnf(c_0_162,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_26])).
cnf(c_0_163,negated_conjecture,(r1_tarski(k1_relat_1(esk2_0),X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_160, c_0_47])).
cnf(c_0_164,negated_conjecture,(v1_xboole_0(k1_relat_1(esk4_0))|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_119, c_0_40])).
cnf(c_0_165,negated_conjecture,(r1_tarski(k2_relat_1(esk3_0),X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_161, c_0_138])).
cnf(c_0_166,negated_conjecture,(m1_subset_1(k1_relat_1(esk2_0),k1_zfmisc_1(esk2_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_80, c_0_162])).
cnf(c_0_167,negated_conjecture,(k1_relat_1(esk2_0)=k2_zfmisc_1(esk2_0,esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_115, c_0_163])).
cnf(c_0_168,negated_conjecture,(k1_relat_1(esk4_0)=X1|~v1_xboole_0(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_relat_1(esk4_0)))), inference(spm,[status(thm)],[c_0_81, c_0_128])).
cnf(c_0_169,negated_conjecture,(m1_subset_1(k2_zfmisc_1(X1,k1_relat_1(esk4_0)),k1_zfmisc_1(X2))|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_103, c_0_164])).
cnf(c_0_170,negated_conjecture,(v1_xboole_0(esk2_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_104, c_0_147])).
cnf(c_0_171,negated_conjecture,(X1=k2_relat_1(esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,k2_relat_1(esk4_0))), inference(spm,[status(thm)],[c_0_60, c_0_116])).
cnf(c_0_172,negated_conjecture,(X1=k2_relat_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_60, c_0_165])).
cnf(c_0_173,negated_conjecture,(m1_subset_1(k2_zfmisc_1(esk2_0,esk3_0),k1_zfmisc_1(esk2_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_26])).
cnf(c_0_174,negated_conjecture,(k2_zfmisc_1(X1,k1_relat_1(esk4_0))=k1_relat_1(esk4_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_168, c_0_169])).
cnf(c_0_175,negated_conjecture,(k1_relat_1(esk4_0)=k2_zfmisc_1(esk2_0,esk3_0)|~v1_xboole_0(esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_158]), c_0_57])).
cnf(c_0_176,negated_conjecture,(~r2_hidden(X1,k1_relat_1(esk2_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_162]), c_0_170])).
cnf(c_0_177,negated_conjecture,(k1_relat_1(esk2_0)=k2_relat_1(esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_171, c_0_163])).
cnf(c_0_178,negated_conjecture,(k1_relat_1(esk2_0)=k2_relat_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_172, c_0_163])).
cnf(c_0_179,negated_conjecture,(~r2_hidden(X1,k2_zfmisc_1(esk2_0,esk3_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_173]), c_0_170])).
cnf(c_0_180,negated_conjecture,(k2_zfmisc_1(X1,k2_zfmisc_1(esk2_0,esk3_0))=k2_zfmisc_1(esk2_0,esk3_0)|~v1_xboole_0(esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175]), c_0_104])).
cnf(c_0_181,negated_conjecture,(~r2_hidden(X1,k2_relat_1(esk4_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_26])).
cnf(c_0_182,negated_conjecture,(~r2_hidden(X1,k2_relat_1(esk3_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_178]), c_0_26])).
cnf(c_0_183,negated_conjecture,(r1_tarski(k2_zfmisc_1(esk2_0,esk3_0),X1)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_179, c_0_47])).
cnf(c_0_184,negated_conjecture,(k2_zfmisc_1(X1,k2_zfmisc_1(esk2_0,esk3_0))=k2_zfmisc_1(esk2_0,esk3_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_180, c_0_147])).
fof(c_0_185, plain, (![X4]:![X5]:![X6]:(((X5=k1_xboole_0|r2_hidden(X6,k1_funct_2(X4,X5)))|(~v1_funct_1(X6)|(~v1_funct_2(X6,X4,X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))&((X4!=k1_xboole_0|r2_hidden(X6,k1_funct_2(X4,X5)))|(~v1_funct_1(X6)|(~v1_funct_2(X6,X4,X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_funct_2])])])).
fof(c_0_186, plain, (![X5]:![X6]:![X7]:![X8]:(((v1_funct_1(X8)|~r2_hidden(X8,k5_partfun1(X5,X6,X7)))|(~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))&(((v1_funct_2(X8,X5,X6)|~r2_hidden(X8,k5_partfun1(X5,X6,X7)))|(~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))&((m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~r2_hidden(X8,k5_partfun1(X5,X6,X7)))|(~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t158_funct_2])])])])])])).
cnf(c_0_187,negated_conjecture,(X1=k1_relat_1(esk4_0)|~v1_xboole_0(esk2_0)|~r1_tarski(X1,k1_relat_1(esk4_0))), inference(spm,[status(thm)],[c_0_60, c_0_128])).
cnf(c_0_188,negated_conjecture,(r1_tarski(k2_relat_1(esk4_0),X1)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_181, c_0_47])).
cnf(c_0_189,negated_conjecture,(r1_tarski(k2_relat_1(esk3_0),X1)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_182, c_0_47])).
cnf(c_0_190,negated_conjecture,(X1=k2_zfmisc_1(esk2_0,esk3_0)|~r1_tarski(X1,k2_zfmisc_1(esk2_0,esk3_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_60, c_0_183])).
cnf(c_0_191,negated_conjecture,(r1_tarski(k2_relat_1(k2_zfmisc_1(esk2_0,esk3_0)),k2_zfmisc_1(esk2_0,esk3_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_23, c_0_184])).
cnf(c_0_192,plain,(r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_193,plain,(r2_hidden(X1,k1_funct_2(X2,X3))|X3=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_194,plain,(v1_funct_2(X4,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~r2_hidden(X4,k5_partfun1(X2,X3,X1))), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_195,plain,(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~r2_hidden(X4,k5_partfun1(X2,X3,X1))), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_196,plain,(v1_funct_1(X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~r2_hidden(X4,k5_partfun1(X2,X3,X1))), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_197,negated_conjecture,(k1_relat_1(esk4_0)=k2_relat_1(esk4_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_170])).
cnf(c_0_198,negated_conjecture,(k1_relat_1(esk4_0)=k2_relat_1(esk3_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_189]), c_0_170])).
cnf(c_0_199,negated_conjecture,(k2_relat_1(k2_zfmisc_1(esk2_0,esk3_0))=k2_zfmisc_1(esk2_0,esk3_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_190, c_0_191])).
cnf(c_0_200,plain,(X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_81, c_0_72])).
cnf(c_0_201,negated_conjecture,(k2_zfmisc_1(X1,X2)=esk4_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_78, c_0_52])).
cnf(c_0_202,plain,(X1=k1_xboole_0|r1_tarski(X2,k1_funct_2(X3,X1))|~v1_funct_2(esk5_2(X2,k1_funct_2(X3,X1)),X3,X1)|~v1_funct_1(esk5_2(X2,k1_funct_2(X3,X1)))|~m1_subset_1(esk5_2(X2,k1_funct_2(X3,X1)),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(spm,[status(thm)],[c_0_192, c_0_193])).
cnf(c_0_203,plain,(v1_funct_2(esk5_2(k5_partfun1(X1,X2,X3),X4),X1,X2)|r1_tarski(k5_partfun1(X1,X2,X3),X4)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_194, c_0_47])).
cnf(c_0_204,plain,(m1_subset_1(esk5_2(k5_partfun1(X1,X2,X3),X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_tarski(k5_partfun1(X1,X2,X3),X4)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_195, c_0_47])).
cnf(c_0_205,plain,(v1_funct_1(esk5_2(k5_partfun1(X1,X2,X3),X4))|r1_tarski(k5_partfun1(X1,X2,X3),X4)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_196, c_0_47])).
cnf(c_0_206,negated_conjecture,(k2_relat_1(esk3_0)=k2_relat_1(esk4_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_197, c_0_198])).
cnf(c_0_207,negated_conjecture,(k2_relat_1(esk3_0)=esk3_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_141]), c_0_26])).
cnf(c_0_208,negated_conjecture,(k1_relat_1(esk2_0)=esk2_0|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_relat_1(esk2_0)))|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_200, c_0_166])).
cnf(c_0_209,negated_conjecture,(k1_relat_1(esk2_0)=esk4_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_70, c_0_163])).
cnf(c_0_210,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk4_0))|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_201]), c_0_26])).
cnf(c_0_211,plain,(X1=k1_xboole_0|r1_tarski(k5_partfun1(X2,X1,X3),k1_funct_2(X2,X1))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_204]), c_0_205])).
cnf(c_0_212,negated_conjecture,(k2_relat_1(esk4_0)=esk4_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_201]), c_0_26])).
cnf(c_0_213,negated_conjecture,(k2_relat_1(esk4_0)=esk3_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_206, c_0_207])).
fof(c_0_214, plain, (![X3]:![X4]:((~v1_funct_1(X4)|(~v1_funct_2(X4,X3,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))))|r2_hidden(X4,k1_funct_2(X3,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_2])])).
cnf(c_0_215,negated_conjecture,(esk2_0=esk4_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209]), c_0_26]), c_0_210])).
cnf(c_0_216,negated_conjecture,(esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_211]), c_0_95]), c_0_43])])).
cnf(c_0_217,negated_conjecture,(esk3_0=esk4_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_212, c_0_213])).
cnf(c_0_218,plain,(r2_hidden(X1,k1_funct_2(X2,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_2(X1,X2,X2)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_219,negated_conjecture,(esk2_0=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_216]), c_0_61])])).
cnf(c_0_220,negated_conjecture,(esk4_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_217, c_0_216]), c_0_216]), c_0_61])])).
cnf(c_0_221,plain,(r1_tarski(X1,k1_funct_2(X2,X2))|~v1_funct_2(esk5_2(X1,k1_funct_2(X2,X2)),X2,X2)|~v1_funct_1(esk5_2(X1,k1_funct_2(X2,X2)))|~m1_subset_1(esk5_2(X1,k1_funct_2(X2,X2)),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(spm,[status(thm)],[c_0_192, c_0_218])).
cnf(c_0_222,negated_conjecture,(~r1_tarski(k5_partfun1(k1_xboole_0,k1_xboole_0,k1_xboole_0),k1_funct_2(k1_xboole_0,k1_xboole_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_216]), c_0_216]), c_0_219]), c_0_220]), c_0_220]), c_0_219]), c_0_220])).
cnf(c_0_223,plain,(r1_tarski(k5_partfun1(X1,X1,X2),k1_funct_2(X1,X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_203]), c_0_204]), c_0_205])).
cnf(c_0_224,negated_conjecture,(v1_funct_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_95, c_0_220])).
cnf(c_0_225,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223]), c_0_224]), c_0_24]), c_0_37])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 226
# Proof object clause steps            : 188
# Proof object formula steps           : 38
# Proof object conjectures             : 118
# Proof object clause conjectures      : 115
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 153
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 13939
# ...of these trivial                  : 151
# ...subsumed                          : 12429
# ...remaining for further processing  : 1359
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 177
# Backward-rewritten                   : 831
# Generated clauses                    : 61594
# ...of the previous two non-trivial   : 56485
# Contextual simplify-reflections      : 9254
# Paramodulations                      : 61589
# Factorizations                       : 0
# Equation resolutions                 : 5
# Current number of processed clauses  : 315
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 294
# Current number of unprocessed clauses: 4394
# ...number of literals in the above   : 18401
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1039
# Clause-clause subsumption calls (NU) : 1372063
# Rec. Clause-clause subsumption calls : 1050799
# Non-unit clause-clause subsumptions  : 20477
# Unit Clause-clause subsumption calls : 984
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 21
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 738565

# -------------------------------------------------
# User time                : 1.024 s
# System time              : 0.024 s
# Total time               : 1.048 s
# Maximum resident set size: 45476 pages
