# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_orders_2, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3]:![X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4)))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', free_g1_orders_2)).
fof(dt_k1_yellow_1, axiom, (![X1]:(v1_relat_2(k1_yellow_1(X1))&(v4_relat_2(k1_yellow_1(X1))&(v8_relat_2(k1_yellow_1(X1))&(v1_partfun1(k1_yellow_1(X1),X1)&m1_subset_1(k1_yellow_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', dt_k1_yellow_1)).
fof(dt_k2_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1)))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', dt_k2_yellow_1)).
fof(d1_yellow_1, axiom, (![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', d1_yellow_1)).
fof(d4_tarski, axiom, (![X1]:![X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', d4_tarski)).
fof(abstractness_v1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', abstractness_v1_orders_2)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', d9_lattice3)).
fof(d5_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>(v2_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r2_lattice3(X1,u1_struct_0(X1),X2))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', d5_yellow_0)).
fof(t1_subset, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', t1_subset)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', t2_subset)).
fof(t15_yellow_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>(v2_yellow_0(k2_yellow_1(X1))=>r2_hidden(k3_tarski(X1),X1)))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', t15_yellow_1)).
fof(t3_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', t3_yellow_1)).
fof(fc5_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1)))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1))))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', fc6_yellow_1)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', redefinition_r3_orders_2)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t15_yellow_1__17_premises__rerun__rerun.E_input', d3_tarski)).
fof(c_0_16, plain, (![X5]:![X6]:![X7]:![X8]:(((X5=X7|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5))))&((X6=X8|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])])])).
fof(c_0_17, plain, (![X2]:![X2]:![X2]:![X2]:![X2]:(v1_relat_2(k1_yellow_1(X2))&(v4_relat_2(k1_yellow_1(X2))&(v8_relat_2(k1_yellow_1(X2))&(v1_partfun1(k1_yellow_1(X2),X2)&m1_subset_1(k1_yellow_1(X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k1_yellow_1])])])).
fof(c_0_18, plain, (![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&l1_orders_2(k2_yellow_1(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k2_yellow_1])])])).
fof(c_0_19, plain, (![X2]:k2_yellow_1(X2)=g1_orders_2(X2,k1_yellow_1(X2))), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:(((((r2_hidden(X7,esk1_3(X5,X6,X7))|~r2_hidden(X7,X6))|X6!=k3_tarski(X5))&((r2_hidden(esk1_3(X5,X6,X7),X5)|~r2_hidden(X7,X6))|X6!=k3_tarski(X5)))&(((~r2_hidden(X7,X9)|~r2_hidden(X9,X5))|r2_hidden(X7,X6))|X6!=k3_tarski(X5)))&(((~r2_hidden(esk2_2(X5,X6),X6)|(~r2_hidden(esk2_2(X5,X6),X11)|~r2_hidden(X11,X5)))|X6=k3_tarski(X5))&(((r2_hidden(esk2_2(X5,X6),esk3_2(X5,X6))|r2_hidden(esk2_2(X5,X6),X6))|X6=k3_tarski(X5))&((r2_hidden(esk3_2(X5,X6),X5)|r2_hidden(esk2_2(X5,X6),X6))|X6=k3_tarski(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])])).
cnf(c_0_21,plain,(X2=X3|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|g1_orders_2(X2,X1)!=g1_orders_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(m1_subset_1(k1_yellow_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, (![X2]:(~l1_orders_2(X2)|(~v1_orders_2(X2)|X2=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
cnf(c_0_24,plain,(l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,plain,(k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(v1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_27, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk4_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk4_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
fof(c_0_28, plain, (![X3]:![X5]:((((m1_subset_1(esk5_1(X3),u1_struct_0(X3))|~v2_yellow_0(X3))|~l1_orders_2(X3))&((r2_lattice3(X3,u1_struct_0(X3),esk5_1(X3))|~v2_yellow_0(X3))|~l1_orders_2(X3)))&(((~m1_subset_1(X5,u1_struct_0(X3))|~r2_lattice3(X3,u1_struct_0(X3),X5))|v2_yellow_0(X3))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_0])])])])])])])).
fof(c_0_29, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|m1_subset_1(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_30,plain,(r2_hidden(X3,X1)|X1!=k3_tarski(X2)|~r2_hidden(X4,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_31, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_32,plain,(X1=X2|g1_orders_2(X1,k1_yellow_1(X1))!=g1_orders_2(X2,X3)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_33,plain,(X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34,plain,(l1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_35,plain,(v1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_36,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37,plain,(r2_lattice3(X1,u1_struct_0(X1),esk5_1(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38,plain,(m1_subset_1(esk5_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40,plain,(r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_41,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42,plain,(u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33])]), c_0_34]), c_0_35])])).
fof(c_0_43, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>(v2_yellow_0(k2_yellow_1(X1))=>r2_hidden(k3_tarski(X1),X1))))), inference(assume_negation,[status(cth)],[t15_yellow_1])).
cnf(c_0_44,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk4_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_45,plain,(r1_orders_2(X1,X2,esk5_1(X1))|~v2_yellow_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_46,plain,(r2_hidden(X3,esk1_3(X2,X1,X3))|X1!=k3_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_47,plain,(v1_xboole_0(X1)|r2_hidden(X2,k3_tarski(X1))|~r2_hidden(X2,X3)|~m1_subset_1(X3,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48,plain,(m1_subset_1(esk5_1(g1_orders_2(X1,k1_yellow_1(X1))),X1)|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_42]), c_0_34])])).
fof(c_0_49, negated_conjecture, ((~v1_xboole_0(esk6_0)&(v2_yellow_0(k2_yellow_1(esk6_0))&~r2_hidden(k3_tarski(esk6_0),esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_43])])])])).
fof(c_0_50, plain, (![X4]:![X5]:![X6]:(((((~r3_orders_2(k2_yellow_1(X4),X5,X6)|r1_tarski(X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4))&((((~r1_tarski(X5,X6)|r3_orders_2(k2_yellow_1(X4),X5,X6))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X4))))|~m1_subset_1(X5,u1_struct_0(k2_yellow_1(X4))))|v1_xboole_0(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])])])).
cnf(c_0_51,plain,(r2_lattice3(X1,X2,esk5_1(X1))|~v2_yellow_0(X1)|~r2_hidden(esk4_3(X1,X2,esk5_1(X1)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_38])).
cnf(c_0_52,plain,(X1=k3_tarski(X2)|~r2_hidden(X3,X2)|~r2_hidden(esk2_2(X2,X1),X3)|~r2_hidden(esk2_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_53,plain,(r2_hidden(X1,esk1_3(X2,k3_tarski(X2),X1))|~r2_hidden(X1,k3_tarski(X2))), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_54,plain,(r2_hidden(esk1_3(X2,X1,X3),X2)|X1!=k3_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_55,plain,(v1_xboole_0(X1)|r2_hidden(X2,k3_tarski(X1))|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))|~r2_hidden(X2,esk5_1(g1_orders_2(X1,k1_yellow_1(X1))))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_56,plain,(X1=k3_tarski(X2)|r2_hidden(esk2_2(X2,X1),X1)|r2_hidden(esk3_2(X2,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_57,negated_conjecture,(v2_yellow_0(k2_yellow_1(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_58,plain,(v1_xboole_0(X1)|r1_tarski(X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r3_orders_2(k2_yellow_1(X1),X2,X3)), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_59, plain, (![X2]:![X2]:![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&(v3_orders_2(k2_yellow_1(X2))&(v4_orders_2(k2_yellow_1(X2))&v5_orders_2(k2_yellow_1(X2)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc5_yellow_1])])])).
fof(c_0_60, plain, (![X2]:((~v2_struct_0(k2_yellow_1(X2))|v1_xboole_0(X2))&(v1_orders_2(k2_yellow_1(X2))|v1_xboole_0(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
cnf(c_0_61,plain,(v1_xboole_0(u1_struct_0(X1))|r2_lattice3(X1,X2,esk5_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(esk4_3(X1,X2,esk5_1(X1)),u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_51, c_0_41])).
cnf(c_0_62,plain,(r2_lattice3(X1,X3,X2)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_63,plain,(X1=k3_tarski(X2)|~r2_hidden(esk1_3(X3,k3_tarski(X3),esk2_2(X2,X1)),X2)|~r2_hidden(esk2_2(X2,X1),k3_tarski(X3))|~r2_hidden(esk2_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_64,plain,(r2_hidden(esk1_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_65,plain,(esk5_1(g1_orders_2(X1,k1_yellow_1(X1)))=k3_tarski(X2)|v1_xboole_0(X1)|r2_hidden(esk2_2(X2,esk5_1(g1_orders_2(X1,k1_yellow_1(X1)))),k3_tarski(X1))|r2_hidden(esk3_2(X2,esk5_1(g1_orders_2(X1,k1_yellow_1(X1)))),X2)|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_66,negated_conjecture,(v2_yellow_0(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))), inference(rw,[status(thm)],[c_0_57, c_0_25])).
cnf(c_0_67,negated_conjecture,(~v1_xboole_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_68,plain,(v1_xboole_0(X1)|r1_tarski(X2,X3)|~r3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_25]), c_0_25]), c_0_25])).
fof(c_0_69, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_70,plain,(v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_71,plain,(v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_72,plain,(v1_xboole_0(u1_struct_0(X1))|r2_lattice3(X1,X2,esk5_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_38])).
cnf(c_0_73,plain,(X1=k3_tarski(X2)|~r2_hidden(esk2_2(X2,X1),k3_tarski(X2))|~r2_hidden(esk2_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_74,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(X1)|r2_hidden(esk2_2(X1,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),k3_tarski(esk6_0))|r2_hidden(esk3_2(X1,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_75,plain,(r1_tarski(X1,X2)|v1_xboole_0(X3)|~r3_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_42]), c_0_42])).
cnf(c_0_76,plain,(v2_struct_0(X2)|r3_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r1_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77,plain,(v3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_70, c_0_25])).
cnf(c_0_78,plain,(v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_71, c_0_25])).
cnf(c_0_79,plain,(v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,X2,esk5_1(X1))|~v2_yellow_0(X1)|~r2_hidden(X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_72]), c_0_38])).
cnf(c_0_80,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|r2_hidden(esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),esk6_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_56])).
fof(c_0_81, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk7_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk7_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_82,plain,(r1_tarski(X1,X2)|v1_xboole_0(X3)|~r1_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_34]), c_0_42]), c_0_42])]), c_0_78])).
cnf(c_0_83,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),esk5_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84,plain,(X1=k3_tarski(X2)|r2_hidden(esk3_2(X2,X1),X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_56]), c_0_56])).
cnf(c_0_85,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_86,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|r1_tarski(esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),esk5_1(g1_orders_2(X1,k1_yellow_1(X1))))|v1_xboole_0(X1)|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))|~m1_subset_1(esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_42]), c_0_34]), c_0_42])]), c_0_48])).
cnf(c_0_87,plain,(X1=k3_tarski(X2)|r2_hidden(X3,k3_tarski(X2))|~r2_hidden(X3,esk3_2(X2,X1))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_40, c_0_84])).
cnf(c_0_88,plain,(X1=k3_tarski(X2)|r2_hidden(esk2_2(X2,X1),X1)|r2_hidden(esk2_2(X2,X1),esk3_2(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_89,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|v1_xboole_0(X1)|r2_hidden(X2,esk5_1(g1_orders_2(X1,k1_yellow_1(X1))))|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))|~r2_hidden(X2,esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))))|~m1_subset_1(esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_90,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|m1_subset_1(esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),esk6_0)), inference(spm,[status(thm)],[c_0_39, c_0_80])).
cnf(c_0_91,plain,(X1=k3_tarski(X2)|r2_hidden(esk2_2(X2,X1),k3_tarski(X2))|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_40])).
cnf(c_0_92,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|r2_hidden(X1,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0))))|~r2_hidden(X1,esk3_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_66])]), c_0_67])).
cnf(c_0_93,plain,(X1=k3_tarski(X2)|~r2_hidden(esk2_2(X2,X1),X1)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_73, c_0_91])).
cnf(c_0_94,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|r2_hidden(esk2_2(esk6_0,esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))),esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0))))), inference(spm,[status(thm)],[c_0_92, c_0_88])).
cnf(c_0_95,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|~r2_hidden(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0))),esk6_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_96,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)|~m1_subset_1(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0))),esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_41]), c_0_67])).
cnf(c_0_97,negated_conjecture,(~r2_hidden(k3_tarski(esk6_0),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_98,negated_conjecture,(esk5_1(g1_orders_2(esk6_0,k1_yellow_1(esk6_0)))=k3_tarski(esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_48]), c_0_66])])).
cnf(c_0_99,negated_conjecture,(~m1_subset_1(k3_tarski(esk6_0),esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_41]), c_0_67])).
cnf(c_0_100,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_98]), c_0_42]), c_0_66]), c_0_34])]), c_0_99]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 101
# Proof object clause steps            : 68
# Proof object formula steps           : 33
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 41
# Processed clauses                    : 488
# ...of these trivial                  : 3
# ...subsumed                          : 159
# ...remaining for further processing  : 326
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 88
# Generated clauses                    : 3040
# ...of the previous two non-trivial   : 3023
# Contextual simplify-reflections      : 204
# Paramodulations                      : 3013
# Factorizations                       : 6
# Equation resolutions                 : 21
# Current number of processed clauses  : 191
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 173
# Current number of unprocessed clauses: 1266
# ...number of literals in the above   : 7045
# Current number of archived formulas  : 0
# Current number of archived clauses   : 133
# Clause-clause subsumption calls (NU) : 27607
# Rec. Clause-clause subsumption calls : 6764
# Non-unit clause-clause subsumptions  : 368
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 99241

# -------------------------------------------------
# User time                : 0.108 s
# System time              : 0.000 s
# Total time               : 0.108 s
# Maximum resident set size: 43328 pages
