# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t15_yellow_0, conjecture, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('out/t15_yellow_0__2_premises__rerun.E_input', t15_yellow_0)).
fof(d7_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3)))))))), file('out/t15_yellow_0__2_premises__rerun.E_input', d7_yellow_0)).
fof(d3_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>(v5_orders_2(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X2))=>X2=X3)))))), file('out/t15_yellow_0__2_premises__rerun.E_input', d3_yellow_0)).
fof(c_0_3, negated_conjecture, (~(![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))))), inference(assume_negation,[status(cth)],[t15_yellow_0])).
fof(c_0_4, negated_conjecture, (![X7]:![X10]:((v5_orders_2(esk1_0)&l1_orders_2(esk1_0))&(((((m1_subset_1(esk3_1(X7),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X7))|~m1_subset_1(X7,u1_struct_0(esk1_0)))|~r1_yellow_0(esk1_0,esk2_0))&((((r2_lattice3(esk1_0,esk2_0,esk3_1(X7))|~r2_lattice3(esk1_0,esk2_0,X7))|~m1_subset_1(X7,u1_struct_0(esk1_0)))|~r1_yellow_0(esk1_0,esk2_0))&(((~r1_orders_2(esk1_0,X7,esk3_1(X7))|~r2_lattice3(esk1_0,esk2_0,X7))|~m1_subset_1(X7,u1_struct_0(esk1_0)))|~r1_yellow_0(esk1_0,esk2_0))))&((m1_subset_1(esk4_0,u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,esk2_0))&((r2_lattice3(esk1_0,esk2_0,esk4_0)|r1_yellow_0(esk1_0,esk2_0))&((~m1_subset_1(X10,u1_struct_0(esk1_0))|(~r2_lattice3(esk1_0,esk2_0,X10)|r1_orders_2(esk1_0,esk4_0,X10)))|r1_yellow_0(esk1_0,esk2_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])])])).
fof(c_0_5, plain, (![X6]:![X7]:![X9]:![X10]:![X7]:![X12]:![X15]:((((m1_subset_1(esk5_2(X6,X7),u1_struct_0(X6))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((r2_lattice3(X6,X7,esk5_2(X6,X7))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((~m1_subset_1(X9,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X9)|r1_orders_2(X6,esk5_2(X6,X7),X9)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((m1_subset_1(esk6_3(X6,X7,X10),u1_struct_0(X6))|~r2_lattice3(X6,X7,X10))|X10=esk5_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((r2_lattice3(X6,X7,esk6_3(X6,X7,X10))|~r2_lattice3(X6,X7,X10))|X10=esk5_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((~r1_orders_2(X6,X10,esk6_3(X6,X7,X10))|~r2_lattice3(X6,X7,X10))|X10=esk5_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6)))))))&(((((((m1_subset_1(esk8_3(X6,X7,X12),u1_struct_0(X6))|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r2_lattice3(X6,X7,esk8_3(X6,X7,X12))|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X15)|r1_orders_2(X6,esk8_3(X6,X7,X12),X15)))|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk8_3(X6,X7,X12)!=X12|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))))&(((((((m1_subset_1(esk8_3(X6,X7,X12),u1_struct_0(X6))|r2_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r2_lattice3(X6,X7,esk8_3(X6,X7,X12))|r2_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X15)|r1_orders_2(X6,esk8_3(X6,X7,X12),X15)))|r2_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk8_3(X6,X7,X12)!=X12|r2_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))))&((((((m1_subset_1(esk8_3(X6,X7,X12),u1_struct_0(X6))|~r1_orders_2(X6,X12,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r2_lattice3(X6,X7,esk8_3(X6,X7,X12))|~r1_orders_2(X6,X12,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X15)|r1_orders_2(X6,esk8_3(X6,X7,X12),X15)))|~r1_orders_2(X6,X12,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk8_3(X6,X7,X12)!=X12|~r1_orders_2(X6,X12,esk7_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])])).
cnf(c_0_6,negated_conjecture,(~r1_yellow_0(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(r1_orders_2(X1,esk5_2(X1,X2),X3)|~l1_orders_2(X1)|~r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(m1_subset_1(esk3_1(X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(~r2_lattice3(esk1_0,X1,esk3_1(esk5_2(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,X1))|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_11,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk3_1(X1))|~r1_yellow_0(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13,plain,(r2_lattice3(X1,X2,esk5_2(X1,X2))|~l1_orders_2(X1)|~r1_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,negated_conjecture,(~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_8])])).
cnf(c_0_15,plain,(m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~r1_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16,negated_conjecture,(r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(~r1_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_8])])).
cnf(c_0_18,plain,(r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_orders_2(X1,esk8_3(X1,X2,X3),X4)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r2_lattice3(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,negated_conjecture,(r1_orders_2(esk1_0,esk8_3(esk1_0,X1,X2),esk4_0)|r2_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,X2))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_8])])).
cnf(c_0_21,negated_conjecture,(r1_yellow_0(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|((~r1_orders_2(X4,X5,X6)|~r1_orders_2(X4,X6,X5))|X5=X6))))|~l1_orders_2(X4))&(((m1_subset_1(esk9_1(X4),u1_struct_0(X4))|v5_orders_2(X4))|~l1_orders_2(X4))&(((m1_subset_1(esk10_1(X4),u1_struct_0(X4))|v5_orders_2(X4))|~l1_orders_2(X4))&((((r1_orders_2(X4,esk9_1(X4),esk10_1(X4))|v5_orders_2(X4))|~l1_orders_2(X4))&((r1_orders_2(X4,esk10_1(X4),esk9_1(X4))|v5_orders_2(X4))|~l1_orders_2(X4)))&((esk9_1(X4)!=esk10_1(X4)|v5_orders_2(X4))|~l1_orders_2(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_yellow_0])])])])])])])).
cnf(c_0_23,negated_conjecture,(r1_orders_2(esk1_0,esk8_3(esk1_0,X1,esk4_0),esk4_0)|r2_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,esk4_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_20, c_0_19])).
cnf(c_0_24,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(sr,[status(thm)],[c_0_21, c_0_17])).
cnf(c_0_25,plain,(X2=X3|~l1_orders_2(X1)|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,negated_conjecture,(r1_orders_2(esk1_0,esk8_3(esk1_0,esk2_0,esk4_0),esk4_0)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_17])).
cnf(c_0_27,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_28,plain,(r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|esk8_3(X1,X2,X3)!=X3), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29,negated_conjecture,(r1_yellow_0(esk1_0,esk2_0)|r1_orders_2(esk1_0,esk4_0,X1)|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_30,negated_conjecture,(esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~r1_orders_2(esk1_0,esk4_0,esk8_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_19]), c_0_8]), c_0_27])])).
cnf(c_0_31,negated_conjecture,(r1_yellow_0(esk1_0,X1)|esk8_3(esk1_0,X1,esk4_0)!=esk4_0|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_8])]), c_0_19])]), c_0_17])).
cnf(c_0_32,plain,(r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|esk8_3(X1,X2,X3)!=X3), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_33,plain,(r1_yellow_0(X1,X2)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_34,plain,(r1_yellow_0(X1,X2)|r1_orders_2(X1,esk8_3(X1,X2,X3),X4)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|~r2_lattice3(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_35,plain,(r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_36,negated_conjecture,(esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_29]), c_0_17])).
cnf(c_0_37,plain,(r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_38,negated_conjecture,(esk8_3(esk1_0,esk2_0,esk4_0)!=esk4_0|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_24]), c_0_19]), c_0_8])]), c_0_17])).
cnf(c_0_39,plain,(r1_yellow_0(X1,X2)|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|esk8_3(X1,X2,X3)!=X3), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk8_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_29]), c_0_8])]), c_0_19])]), c_0_17])).
cnf(c_0_41,plain,(r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_42,negated_conjecture,(X1=esk4_0|~r1_orders_2(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_29]), c_0_8]), c_0_27])]), c_0_19])]), c_0_17])).
cnf(c_0_43,negated_conjecture,(r1_orders_2(esk1_0,esk8_3(esk1_0,X1,esk4_0),X2)|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_29]), c_0_8])]), c_0_19])]), c_0_17])).
cnf(c_0_44,negated_conjecture,(r2_lattice3(esk1_0,X1,esk8_3(esk1_0,X1,esk4_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_29]), c_0_8])]), c_0_19])]), c_0_17])).
cnf(c_0_45,negated_conjecture,(esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_24]), c_0_19]), c_0_8])]), c_0_17])).
cnf(c_0_46,negated_conjecture,(esk8_3(esk1_0,esk2_0,esk4_0)!=esk4_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_24]), c_0_19]), c_0_8])]), c_0_17])).
cnf(c_0_47,negated_conjecture,(m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_24]), c_0_19]), c_0_8])]), c_0_17])).
cnf(c_0_48,plain,(r1_yellow_0(X1,X2)|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_orders_2(X1,esk8_3(X1,X2,X3),X4)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r2_lattice3(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_49,negated_conjecture,(r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_19])]), c_0_40]), c_0_31])).
cnf(c_0_50,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_24])]), c_0_17]), c_0_46]), c_0_47])).
cnf(c_0_51,negated_conjecture,(esk8_3(esk1_0,X1,X2)=esk4_0|m1_subset_1(esk7_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,X1,X2))|~r2_lattice3(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(esk8_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_48]), c_0_19]), c_0_8])])).
cnf(c_0_52,plain,(r1_yellow_0(X1,X2)|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_53,negated_conjecture,(~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_24])]), c_0_17])).
cnf(c_0_54,negated_conjecture,(esk8_3(esk1_0,esk2_0,X1)=esk4_0|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_24]), c_0_8])]), c_0_17])).
cnf(c_0_55,plain,(r1_yellow_0(X1,X2)|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_56,negated_conjecture,(~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_57,negated_conjecture,(esk8_3(esk1_0,esk2_0,X1)=esk4_0|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_8])]), c_0_17])).
cnf(c_0_58,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_24]), c_0_19])]), c_0_46]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 52
# Proof object formula steps           : 7
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 91
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 109
# ...of these trivial                  : 0
# ...subsumed                          : 21
# ...remaining for further processing  : 88
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 0
# Generated clauses                    : 150
# ...of the previous two non-trivial   : 124
# Contextual simplify-reflections      : 35
# Paramodulations                      : 148
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 77
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 70
# Current number of unprocessed clauses: 39
# ...number of literals in the above   : 362
# Current number of archived formulas  : 0
# Current number of archived clauses   : 11
# Clause-clause subsumption calls (NU) : 3296
# Rec. Clause-clause subsumption calls : 196
# Non-unit clause-clause subsumptions  : 52
# Unit Clause-clause subsumption calls : 54
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9123

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43184 pages
