# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_funct_2, axiom, (![X1]:![X2]:![X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2))))))), file('out/t169_funct_2__1_premises__rerun.E_input', d2_funct_2)).
fof(t169_funct_2, conjecture, (![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X3,k1_funct_2(X1,X2))=>(k1_relat_1(X3)=X1&r1_tarski(k2_relat_1(X3),X2))))), file('out/t169_funct_2__1_premises__rerun.E_input', t169_funct_2)).
fof(c_0_2, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:![X6]:![X7]:![X8]:![X13]:((((((v1_relat_1(esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X8))|X8!=k1_funct_2(X6,X7))&((v1_funct_1(esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X8))|X8!=k1_funct_2(X6,X7)))&(((X9=esk1_4(X6,X7,X8,X9)|~r2_hidden(X9,X8))|X8!=k1_funct_2(X6,X7))&(((k1_relat_1(esk1_4(X6,X7,X8,X9))=X6|~r2_hidden(X9,X8))|X8!=k1_funct_2(X6,X7))&((r1_tarski(k2_relat_1(esk1_4(X6,X7,X8,X9)),X7)|~r2_hidden(X9,X8))|X8!=k1_funct_2(X6,X7)))))&((((~v1_relat_1(X11)|~v1_funct_1(X11))|(X9!=X11|(k1_relat_1(X11)!=X6|~r1_tarski(k2_relat_1(X11),X7))))|r2_hidden(X9,X8))|X8!=k1_funct_2(X6,X7)))&(((~r2_hidden(esk2_3(X6,X7,X8),X8)|((~v1_relat_1(X13)|~v1_funct_1(X13))|(esk2_3(X6,X7,X8)!=X13|(k1_relat_1(X13)!=X6|~r1_tarski(k2_relat_1(X13),X7)))))|X8=k1_funct_2(X6,X7))&((((v1_relat_1(esk3_3(X6,X7,X8))|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k1_funct_2(X6,X7))&((v1_funct_1(esk3_3(X6,X7,X8))|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k1_funct_2(X6,X7)))&(((esk2_3(X6,X7,X8)=esk3_3(X6,X7,X8)|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k1_funct_2(X6,X7))&(((k1_relat_1(esk3_3(X6,X7,X8))=X6|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k1_funct_2(X6,X7))&((r1_tarski(k2_relat_1(esk3_3(X6,X7,X8)),X7)|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k1_funct_2(X6,X7)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])])).
cnf(c_0_3,plain,(k1_relat_1(esk1_4(X2,X3,X1,X4))=X2|X1!=k1_funct_2(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_4,plain,(X4=esk1_4(X2,X3,X1,X4)|X1!=k1_funct_2(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
fof(c_0_5, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X3,k1_funct_2(X1,X2))=>(k1_relat_1(X3)=X1&r1_tarski(k2_relat_1(X3),X2)))))), inference(assume_negation,[status(cth)],[t169_funct_2])).
cnf(c_0_6,plain,(k1_relat_1(esk1_4(X1,X2,k1_funct_2(X1,X2),X3))=X1|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_3])).
cnf(c_0_7,plain,(esk1_4(X1,X2,k1_funct_2(X1,X2),X3)=X3|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_4])).
fof(c_0_8, negated_conjecture, (((v1_relat_1(esk6_0)&v1_funct_1(esk6_0))&(r2_hidden(esk6_0,k1_funct_2(esk4_0,esk5_0))&(k1_relat_1(esk6_0)!=esk4_0|~r1_tarski(k2_relat_1(esk6_0),esk5_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_9,plain,(r1_tarski(k2_relat_1(esk1_4(X2,X3,X1,X4)),X3)|X1!=k1_funct_2(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_10,plain,(k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11,negated_conjecture,(r2_hidden(esk6_0,k1_funct_2(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(r1_tarski(k2_relat_1(esk1_4(X1,X2,k1_funct_2(X1,X2),X3)),X2)|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(~r1_tarski(k2_relat_1(esk6_0),esk5_0)|k1_relat_1(esk6_0)!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,negated_conjecture,(k1_relat_1(esk6_0)=esk4_0), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15,plain,(r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(spm,[status(thm)],[c_0_12, c_0_7])).
cnf(c_0_16,negated_conjecture,(~r1_tarski(k2_relat_1(esk6_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14])])).
cnf(c_0_17,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_11]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 18
# Proof object clause steps            : 13
# Proof object formula steps           : 5
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 45
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 37
# ...of the previous two non-trivial   : 36
# Contextual simplify-reflections      : 0
# Paramodulations                      : 30
# Factorizations                       : 0
# Equation resolutions                 : 9
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 22
# ...number of literals in the above   : 69
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 50
# Rec. Clause-clause subsumption calls : 27
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1643

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45200 pages
