# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_funct_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((~(X2=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X3)))&(X2=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('out/t16_yellow_6__15_premises__rerun.E_input', d1_funct_2)).
fof(dt_u1_waybel_0, axiom, (![X1]:![X2]:((l1_struct_0(X1)&l1_waybel_0(X2,X1))=>(v1_funct_1(u1_waybel_0(X1,X2))&(v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1)))))))), file('out/t16_yellow_6__15_premises__rerun.E_input', dt_u1_waybel_0)).
fof(t16_yellow_6, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v1_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>k2_waybel_0(X1,X2,X3)=k4_yellow_6(X1,X2))))), file('out/t16_yellow_6__15_premises__rerun.E_input', t16_yellow_6)).
fof(cc2_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2)))), file('out/t16_yellow_6__15_premises__rerun.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('out/t16_yellow_6__15_premises__rerun.E_input', cc1_relset_1)).
fof(redefinition_k1_relset_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2))), file('out/t16_yellow_6__15_premises__rerun.E_input', redefinition_k1_relset_1)).
fof(d12_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v3_funct_1(X1)=>(v1_xboole_0(X1)|![X2]:(X2=k3_funct_1(X1)<=>?[X3]:(r2_hidden(X3,k9_xtuple_0(X1))&X2=k1_funct_1(X1,X3))))))), file('out/t16_yellow_6__15_premises__rerun.E_input', d12_funct_1)).
fof(d8_yellow_6, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:(l1_waybel_0(X2,X1)=>(v1_yellow_6(X2,X1)=>(v2_struct_0(X2)|k4_yellow_6(X1,X2)=k3_funct_1(u1_waybel_0(X1,X2))))))), file('out/t16_yellow_6__15_premises__rerun.E_input', d8_yellow_6)).
fof(fc16_yellow_6, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(v1_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))=>(v1_funct_1(u1_waybel_0(X1,X2))&(v3_funct_1(u1_waybel_0(X1,X2))&v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1)))))), file('out/t16_yellow_6__15_premises__rerun.E_input', fc16_yellow_6)).
fof(redefinition_k3_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4))), file('out/t16_yellow_6__15_premises__rerun.E_input', redefinition_k3_funct_2)).
fof(d8_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>k2_waybel_0(X1,X2,X3)=k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),X3))))), file('out/t16_yellow_6__15_premises__rerun.E_input', d8_waybel_0)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t16_yellow_6__15_premises__rerun.E_input', t2_subset)).
fof(cc3_relset_1, axiom, (![X1]:![X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3)))), file('out/t16_yellow_6__15_premises__rerun.E_input', cc3_relset_1)).
fof(fc15_yellow_6, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&l1_waybel_0(X2,X1)))=>(v1_funct_1(u1_waybel_0(X1,X2))&(~(v1_xboole_0(u1_waybel_0(X1,X2)))&v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1)))))), file('out/t16_yellow_6__15_premises__rerun.E_input', fc15_yellow_6)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t16_yellow_6__15_premises__rerun.E_input', fc2_struct_0)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t16_yellow_6__15_premises__rerun.E_input', fc1_xboole_0)).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(((((~v1_funct_2(X6,X4,X5)|X4=k1_relset_1(X4,X6))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X4!=k1_relset_1(X4,X6)|v1_funct_2(X6,X4,X5))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))&((((~v1_funct_2(X6,X4,X5)|X6=k1_xboole_0)|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X6!=k1_xboole_0|v1_funct_2(X6,X4,X5))|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_funct_2])])])])).
fof(c_0_17, plain, (![X3]:![X4]:((v1_funct_1(u1_waybel_0(X3,X4))|(~l1_struct_0(X3)|~l1_waybel_0(X4,X3)))&((v1_funct_2(u1_waybel_0(X3,X4),u1_struct_0(X4),u1_struct_0(X3))|(~l1_struct_0(X3)|~l1_waybel_0(X4,X3)))&(m1_subset_1(u1_waybel_0(X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|(~l1_struct_0(X3)|~l1_waybel_0(X4,X3)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_waybel_0])])])).
fof(c_0_18, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v1_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>k2_waybel_0(X1,X2,X3)=k4_yellow_6(X1,X2)))))), inference(assume_negation,[status(cth)],[t16_yellow_6])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:((v4_relat_1(X6,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(v5_relat_1(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_21,plain,(X3=k1_xboole_0|X2=k1_relset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(m1_subset_1(u1_waybel_0(X2,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(v1_funct_2(u1_waybel_0(X2,X1),u1_struct_0(X1),u1_struct_0(X2))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, negated_conjecture, (((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&(v1_yellow_6(esk2_0,esk1_0)&l1_waybel_0(esk2_0,esk1_0)))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&k2_waybel_0(esk1_0,esk2_0,esk3_0)!=k4_yellow_6(esk1_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])])])).
fof(c_0_25, plain, (![X3]:![X4]:((~v1_relat_1(X4)|~v4_relat_1(X4,X3))|k1_relset_1(X3,X4)=k9_xtuple_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_26,plain,(v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, (![X4]:![X5]:![X5]:![X7]:((((((r2_hidden(esk4_2(X4,X5),k9_xtuple_0(X4))|X5!=k3_funct_1(X4))|v1_xboole_0(X4))|~v3_funct_1(X4))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&((((X5=k1_funct_1(X4,esk4_2(X4,X5))|X5!=k3_funct_1(X4))|v1_xboole_0(X4))|~v3_funct_1(X4))|(~v1_relat_1(X4)|~v1_funct_1(X4))))&(((((~r2_hidden(X7,k9_xtuple_0(X4))|X5!=k1_funct_1(X4,X7))|X5=k3_funct_1(X4))|v1_xboole_0(X4))|~v3_funct_1(X4))|(~v1_relat_1(X4)|~v1_funct_1(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_funct_1])])])])])])])).
cnf(c_0_29,plain,(k1_relset_1(u1_struct_0(X1),u1_waybel_0(X2,X1))=u1_struct_0(X1)|u1_struct_0(X2)=k1_xboole_0|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_30,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,negated_conjecture,(l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(k1_relset_1(X1,X2)=k9_xtuple_0(X2)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,plain,(v4_relat_1(u1_waybel_0(X1,X2),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_22])).
cnf(c_0_34,plain,(v1_relat_1(u1_waybel_0(X1,X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_27, c_0_22])).
fof(c_0_35, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_struct_0(X3))|(~l1_waybel_0(X4,X3)|(~v1_yellow_6(X4,X3)|(v2_struct_0(X4)|k4_yellow_6(X3,X4)=k3_funct_1(u1_waybel_0(X3,X4))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_yellow_6])])])])])])).
fof(c_0_36, plain, (![X3]:![X4]:((v1_funct_1(u1_waybel_0(X3,X4))|((v2_struct_0(X3)|~l1_struct_0(X3))|(~v1_yellow_6(X4,X3)|~l1_waybel_0(X4,X3))))&((v3_funct_1(u1_waybel_0(X3,X4))|((v2_struct_0(X3)|~l1_struct_0(X3))|(~v1_yellow_6(X4,X3)|~l1_waybel_0(X4,X3))))&(v1_funct_2(u1_waybel_0(X3,X4),u1_struct_0(X4),u1_struct_0(X3))|((v2_struct_0(X3)|~l1_struct_0(X3))|(~v1_yellow_6(X4,X3)|~l1_waybel_0(X4,X3))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc16_yellow_6])])])])).
cnf(c_0_37,plain,(v1_xboole_0(X1)|X2=k3_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_funct_1(X1)|X2!=k1_funct_1(X1,X3)|~r2_hidden(X3,k9_xtuple_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38,negated_conjecture,(k1_relset_1(u1_struct_0(esk2_0),u1_waybel_0(esk1_0,esk2_0))=u1_struct_0(esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_39,plain,(k1_relset_1(u1_struct_0(X1),u1_waybel_0(X2,X1))=k9_xtuple_0(u1_waybel_0(X2,X1))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_40,plain,(k4_yellow_6(X1,X2)=k3_funct_1(u1_waybel_0(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_yellow_6(X2,X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41,negated_conjecture,(v1_yellow_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_43,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_44,plain,(v2_struct_0(X2)|v3_funct_1(u1_waybel_0(X2,X1))|~l1_waybel_0(X1,X2)|~v1_yellow_6(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_45, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X8,X5)))|k3_funct_2(X5,X6,X7,X8)=k1_funct_1(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
fof(c_0_46, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|~l1_struct_0(X4))|((v2_struct_0(X5)|~l1_waybel_0(X5,X4))|(~m1_subset_1(X6,u1_struct_0(X5))|k2_waybel_0(X4,X5,X6)=k3_funct_2(u1_struct_0(X5),u1_struct_0(X4),u1_waybel_0(X4,X5),X6))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_waybel_0])])])])])])).
cnf(c_0_47,plain,(k1_funct_1(X1,X2)=k3_funct_1(X1)|v1_xboole_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1))|~v3_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_48,negated_conjecture,(k9_xtuple_0(u1_waybel_0(esk1_0,esk2_0))=u1_struct_0(esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_30]), c_0_31])])).
cnf(c_0_49,negated_conjecture,(k3_funct_1(u1_waybel_0(esk1_0,esk2_0))=k4_yellow_6(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_30]), c_0_31])]), c_0_42]), c_0_43])).
cnf(c_0_50,negated_conjecture,(v3_funct_1(u1_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_41]), c_0_30]), c_0_31])]), c_0_43])).
cnf(c_0_51,plain,(k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52,plain,(v1_funct_1(u1_waybel_0(X2,X1))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_53,plain,(k2_waybel_0(X1,X2,X3)=k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_55,negated_conjecture,(k1_funct_1(u1_waybel_0(esk1_0,esk2_0),X1)=k4_yellow_6(esk1_0,esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0|v1_xboole_0(u1_waybel_0(esk1_0,esk2_0))|~r2_hidden(X1,u1_struct_0(esk2_0))|~v1_relat_1(u1_waybel_0(esk1_0,esk2_0))|~v1_funct_1(u1_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50])])).
fof(c_0_56, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_57,plain,(k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),u1_waybel_0(X2,X1),X3)=k1_funct_1(u1_waybel_0(X2,X1),X3)|v1_xboole_0(u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_22]), c_0_23]), c_0_52])).
cnf(c_0_58,negated_conjecture,(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),u1_waybel_0(X1,esk2_0),esk3_0)=k2_waybel_0(X1,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_42])).
cnf(c_0_59,negated_conjecture,(k1_funct_1(u1_waybel_0(esk1_0,esk2_0),X1)=k4_yellow_6(esk1_0,esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0|v1_xboole_0(u1_waybel_0(esk1_0,esk2_0))|~r2_hidden(X1,u1_struct_0(esk2_0))|~v1_funct_1(u1_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_34]), c_0_30]), c_0_31])])).
cnf(c_0_60,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_61, plain, (![X4]:![X5]:![X6]:(~v1_xboole_0(X4)|(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])])])).
cnf(c_0_62,negated_conjecture,(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),u1_waybel_0(X1,esk2_0),esk3_0)=k1_funct_1(u1_waybel_0(X1,esk2_0),esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_57, c_0_54])).
cnf(c_0_63,negated_conjecture,(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0),esk3_0)=k2_waybel_0(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_30]), c_0_31])]), c_0_43])).
cnf(c_0_64,negated_conjecture,(k1_funct_1(u1_waybel_0(esk1_0,esk2_0),X1)=k4_yellow_6(esk1_0,esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0|v1_xboole_0(u1_waybel_0(esk1_0,esk2_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_52]), c_0_30]), c_0_31])])).
cnf(c_0_65,negated_conjecture,(r2_hidden(esk3_0,u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_60, c_0_54])).
fof(c_0_66, plain, (![X3]:![X4]:((v1_funct_1(u1_waybel_0(X3,X4))|((v2_struct_0(X3)|~l1_struct_0(X3))|(v2_struct_0(X4)|~l1_waybel_0(X4,X3))))&((~v1_xboole_0(u1_waybel_0(X3,X4))|((v2_struct_0(X3)|~l1_struct_0(X3))|(v2_struct_0(X4)|~l1_waybel_0(X4,X3))))&(v1_funct_2(u1_waybel_0(X3,X4),u1_struct_0(X4),u1_struct_0(X3))|((v2_struct_0(X3)|~l1_struct_0(X3))|(v2_struct_0(X4)|~l1_waybel_0(X4,X3))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc15_yellow_6])])])])).
cnf(c_0_67,plain,(v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68,negated_conjecture,(k1_funct_1(u1_waybel_0(esk1_0,esk2_0),esk3_0)=k2_waybel_0(esk1_0,esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_30]), c_0_63]), c_0_31])])).
cnf(c_0_69,negated_conjecture,(k1_funct_1(u1_waybel_0(esk1_0,esk2_0),esk3_0)=k4_yellow_6(esk1_0,esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0|v1_xboole_0(u1_waybel_0(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70,negated_conjecture,(k2_waybel_0(esk1_0,esk2_0,esk3_0)!=k4_yellow_6(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_71,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)|~v1_xboole_0(u1_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72,plain,(v1_xboole_0(u1_waybel_0(X1,X2))|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_67, c_0_22])).
cnf(c_0_73,negated_conjecture,(u1_struct_0(esk1_0)=k1_xboole_0|v1_xboole_0(u1_waybel_0(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
cnf(c_0_74,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_75,negated_conjecture,(u1_struct_0(esk1_0)=k1_xboole_0|v1_xboole_0(u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_73]), c_0_30]), c_0_31])]), c_0_43]), c_0_42])).
fof(c_0_76, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_77,negated_conjecture,(u1_struct_0(esk1_0)=k1_xboole_0|v2_struct_0(X1)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_42])).
cnf(c_0_78,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_79,negated_conjecture,(u1_struct_0(esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_30]), c_0_31])]), c_0_43])).
cnf(c_0_80,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_81,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_31])]), c_0_43]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 50
# Proof object formula steps           : 32
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 148
# ...of these trivial                  : 0
# ...subsumed                          : 20
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 19
# Generated clauses                    : 152
# ...of the previous two non-trivial   : 128
# Contextual simplify-reflections      : 25
# Paramodulations                      : 147
# Factorizations                       : 0
# Equation resolutions                 : 5
# Current number of processed clauses  : 68
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 54
# Current number of unprocessed clauses: 23
# ...number of literals in the above   : 59
# Current number of archived formulas  : 0
# Current number of archived clauses   : 60
# Clause-clause subsumption calls (NU) : 639
# Rec. Clause-clause subsumption calls : 183
# Non-unit clause-clause subsumptions  : 47
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6381

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45264 pages
