# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_funct_2, conjecture, (![X1]:![X2]:![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~((r2_hidden(X5,X1)&X4=k1_funct_1(X3,X5)))))=>k2_relset_1(X1,X2,X3)=X2))), file('out/t16_funct_2__8_premises__rerun.E_input', t16_funct_2)).
fof(d12_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:![X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(X5,k1_relat_1(X1))&(r2_hidden(X5,X2)&X4=k1_funct_1(X1,X5))))))), file('out/t16_funct_2__8_premises__rerun.E_input', d12_funct_1)).
fof(cc1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('out/t16_funct_2__8_premises__rerun.E_input', cc1_relset_1)).
fof(redefinition_k1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3))), file('out/t16_funct_2__8_premises__rerun.E_input', redefinition_k1_relset_1)).
fof(d1_funct_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0)))))), file('out/t16_funct_2__8_premises__rerun.E_input', d1_funct_2)).
fof(t23_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X5,X4),X3))))<=>k2_relset_1(X1,X2,X3)=X2))), file('out/t16_funct_2__8_premises__rerun.E_input', t23_relset_1)).
fof(t143_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2)))))), file('out/t16_funct_2__8_premises__rerun.E_input', t143_relat_1)).
fof(redefinition_k2_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3))), file('out/t16_funct_2__8_premises__rerun.E_input', redefinition_k2_relset_1)).
fof(t60_relat_1, axiom, ((k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0)), file('out/t16_funct_2__8_premises__rerun.E_input', t60_relat_1)).
fof(c_0_9, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~((r2_hidden(X5,X1)&X4=k1_funct_1(X3,X5)))))=>k2_relset_1(X1,X2,X3)=X2)))), inference(assume_negation,[status(cth)],[t16_funct_2])).
fof(c_0_10, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:![X7]:![X8]:![X13]:((((((r2_hidden(esk5_4(X6,X7,X8,X9),k1_relat_1(X6))|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((r2_hidden(esk5_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&(((X9=k1_funct_1(X6,esk5_4(X6,X7,X8,X9))|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))))&((((~r2_hidden(X11,k1_relat_1(X6))|(~r2_hidden(X11,X7)|X9!=k1_funct_1(X6,X11)))|r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&((((~r2_hidden(esk6_3(X6,X7,X8),X8)|(~r2_hidden(X13,k1_relat_1(X6))|(~r2_hidden(X13,X7)|esk6_3(X6,X7,X8)!=k1_funct_1(X6,X13))))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((r2_hidden(esk7_3(X6,X7,X8),k1_relat_1(X6))|r2_hidden(esk6_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((r2_hidden(esk7_3(X6,X7,X8),X7)|r2_hidden(esk6_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&(((esk6_3(X6,X7,X8)=k1_funct_1(X6,esk7_3(X6,X7,X8))|r2_hidden(esk6_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|(~v1_relat_1(X6)|~v1_funct_1(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_funct_1])])])])])])])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_12, negated_conjecture, (![X9]:((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(((r2_hidden(esk4_1(X9),esk1_0)|~r2_hidden(X9,esk2_0))&(X9=k1_funct_1(esk3_0,esk4_1(X9))|~r2_hidden(X9,esk2_0)))&k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|k1_relset_1(X4,X5,X6)=k1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:((((((~v1_funct_2(X6,X4,X5)|X4=k1_relset_1(X4,X5,X6))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X4!=k1_relset_1(X4,X5,X6)|v1_funct_2(X6,X4,X5))|X5=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))&((((~v1_funct_2(X6,X4,X5)|X4=k1_relset_1(X4,X5,X6))|X4!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(((X4!=k1_relset_1(X4,X5,X6)|v1_funct_2(X6,X4,X5))|X4!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))&(((((~v1_funct_2(X6,X4,X5)|X6=k1_xboole_0)|X4=k1_xboole_0)|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&((((X6!=k1_xboole_0|v1_funct_2(X6,X4,X5))|X4=k1_xboole_0)|X5!=k1_xboole_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_15,plain,(r2_hidden(X4,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k9_relat_1(X1,X3)|X4!=k1_funct_1(X1,X5)|~r2_hidden(X5,X3)|~r2_hidden(X5,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(k1_relset_1(X1,X2,X3)=k1_relat_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(X3=k1_xboole_0|X2=k1_relset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(r2_hidden(k1_funct_1(X1,X2),k9_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X2,X3)|~v1_funct_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_15])])).
cnf(c_0_21,negated_conjecture,(X1=k1_funct_1(esk3_0,esk4_1(X1))|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,negated_conjecture,(v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,plain,(X1=k1_relat_1(X2)|X3=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_2(X2,X1,X3)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,negated_conjecture,(v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_26, plain, (![X6]:![X7]:![X8]:![X10]:![X11]:((((r2_hidden(esk8_3(X6,X7,X8),X7)|k2_relset_1(X6,X7,X8)=X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))&((~r2_hidden(k4_tarski(X10,esk8_3(X6,X7,X8)),X8)|k2_relset_1(X6,X7,X8)=X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))&((k2_relset_1(X6,X7,X8)!=X7|(~r2_hidden(X11,X7)|r2_hidden(k4_tarski(esk9_4(X6,X7,X8,X11),X11),X8)))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t23_relset_1])])])])])])])])).
fof(c_0_27, plain, (![X5]:![X6]:![X7]:![X9]:((((r2_hidden(esk10_3(X5,X6,X7),k1_relat_1(X7))|~r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7))&(((r2_hidden(k4_tarski(esk10_3(X5,X6,X7),X5),X7)|~r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7))&((r2_hidden(esk10_3(X5,X6,X7),X6)|~r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7))))&(((~r2_hidden(X9,k1_relat_1(X7))|(~r2_hidden(k4_tarski(X9,X5),X7)|~r2_hidden(X9,X6)))|r2_hidden(X5,k9_relat_1(X7,X6)))|~v1_relat_1(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])])])).
cnf(c_0_28,negated_conjecture,(r2_hidden(X1,k9_relat_1(esk3_0,X2))|~r2_hidden(esk4_1(X1),k1_relat_1(esk3_0))|~r2_hidden(esk4_1(X1),X2)|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])])).
cnf(c_0_29,negated_conjecture,(k1_relat_1(esk3_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_17]), c_0_25])])).
cnf(c_0_30,negated_conjecture,(r2_hidden(esk4_1(X1),esk1_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31,plain,(k2_relset_1(X2,X3,X1)=X3|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(k4_tarski(X4,esk8_3(X2,X3,X1)),X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32,plain,(r2_hidden(k4_tarski(esk10_3(X2,X3,X1),X2),X1)|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,negated_conjecture,(esk2_0=k1_xboole_0|r2_hidden(X1,k9_relat_1(esk3_0,X2))|~r2_hidden(esk4_1(X1),X2)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_34,plain,(k2_relset_1(X1,X2,X3)=X2|~r2_hidden(esk8_3(X1,X2,X3),k9_relat_1(X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_16])).
cnf(c_0_35,negated_conjecture,(esk2_0=k1_xboole_0|r2_hidden(X1,k9_relat_1(esk3_0,esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_33, c_0_30])).
cnf(c_0_36,negated_conjecture,(k2_relset_1(X1,X2,esk3_0)=X2|esk2_0=k1_xboole_0|~r2_hidden(esk8_3(X1,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_37,plain,(k2_relset_1(X2,X3,X1)=X3|r2_hidden(esk8_3(X2,X3,X1),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38,negated_conjecture,(k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_39,negated_conjecture,(k2_relset_1(X1,esk2_0,esk3_0)=esk2_0|esk2_0=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40,plain,(X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|X3!=k1_xboole_0|~v1_funct_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_41,negated_conjecture,(esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_17])])).
cnf(c_0_42,plain,(X2=k1_relset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|X2!=k1_xboole_0|~v1_funct_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_43,plain,(X1=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_xboole_0)))|~v1_funct_2(X2,X1,k1_xboole_0)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k1_xboole_0)))), inference(rw,[status(thm)],[c_0_17, c_0_41])).
cnf(c_0_45,negated_conjecture,(v1_funct_2(esk3_0,esk1_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_25, c_0_41])).
cnf(c_0_46,plain,(k1_relset_1(k1_xboole_0,X1,X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X1)))|~v1_funct_2(X2,k1_xboole_0,X1)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_47,negated_conjecture,(esk1_0=k1_xboole_0|esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_48,plain,(k1_relat_1(X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X2)))|~v1_funct_2(X1,k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_18, c_0_46])).
cnf(c_0_49,negated_conjecture,(esk3_0=k1_xboole_0|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_44, c_0_47])).
cnf(c_0_50,negated_conjecture,(esk3_0=k1_xboole_0|v1_funct_2(esk3_0,k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_45, c_0_47])).
cnf(c_0_51,plain,(r2_hidden(X2,k9_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X4,X3)|~r2_hidden(k4_tarski(X4,X2),X1)|~r2_hidden(X4,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_52,plain,(r2_hidden(esk10_3(X2,X3,X1),k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_53,negated_conjecture,(k1_relat_1(esk3_0)=k1_xboole_0|esk3_0=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_54,plain,(r2_hidden(X1,k9_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(esk10_3(X1,X4,X2),X3)|~r2_hidden(X1,k9_relat_1(X2,X4))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_32]), c_0_52])).
cnf(c_0_55,negated_conjecture,(esk3_0=k1_xboole_0|r2_hidden(esk10_3(X1,X2,esk3_0),k1_xboole_0)|~r2_hidden(X1,k9_relat_1(esk3_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_22])])).
cnf(c_0_56,negated_conjecture,(esk3_0=k1_xboole_0|r2_hidden(X1,k9_relat_1(esk3_0,k1_xboole_0))|~r2_hidden(X1,k9_relat_1(esk3_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_22])])).
cnf(c_0_57,negated_conjecture,(esk3_0=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,X1),k9_relat_1(esk3_0,k1_xboole_0))|~r2_hidden(X1,k1_relat_1(esk3_0))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_20]), c_0_22]), c_0_23])])).
cnf(c_0_58,negated_conjecture,(r2_hidden(esk4_1(X1),esk1_0)|~r2_hidden(X1,k1_xboole_0)), inference(rw,[status(thm)],[c_0_30, c_0_41])).
cnf(c_0_59,negated_conjecture,(esk3_0=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk4_1(X1)),k9_relat_1(esk3_0,k1_xboole_0))|~r2_hidden(esk4_1(X1),k1_relat_1(esk3_0))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_60,negated_conjecture,(k1_funct_1(esk3_0,esk4_1(X1))=X1|~r2_hidden(X1,k1_xboole_0)), inference(rw,[status(thm)],[c_0_21, c_0_41])).
fof(c_0_61, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|k2_relset_1(X4,X5,X6)=k2_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_62,negated_conjecture,(esk3_0=k1_xboole_0|r2_hidden(X1,k9_relat_1(esk3_0,k1_xboole_0))|~r2_hidden(esk4_1(X1),k1_relat_1(esk3_0))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_63,negated_conjecture,(esk3_0=k1_xboole_0|r2_hidden(esk4_1(X1),k1_xboole_0)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_58, c_0_47])).
cnf(c_0_64,plain,(k2_relset_1(X1,X2,X3)=k2_relat_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_65,negated_conjecture,(esk3_0=k1_xboole_0|r2_hidden(X1,k9_relat_1(esk3_0,k1_xboole_0))|~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_53]), c_0_63])).
cnf(c_0_66,negated_conjecture,(k2_relset_1(esk1_0,k1_xboole_0,esk3_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_41]), c_0_41])).
cnf(c_0_67,plain,(k2_relset_1(X1,X2,X3)=k2_relset_1(X4,X5,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))), inference(spm,[status(thm)],[c_0_64, c_0_64])).
cnf(c_0_68,negated_conjecture,(k2_relset_1(X1,X2,esk3_0)=X2|esk3_0=k1_xboole_0|~r2_hidden(esk8_3(X1,X2,esk3_0),k1_xboole_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_34, c_0_65])).
cnf(c_0_69,negated_conjecture,(k2_relset_1(X1,X2,esk3_0)!=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_44])])).
cnf(c_0_70,negated_conjecture,(esk3_0=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,k1_xboole_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_37]), c_0_69])).
cnf(c_0_71,negated_conjecture,(esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_70, c_0_44])).
cnf(c_0_72,plain,(k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_73,negated_conjecture,(k2_relset_1(esk1_0,k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_66, c_0_71])).
cnf(c_0_74,plain,(k2_relset_1(X1,X2,k1_xboole_0)=k1_xboole_0|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_72, c_0_64])).
cnf(c_0_75,negated_conjecture,(~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_76,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_71]), c_0_75]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 77
# Proof object clause steps            : 59
# Proof object formula steps           : 18
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 183
# ...of these trivial                  : 0
# ...subsumed                          : 38
# ...remaining for further processing  : 144
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 49
# Generated clauses                    : 272
# ...of the previous two non-trivial   : 278
# Contextual simplify-reflections      : 29
# Paramodulations                      : 259
# Factorizations                       : 0
# Equation resolutions                 : 15
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 43
# Current number of unprocessed clauses: 38
# ...number of literals in the above   : 162
# Current number of archived formulas  : 0
# Current number of archived clauses   : 83
# Clause-clause subsumption calls (NU) : 1315
# Rec. Clause-clause subsumption calls : 706
# Non-unit clause-clause subsumptions  : 68
# Unit Clause-clause subsumption calls : 41
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8469

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45476 pages
