# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_funct_1, conjecture, (![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0)), file('out/t16_funct_1__10_premises__rerun.E_input', t16_funct_1)).
fof(s3_funct_1__e7_25__funct_1, axiom, (![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=1)))), file('out/t16_funct_1__10_premises__rerun.E_input', s3_funct_1__e7_25__funct_1)).
fof(d5_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4)))))), file('out/t16_funct_1__10_premises__rerun.E_input', d5_funct_1)).
fof(t64_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0))), file('out/t16_funct_1__10_premises__rerun.E_input', t64_relat_1)).
fof(s3_funct_1__e2_25__funct_1, axiom, (![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=k1_xboole_0)))), file('out/t16_funct_1__10_premises__rerun.E_input', s3_funct_1__e2_25__funct_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t16_funct_1__10_premises__rerun.E_input', d1_xboole_0)).
fof(cc1_funct_1, axiom, (![X1]:(v1_xboole_0(X1)=>v1_funct_1(X1))), file('out/t16_funct_1__10_premises__rerun.E_input', cc1_funct_1)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t16_funct_1__10_premises__rerun.E_input', fc1_xboole_0)).
fof(t60_relat_1, axiom, ((k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0)), file('out/t16_funct_1__10_premises__rerun.E_input', t60_relat_1)).
fof(t2_tarski, axiom, (![X1]:![X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2)), file('out/t16_funct_1__10_premises__rerun.E_input', t2_tarski)).
fof(spc1_boole, axiom, (~(v1_xboole_0(1))), file('out/t16_funct_1__10_premises__rerun.E_input', spc1_boole)).
fof(c_0_11, negated_conjecture, (~(![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0))), inference(assume_negation,[status(cth)],[t16_funct_1])).
fof(c_0_12, negated_conjecture, (![X5]:![X6]:(((~v1_relat_1(X5)|~v1_funct_1(X5))|((~v1_relat_1(X6)|~v1_funct_1(X6))|((k1_relat_1(X5)!=esk7_0|k1_relat_1(X6)!=esk7_0)|X5=X6)))&esk7_0!=k1_xboole_0)), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])])).
fof(c_0_13, plain, (![X4]:![X6]:((v1_relat_1(esk8_1(X4))&v1_funct_1(esk8_1(X4)))&(k1_relat_1(esk8_1(X4))=X4&(~r2_hidden(X6,X4)|k1_funct_1(esk8_1(X4),X6)=1)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e7_25__funct_1])])])])])])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X6]:![X11]:((((((r2_hidden(esk1_3(X5,X6,X7),k1_relat_1(X5))|~r2_hidden(X7,X6))|X6!=k2_relat_1(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((X7=k1_funct_1(X5,esk1_3(X5,X6,X7))|~r2_hidden(X7,X6))|X6!=k2_relat_1(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(X9,k1_relat_1(X5))|X7!=k1_funct_1(X5,X9))|r2_hidden(X7,X6))|X6!=k2_relat_1(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(esk2_2(X5,X6),X6)|(~r2_hidden(X11,k1_relat_1(X5))|esk2_2(X5,X6)!=k1_funct_1(X5,X11)))|X6=k2_relat_1(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&((((r2_hidden(esk3_2(X5,X6),k1_relat_1(X5))|r2_hidden(esk2_2(X5,X6),X6))|X6=k2_relat_1(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((esk2_2(X5,X6)=k1_funct_1(X5,esk3_2(X5,X6))|r2_hidden(esk2_2(X5,X6),X6))|X6=k2_relat_1(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])])).
fof(c_0_15, plain, (![X2]:(((k1_relat_1(X2)!=k1_xboole_0|X2=k1_xboole_0)|~v1_relat_1(X2))&((k2_relat_1(X2)!=k1_xboole_0|X2=k1_xboole_0)|~v1_relat_1(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
fof(c_0_16, plain, (![X4]:![X6]:((v1_relat_1(esk4_1(X4))&v1_funct_1(esk4_1(X4)))&(k1_relat_1(esk4_1(X4))=X4&(~r2_hidden(X6,X4)|k1_funct_1(esk4_1(X4),X6)=k1_xboole_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_25__funct_1])])])])])])).
cnf(c_0_17,negated_conjecture,(X1=X2|k1_relat_1(X2)!=esk7_0|k1_relat_1(X1)!=esk7_0|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(k1_relat_1(esk8_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(v1_funct_1(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(v1_relat_1(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_21, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk5_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_22,plain,(r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k2_relat_1(X1)|X3!=k1_funct_1(X1,X4)|~r2_hidden(X4,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_23, plain, (![X2]:(~v1_xboole_0(X2)|v1_funct_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_1])])).
cnf(c_0_24,plain,(X1=k1_xboole_0|~v1_relat_1(X1)|k1_relat_1(X1)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(k1_relat_1(esk4_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26,plain,(v1_relat_1(esk4_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27,negated_conjecture,(X1=esk8_1(esk7_0)|k1_relat_1(X1)!=esk7_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])])])).
cnf(c_0_28,plain,(v1_funct_1(esk4_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,plain,(r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_22])])).
cnf(c_0_31,plain,(v1_funct_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_33,plain,(esk4_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])])).
cnf(c_0_34,plain,(k1_funct_1(esk8_1(X1),X2)=1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35,negated_conjecture,(esk8_1(esk7_0)=esk4_1(esk7_0)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_25]), c_0_28]), c_0_26])])])).
cnf(c_0_36,plain,(~v1_xboole_0(k2_relat_1(X1))|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37,plain,(k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_38,plain,(k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_39,plain,(v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40,plain,(v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_26, c_0_33])).
fof(c_0_41, plain, (![X4]:![X5]:(((~r2_hidden(esk6_2(X4,X5),X4)|~r2_hidden(esk6_2(X4,X5),X5))|X4=X5)&((r2_hidden(esk6_2(X4,X5),X4)|r2_hidden(esk6_2(X4,X5),X5))|X4=X5))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])])])).
cnf(c_0_42,plain,(k1_funct_1(esk4_1(X1),X2)=k1_xboole_0|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_43,negated_conjecture,(k1_funct_1(esk4_1(esk7_0),X1)=1|~r2_hidden(X1,esk7_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44,plain,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_32]), c_0_39]), c_0_40])])).
cnf(c_0_45,plain,(X1=X2|r2_hidden(esk6_2(X1,X2),X2)|r2_hidden(esk6_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_46, plain, (~v1_xboole_0(1)), inference(fof_simplification,[status(thm)],[spc1_boole])).
cnf(c_0_47,negated_conjecture,(1=k1_xboole_0|~r2_hidden(X1,esk7_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48,plain,(k1_xboole_0=X1|r2_hidden(esk6_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49,negated_conjecture,(esk7_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50,plain,(~v1_xboole_0(1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_51,negated_conjecture,(1=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_52,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_51]), c_0_32])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 32
# Proof object formula steps           : 21
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 126
# ...of these trivial                  : 0
# ...subsumed                          : 38
# ...remaining for further processing  : 88
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 5
# Generated clauses                    : 147
# ...of the previous two non-trivial   : 108
# Contextual simplify-reflections      : 6
# Paramodulations                      : 134
# Factorizations                       : 2
# Equation resolutions                 : 12
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 25
# ...number of literals in the above   : 78
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 483
# Rec. Clause-clause subsumption calls : 256
# Non-unit clause-clause subsumptions  : 29
# Unit Clause-clause subsumption calls : 44
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3599

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45488 pages
