# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_funct_1, conjecture, (![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0)), file('out/t16_funct_1__6_premises__rerun.E_input', t16_funct_1)).
fof(s3_funct_1__e7_25__funct_1, axiom, (![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=1)))), file('out/t16_funct_1__6_premises__rerun.E_input', s3_funct_1__e7_25__funct_1)).
fof(d2_tarski, axiom, (![X1]:![X2]:![X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2)))), file('out/t16_funct_1__6_premises__rerun.E_input', d2_tarski)).
fof(s3_funct_1__e2_24__funct_1, axiom, (![X1]:![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>k1_funct_1(X3,X4)=X2)))), file('out/t16_funct_1__6_premises__rerun.E_input', s3_funct_1__e2_24__funct_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t16_funct_1__6_premises__rerun.E_input', d1_xboole_0)).
fof(d4_relat_1, axiom, (![X1]:![X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t16_funct_1__6_premises__rerun.E_input', d4_relat_1)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t16_funct_1__6_premises__rerun.E_input', fc1_xboole_0)).
fof(c_0_7, negated_conjecture, (~(![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0))), inference(assume_negation,[status(cth)],[t16_funct_1])).
fof(c_0_8, negated_conjecture, (![X5]:![X6]:(((~v1_relat_1(X5)|~v1_funct_1(X5))|((~v1_relat_1(X6)|~v1_funct_1(X6))|((k1_relat_1(X5)!=esk1_0|k1_relat_1(X6)!=esk1_0)|X5=X6)))&esk1_0!=k1_xboole_0)), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
fof(c_0_9, plain, (![X4]:![X6]:((v1_relat_1(esk2_1(X4))&v1_funct_1(esk2_1(X4)))&(k1_relat_1(esk2_1(X4))=X4&(~r2_hidden(X6,X4)|k1_funct_1(esk2_1(X4),X6)=1)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e7_25__funct_1])])])])])])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(X8=X5|X8=X6))|X7!=k2_tarski(X5,X6))&(((X8!=X5|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))&((X8!=X6|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))))&((((esk3_3(X5,X6,X7)!=X5|~r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6))&((esk3_3(X5,X6,X7)!=X6|~r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6)))&((r2_hidden(esk3_3(X5,X6,X7),X7)|(esk3_3(X5,X6,X7)=X5|esk3_3(X5,X6,X7)=X6))|X7=k2_tarski(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])])).
cnf(c_0_11,negated_conjecture,(X1=X2|k1_relat_1(X2)!=esk1_0|k1_relat_1(X1)!=esk1_0|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(v1_funct_1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(k1_relat_1(esk2_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(v1_relat_1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_15, plain, (![X5]:![X6]:![X8]:((v1_relat_1(esk4_2(X5,X6))&v1_funct_1(esk4_2(X5,X6)))&(k1_relat_1(esk4_2(X5,X6))=X5&(~r2_hidden(X8,X5)|k1_funct_1(esk4_2(X5,X6),X8)=X6)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_24__funct_1])])])])])])).
fof(c_0_16, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk5_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_17,plain,(r2_hidden(X4,X1)|X1!=k2_tarski(X2,X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,negated_conjecture,(X1=esk2_1(esk1_0)|k1_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])])])).
cnf(c_0_19,plain,(v1_funct_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,plain,(k1_relat_1(esk4_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(v1_relat_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(r2_hidden(X1,k2_tarski(X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_17])])).
cnf(c_0_24,plain,(k1_funct_1(esk4_2(X1,X2),X3)=X2|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,negated_conjecture,(esk4_2(esk1_0,X1)=esk2_1(esk1_0)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])])])).
cnf(c_0_26,plain,(~v1_xboole_0(k2_tarski(X1,X2))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27,negated_conjecture,(k1_funct_1(esk2_1(esk1_0),X1)=X2|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28,negated_conjecture,(~v1_xboole_0(k1_funct_1(esk2_1(esk1_0),X1))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_29, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk6_3(X5,X6,X7)),X5))|X6!=k1_relat_1(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k1_relat_1(X5)))&(((~r2_hidden(esk7_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk7_2(X5,X6),X11),X5))|X6=k1_relat_1(X5))&((r2_hidden(esk7_2(X5,X6),X6)|r2_hidden(k4_tarski(esk7_2(X5,X6),esk8_2(X5,X6)),X5))|X6=k1_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])])).
cnf(c_0_30,negated_conjecture,(~v1_xboole_0(X1)|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_27])).
cnf(c_0_31,plain,(v1_xboole_0(X1)|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_32,plain,(X1=k1_relat_1(X2)|r2_hidden(k4_tarski(esk7_2(X2,X1),esk8_2(X2,X1)),X2)|r2_hidden(esk7_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33,negated_conjecture,(v1_xboole_0(esk1_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_35,negated_conjecture,(X1=k1_relat_1(esk1_0)|r2_hidden(esk7_2(esk1_0,X1),X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_30, c_0_32])).
cnf(c_0_36,negated_conjecture,(v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37,negated_conjecture,(X1=k1_relat_1(esk1_0)|r2_hidden(esk7_2(esk1_0,X1),X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_38,negated_conjecture,(k1_relat_1(esk1_0)=esk1_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_37])).
cnf(c_0_39,negated_conjecture,(k1_relat_1(esk1_0)=esk1_0), inference(spm,[status(thm)],[c_0_38, c_0_36])).
cnf(c_0_40,negated_conjecture,(X1=esk1_0|r2_hidden(esk7_2(esk1_0,X1),X1)), inference(rw,[status(thm)],[c_0_37, c_0_39])).
cnf(c_0_41,negated_conjecture,(X1=esk1_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_22, c_0_40])).
cnf(c_0_42,negated_conjecture,(esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_43,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_34]), c_0_42]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 30
# Proof object formula steps           : 14
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 216
# ...of these trivial                  : 2
# ...subsumed                          : 37
# ...remaining for further processing  : 177
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 54
# Generated clauses                    : 623
# ...of the previous two non-trivial   : 599
# Contextual simplify-reflections      : 0
# Paramodulations                      : 610
# Factorizations                       : 2
# Equation resolutions                 : 13
# Current number of processed clauses  : 90
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 62
# Current number of unprocessed clauses: 140
# ...number of literals in the above   : 312
# Current number of archived formulas  : 0
# Current number of archived clauses   : 82
# Clause-clause subsumption calls (NU) : 2002
# Rec. Clause-clause subsumption calls : 1797
# Non-unit clause-clause subsumptions  : 26
# Unit Clause-clause subsumption calls : 138
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 156
# BW rewrite match successes           : 16
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9134

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 45488 pages
