# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d12_xtuple_0, axiom, (![X1]:![X2]:(X2=k9_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', d12_xtuple_0)).
fof(d5_tarski, axiom, (![X1]:![X2]:k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', d5_tarski)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(t6_boole, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', t6_boole)).
fof(fc10_relat_1, axiom, (![X1]:(v1_xboole_0(X1)=>v1_xboole_0(k9_xtuple_0(X1)))), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', fc10_relat_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', d1_xboole_0)).
fof(t169_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k9_relat_1(X3,X1))))), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', t169_relat_1)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', fc1_xboole_0)).
fof(t170_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k9_xtuple_0(X2))<=>~(k9_relat_1(X2,X1)=k1_xboole_0)))), file('out/t170_relat_1__47_premises__rerun__rerun.E_input', t170_relat_1)).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk1_3(X5,X6,X7)),X5))|X6!=k9_xtuple_0(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k9_xtuple_0(X5)))&(((~r2_hidden(esk2_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk2_2(X5,X6),X11),X5))|X6=k9_xtuple_0(X5))&((r2_hidden(esk2_2(X5,X6),X6)|r2_hidden(k4_tarski(esk2_2(X5,X6),esk3_2(X5,X6)),X5))|X6=k9_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_xtuple_0])])])])])])])).
fof(c_0_10, plain, (![X3]:![X4]:k4_tarski(X3,X4)=k2_tarski(k2_tarski(X3,X4),k1_tarski(X3))), inference(variable_rename,[status(thm)],[d5_tarski])).
cnf(c_0_11,plain,(r2_hidden(k4_tarski(X3,esk1_3(X2,X1,X3)),X2)|X1!=k9_xtuple_0(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_14, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_15, plain, (![X2]:(~v1_xboole_0(X2)|v1_xboole_0(k9_xtuple_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc10_relat_1])])).
fof(c_0_16, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk4_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_17,plain,(r2_hidden(k2_tarski(k2_tarski(X3,esk1_3(X2,X1,X3)),k1_tarski(X3)),X2)|X1!=k9_xtuple_0(X2)|~r2_hidden(X3,X1)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(r2_hidden(X3,X1)|X1!=k9_xtuple_0(X2)|~r2_hidden(k4_tarski(X3,X4),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:(((~r2_hidden(k4_tarski(X4,X5),X6)|r2_hidden(X5,k9_relat_1(X6,X4)))|~v1_relat_1(X6))&((~r2_hidden(X5,k9_relat_1(X6,X4))|r2_hidden(k4_tarski(X4,X5),X6))|~v1_relat_1(X6)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])])).
cnf(c_0_23,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,plain,(r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk1_3(X2,k9_xtuple_0(X2),X1))),X2)|~r2_hidden(X1,k9_xtuple_0(X2))), inference(er,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18])])).
cnf(c_0_25,plain,(k9_xtuple_0(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_27,plain,(X1=k9_xtuple_0(X2)|r2_hidden(k4_tarski(esk2_2(X2,X1),esk3_2(X2,X1)),X2)|r2_hidden(esk2_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28,plain,(r2_hidden(X3,X1)|X1!=k9_xtuple_0(X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),X2)), inference(rw,[status(thm)],[c_0_21, c_0_12])).
cnf(c_0_29,plain,(r2_hidden(k4_tarski(X2,X3),X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31,plain,(k9_xtuple_0(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32,plain,(X1=k9_xtuple_0(X2)|r2_hidden(esk2_2(X2,X1),X1)|r2_hidden(k2_tarski(k2_tarski(esk2_2(X2,X1),esk3_2(X2,X1)),k1_tarski(esk2_2(X2,X1))),X2)), inference(rw,[status(thm)],[c_0_27, c_0_12])).
cnf(c_0_33,plain,(r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),X2)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_34,plain,(r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,X2))), inference(rw,[status(thm)],[c_0_29, c_0_12])).
cnf(c_0_35,plain,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_26])])).
cnf(c_0_36,plain,(X1=k9_xtuple_0(X2)|r2_hidden(k2_tarski(k1_tarski(esk2_2(X2,X1)),k2_tarski(esk2_2(X2,X1),esk3_2(X2,X1))),X2)|r2_hidden(esk2_2(X2,X1),X1)), inference(rw,[status(thm)],[c_0_32, c_0_18])).
fof(c_0_37, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k9_xtuple_0(X2))<=>~(k9_relat_1(X2,X1)=k1_xboole_0))))), inference(assume_negation,[status(cth)],[t170_relat_1])).
cnf(c_0_38,plain,(r2_hidden(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)|~r2_hidden(X3,k9_relat_1(X2,X1))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39,plain,(X1=k1_xboole_0|r2_hidden(esk2_2(k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_31])).
fof(c_0_40, negated_conjecture, ((v1_relat_1(esk6_0)&((~r2_hidden(esk5_0,k9_xtuple_0(esk6_0))|k9_relat_1(esk6_0,esk5_0)=k1_xboole_0)&(r2_hidden(esk5_0,k9_xtuple_0(esk6_0))|k9_relat_1(esk6_0,esk5_0)!=k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_37])])])])).
cnf(c_0_41,plain,(r2_hidden(X2,k9_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X3,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_42,plain,(k9_relat_1(X1,X2)=k1_xboole_0|r2_hidden(X2,k9_xtuple_0(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43,negated_conjecture,(v1_relat_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44,plain,(r2_hidden(X2,k9_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(k2_tarski(k2_tarski(X3,X2),k1_tarski(X3)),X1)), inference(rw,[status(thm)],[c_0_41, c_0_12])).
cnf(c_0_45,negated_conjecture,(r2_hidden(esk5_0,k9_xtuple_0(esk6_0))|k9_relat_1(esk6_0,esk5_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46,negated_conjecture,(k9_relat_1(esk6_0,X1)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk6_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47,plain,(r2_hidden(X1,k9_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),X2)), inference(spm,[status(thm)],[c_0_44, c_0_18])).
cnf(c_0_48,negated_conjecture,(k9_relat_1(esk6_0,esk5_0)=k1_xboole_0|~r2_hidden(esk5_0,k9_xtuple_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49,negated_conjecture,(r2_hidden(esk5_0,k9_xtuple_0(esk6_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50,plain,(r2_hidden(esk1_3(X1,k9_xtuple_0(X1),X2),k9_relat_1(X1,X2))|~v1_relat_1(X1)|~r2_hidden(X2,k9_xtuple_0(X1))), inference(spm,[status(thm)],[c_0_47, c_0_24])).
cnf(c_0_51,negated_conjecture,(k9_relat_1(esk6_0,esk5_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49])])).
cnf(c_0_52,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_43]), c_0_49])]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 35
# Proof object formula steps           : 18
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 15
# Processed clauses                    : 1029
# ...of these trivial                  : 0
# ...subsumed                          : 872
# ...remaining for further processing  : 157
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 3241
# ...of the previous two non-trivial   : 2602
# Contextual simplify-reflections      : 424
# Paramodulations                      : 3237
# Factorizations                       : 2
# Equation resolutions                 : 2
# Current number of processed clauses  : 137
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 128
# Current number of unprocessed clauses: 1603
# ...number of literals in the above   : 5415
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 5851
# Rec. Clause-clause subsumption calls : 5090
# Non-unit clause-clause subsumptions  : 883
# Unit Clause-clause subsumption calls : 30
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 49591

# -------------------------------------------------
# User time                : 0.048 s
# System time              : 0.008 s
# Total time               : 0.056 s
# Maximum resident set size: 45320 pages
