# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t173_relat_1__7_premises__rerun.E_input', t3_xboole_0)).
fof(t166_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2)))))), file('out/t173_relat_1__7_premises__rerun.E_input', t166_relat_1)).
fof(t7_xboole_0, axiom, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1))))), file('out/t173_relat_1__7_premises__rerun.E_input', t7_xboole_0)).
fof(t173_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), file('out/t173_relat_1__7_premises__rerun.E_input', t173_relat_1)).
fof(t152_zfmisc_1, axiom, (![X1]:![X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2)))), file('out/t173_relat_1__7_premises__rerun.E_input', t152_zfmisc_1)).
fof(t113_zfmisc_1, axiom, (![X1]:![X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), file('out/t173_relat_1__7_premises__rerun.E_input', t113_zfmisc_1)).
fof(symmetry_r1_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1))), file('out/t173_relat_1__7_premises__rerun.E_input', symmetry_r1_xboole_0)).
fof(d5_relat_1, axiom, (![X1]:![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1)))), file('out/t173_relat_1__7_premises__rerun.E_input', d5_relat_1)).
fof(c_0_8, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk1_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk1_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X9]:((((r2_hidden(esk2_3(X5,X6,X7),k2_relat_1(X7))|~r2_hidden(X5,k10_relat_1(X7,X6)))|~v1_relat_1(X7))&(((r2_hidden(k4_tarski(X5,esk2_3(X5,X6,X7)),X7)|~r2_hidden(X5,k10_relat_1(X7,X6)))|~v1_relat_1(X7))&((r2_hidden(esk2_3(X5,X6,X7),X6)|~r2_hidden(X5,k10_relat_1(X7,X6)))|~v1_relat_1(X7))))&(((~r2_hidden(X9,k2_relat_1(X7))|(~r2_hidden(k4_tarski(X5,X9),X7)|~r2_hidden(X9,X6)))|r2_hidden(X5,k10_relat_1(X7,X6)))|~v1_relat_1(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])])])).
fof(c_0_10, plain, (![X3]:(X3=k1_xboole_0|r2_hidden(esk3_1(X3),X3))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])])])).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))))), inference(assume_negation,[status(cth)],[t173_relat_1])).
fof(c_0_12, plain, (![X3]:![X4]:~r2_hidden(X3,k2_zfmisc_1(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
fof(c_0_13, plain, (![X3]:![X4]:![X3]:![X4]:((k2_zfmisc_1(X3,X4)!=k1_xboole_0|(X3=k1_xboole_0|X4=k1_xboole_0))&((X3!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)&(X4!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])])])).
fof(c_0_14, plain, ((~epred1_0<=>![X2]:X2!=k1_xboole_0)), introduced(definition)).
fof(c_0_15, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|r1_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_16,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17,plain,(r2_hidden(esk2_3(X2,X3,X1),k2_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k10_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,plain,(r2_hidden(esk2_3(X2,X3,X1),X3)|~v1_relat_1(X1)|~r2_hidden(X2,k10_relat_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19,plain,(r2_hidden(esk3_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_20, negated_conjecture, ((v1_relat_1(esk5_0)&((k10_relat_1(esk5_0,esk4_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk5_0),esk4_0))&(k10_relat_1(esk5_0,esk4_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk5_0),esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_21, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(esk6_3(X5,X6,X7),X7),X5))|X6!=k2_relat_1(X5))&((~r2_hidden(k4_tarski(X9,X7),X5)|r2_hidden(X7,X6))|X6!=k2_relat_1(X5)))&(((~r2_hidden(esk7_2(X5,X6),X6)|~r2_hidden(k4_tarski(X11,esk7_2(X5,X6)),X5))|X6=k2_relat_1(X5))&((r2_hidden(esk7_2(X5,X6),X6)|r2_hidden(k4_tarski(esk8_2(X5,X6),esk7_2(X5,X6)),X5))|X6=k2_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])])).
fof(c_0_22, plain, ((~epred2_0<=>![X1]:~r2_hidden(X1,k1_xboole_0))), introduced(definition)).
cnf(c_0_23,plain,(~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,plain,(epred1_0|X1!=k1_xboole_0), inference(split_equiv,[status(thm)],[c_0_14])).
cnf(c_0_26,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28,plain,(~v1_relat_1(X1)|~r2_hidden(esk2_3(X2,X3,X1),X4)|~r2_hidden(X2,k10_relat_1(X1,X3))|~r1_xboole_0(X4,k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_29,plain,(k10_relat_1(X1,X2)=k1_xboole_0|r2_hidden(esk2_3(esk3_1(k10_relat_1(X1,X2)),X2,X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_30,negated_conjecture,(r1_xboole_0(k2_relat_1(esk5_0),esk4_0)|k10_relat_1(esk5_0,esk4_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31,plain,(r2_hidden(k4_tarski(esk6_3(X2,X1,X3),X3),X2)|X1!=k2_relat_1(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32,plain,(~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_14]), c_0_22])).
cnf(c_0_33,plain,(epred1_0), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_35,negated_conjecture,(~r1_xboole_0(k2_relat_1(esk5_0),esk4_0)|k10_relat_1(esk5_0,esk4_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36,plain,(r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_37,plain,(k10_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(X2,k2_relat_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_19])).
cnf(c_0_38,negated_conjecture,(k10_relat_1(esk5_0,esk4_0)=k1_xboole_0|r1_xboole_0(esk4_0,k2_relat_1(esk5_0))), inference(spm,[status(thm)],[c_0_26, c_0_30])).
cnf(c_0_39,negated_conjecture,(v1_relat_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_40,plain,(r2_hidden(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X4,X3)|~r2_hidden(k4_tarski(X2,X4),X1)|~r2_hidden(X4,k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41,plain,(r2_hidden(k4_tarski(esk6_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_42,plain,(~epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33])])).
cnf(c_0_43,plain,(r2_hidden(esk1_2(X1,X2),X2)|r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_26, c_0_34])).
cnf(c_0_44,negated_conjecture,(r2_hidden(esk1_2(esk4_0,k2_relat_1(esk5_0)),esk4_0)|k10_relat_1(esk5_0,esk4_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45,negated_conjecture,(k10_relat_1(esk5_0,esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
cnf(c_0_46,plain,(r2_hidden(esk6_3(X1,k2_relat_1(X1),X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47,plain,(~r2_hidden(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_22]), c_0_42])).
cnf(c_0_48,plain,(r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X2,X1),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_16, c_0_43])).
cnf(c_0_49,negated_conjecture,(r2_hidden(esk1_2(esk4_0,k2_relat_1(esk5_0)),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45])])).
cnf(c_0_50,negated_conjecture,(~r1_xboole_0(k2_relat_1(esk5_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_45])])).
cnf(c_0_51,negated_conjecture,(~r2_hidden(X1,k2_relat_1(esk5_0))|~r2_hidden(X1,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_45]), c_0_39])]), c_0_47])).
cnf(c_0_52,negated_conjecture,(~r1_xboole_0(esk4_0,k2_relat_1(esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_53,negated_conjecture,(r1_xboole_0(X1,k2_relat_1(esk5_0))|~r2_hidden(esk1_2(X1,k2_relat_1(esk5_0)),esk4_0)), inference(spm,[status(thm)],[c_0_51, c_0_34])).
cnf(c_0_54,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 36
# Proof object formula steps           : 19
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 3320
# ...of these trivial                  : 29
# ...subsumed                          : 2642
# ...remaining for further processing  : 649
# Other redundant clauses eliminated   : 101
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 41
# Generated clauses                    : 51949
# ...of the previous two non-trivial   : 38095
# Contextual simplify-reflections      : 1047
# Paramodulations                      : 51760
# Factorizations                       : 60
# Equation resolutions                 : 126
# Current number of processed clauses  : 583
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 552
# Current number of unprocessed clauses: 34417
# ...number of literals in the above   : 225808
# Current number of archived formulas  : 0
# Current number of archived clauses   : 65
# Clause-clause subsumption calls (NU) : 61152
# Rec. Clause-clause subsumption calls : 48063
# Non-unit clause-clause subsumptions  : 3073
# Unit Clause-clause subsumption calls : 98
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1305609

# -------------------------------------------------
# User time                : 1.076 s
# System time              : 0.016 s
# Total time               : 1.092 s
# Maximum resident set size: 45460 pages
