# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t74_xboole_1, axiom, (![X1]:![X2]:![X3]:~((~(r1_xboole_0(X1,k3_xboole_0(X2,X3)))&r1_xboole_0(X1,X2)))), file('out/t174_funct_1__20_premises__rerun.E_input', t74_xboole_1)).
fof(t89_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))), file('out/t174_funct_1__20_premises__rerun.E_input', t89_xboole_1)).
fof(t2_boole, axiom, (![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), file('out/t174_funct_1__20_premises__rerun.E_input', t2_boole)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t174_funct_1__20_premises__rerun.E_input', t28_xboole_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t174_funct_1__20_premises__rerun.E_input', t7_xboole_1)).
fof(t111_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k4_xboole_0(X1,X3),X2)), file('out/t174_funct_1__20_premises__rerun.E_input', t111_xboole_1)).
fof(t52_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(X1,k4_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X3))), file('out/t174_funct_1__20_premises__rerun.E_input', t52_xboole_1)).
fof(t88_xboole_1, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1)), file('out/t174_funct_1__20_premises__rerun.E_input', t88_xboole_1)).
fof(t43_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3))), file('out/t174_funct_1__20_premises__rerun.E_input', t43_xboole_1)).
fof(t107_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k5_xboole_0(X2,X3))<=>(r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,k3_xboole_0(X2,X3))))), file('out/t174_funct_1__20_premises__rerun.E_input', t107_xboole_1)).
fof(t93_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k2_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), file('out/t174_funct_1__20_premises__rerun.E_input', t93_xboole_1)).
fof(t84_xboole_1, axiom, (![X1]:![X2]:![X3]:~((~(r1_xboole_0(X1,X2))&(r1_xboole_0(X1,X3)&r1_xboole_0(X1,k4_xboole_0(X2,X3)))))), file('out/t174_funct_1__20_premises__rerun.E_input', t84_xboole_1)).
fof(t127_zfmisc_1, axiom, (![X1]:![X2]:![X3]:![X4]:((r1_xboole_0(X1,X2)|r1_xboole_0(X3,X4))=>r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)))), file('out/t174_funct_1__20_premises__rerun.E_input', t127_zfmisc_1)).
fof(t22_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>k3_xboole_0(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))=X1)), file('out/t174_funct_1__20_premises__rerun.E_input', t22_relat_1)).
fof(s3_funct_1__e2_24__funct_1, axiom, (![X1]:![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>k1_funct_1(X3,X4)=X2)))), file('out/t174_funct_1__20_premises__rerun.E_input', s3_funct_1__e2_24__funct_1)).
fof(t174_funct_1, conjecture, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>v5_funct_1(k1_xboole_0,X1))), file('out/t174_funct_1__20_premises__rerun.E_input', t174_funct_1)).
fof(t64_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))<=>(r2_hidden(X1,X2)&~(X1=X3)))), file('out/t174_funct_1__20_premises__rerun.E_input', t64_zfmisc_1)).
fof(t76_enumset1, axiom, (![X1]:k1_enumset1(X1,X1,X1)=k1_tarski(X1)), file('out/t174_funct_1__20_premises__rerun.E_input', t76_enumset1)).
fof(t137_enumset1, axiom, (![X1]:![X2]:![X3]:k2_xboole_0(k2_tarski(X2,X1),k2_tarski(X3,X1))=k1_enumset1(X1,X2,X3)), file('out/t174_funct_1__20_premises__rerun.E_input', t137_enumset1)).
fof(fc1_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X1)=>v1_relat_1(k3_xboole_0(X1,X2)))), file('out/t174_funct_1__20_premises__rerun.E_input', fc1_relat_1)).
fof(d20_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(v5_funct_1(X2,X1)<=>![X3]:(r2_hidden(X3,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X3),k1_funct_1(X1,X3))))))), file('out/t174_funct_1__20_premises__rerun.E_input', d20_funct_1)).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:(r1_xboole_0(X4,k3_xboole_0(X5,X6))|~r1_xboole_0(X4,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t74_xboole_1])])])])])).
fof(c_0_22, plain, (![X3]:![X4]:r1_xboole_0(k3_xboole_0(X3,X4),k4_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t89_xboole_1])).
cnf(c_0_23,plain,(r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_24,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, (![X2]:k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_26, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_27, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_28, plain, (![X4]:![X5]:![X6]:k4_xboole_0(k3_xboole_0(X4,X5),k3_xboole_0(X6,X5))=k3_xboole_0(k4_xboole_0(X4,X6),X5)), inference(variable_rename,[status(thm)],[t111_xboole_1])).
cnf(c_0_29,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(k4_xboole_0(X1,X2),X3))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_33, plain, (![X4]:![X5]:![X6]:k4_xboole_0(X4,k4_xboole_0(X5,X6))=k2_xboole_0(k4_xboole_0(X4,X5),k3_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[t52_xboole_1])).
cnf(c_0_34,plain,(k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k4_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_35, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|k4_xboole_0(k2_xboole_0(X3,X4),X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t88_xboole_1])])).
cnf(c_0_36,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37,plain,(k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:(~r1_tarski(X4,k2_xboole_0(X5,X6))|r1_tarski(k4_xboole_0(X4,X5),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_39,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40,plain,(k4_xboole_0(k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_30]), c_0_30]), c_0_30])).
cnf(c_0_41,plain,(k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44,plain,(k2_xboole_0(k4_xboole_0(X1,k1_xboole_0),k1_xboole_0)=k4_xboole_0(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_30])).
cnf(c_0_45,plain,(k4_xboole_0(k2_xboole_0(X1,k1_xboole_0),k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46,plain,(r1_tarski(k4_xboole_0(X1,X1),X2)), inference(spm,[status(thm)],[c_0_43, c_0_32])).
cnf(c_0_47,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_48,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_46, c_0_40])).
cnf(c_0_49,plain,(r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_47])).
cnf(c_0_50,plain,(r1_tarski(k4_xboole_0(k1_xboole_0,X1),X2)), inference(spm,[status(thm)],[c_0_43, c_0_48])).
cnf(c_0_51,plain,(r1_tarski(k4_xboole_0(k4_xboole_0(k1_xboole_0,X1),X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_52, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(((r1_tarski(X4,k2_xboole_0(X5,X6))|~r1_tarski(X4,k5_xboole_0(X5,X6)))&(r1_xboole_0(X4,k3_xboole_0(X5,X6))|~r1_tarski(X4,k5_xboole_0(X5,X6))))&((~r1_tarski(X4,k2_xboole_0(X5,X6))|~r1_xboole_0(X4,k3_xboole_0(X5,X6)))|r1_tarski(X4,k5_xboole_0(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t107_xboole_1])])])])])).
cnf(c_0_53,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_45, c_0_47])).
cnf(c_0_54,plain,(k4_xboole_0(k4_xboole_0(k1_xboole_0,X1),X2)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_51]), c_0_30])).
cnf(c_0_55,plain,(r1_tarski(X1,k5_xboole_0(X2,X3))|~r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_56,plain,(r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_57,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_32, c_0_47])).
fof(c_0_58, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k2_xboole_0(k5_xboole_0(X3,X4),k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t93_xboole_1])).
cnf(c_0_59,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_60,plain,(r1_tarski(k4_xboole_0(k4_xboole_0(X1,X1),X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_49, c_0_46])).
fof(c_0_61, plain, (![X4]:![X5]:![X6]:(r1_xboole_0(X4,X5)|(~r1_xboole_0(X4,X6)|~r1_xboole_0(X4,k4_xboole_0(X5,X6))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t84_xboole_1])])])])])).
cnf(c_0_62,plain,(r1_tarski(X1,k5_xboole_0(X1,X2))|~r1_xboole_0(X1,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_55, c_0_32])).
cnf(c_0_63,plain,(k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_37, c_0_47])).
cnf(c_0_64,plain,(r1_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65,plain,(k2_xboole_0(X1,X2)=k2_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_66,plain,(k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_59]), c_0_53]), c_0_53])).
cnf(c_0_67,plain,(k4_xboole_0(k4_xboole_0(X1,X1),X2)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_60]), c_0_30])).
fof(c_0_68, plain, (![X5]:![X6]:![X7]:![X8]:((~r1_xboole_0(X5,X6)|r1_xboole_0(k2_zfmisc_1(X5,X7),k2_zfmisc_1(X6,X8)))&(~r1_xboole_0(X7,X8)|r1_xboole_0(k2_zfmisc_1(X5,X7),k2_zfmisc_1(X6,X8))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t127_zfmisc_1])])])).
cnf(c_0_69,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70,plain,(r1_xboole_0(k1_xboole_0,k4_xboole_0(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_24, c_0_30])).
cnf(c_0_71,plain,(r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_40]), c_0_30])).
cnf(c_0_72,plain,(r1_tarski(X1,k5_xboole_0(X1,X1))|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_73,plain,(k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_64]), c_0_65])).
cnf(c_0_74,plain,(k2_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_66, c_0_37])).
cnf(c_0_75,plain,(k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_67])).
cnf(c_0_76,plain,(r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_77,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])])).
fof(c_0_78, plain, (![X2]:(~v1_relat_1(X2)|k3_xboole_0(X2,k2_zfmisc_1(k1_relat_1(X2),k2_relat_1(X2)))=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_relat_1])])).
fof(c_0_79, plain, (![X5]:![X6]:![X8]:((v1_relat_1(esk1_2(X5,X6))&v1_funct_1(esk1_2(X5,X6)))&(k1_relat_1(esk1_2(X5,X6))=X5&(~r2_hidden(X8,X5)|k1_funct_1(esk1_2(X5,X6),X8)=X6)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_24__funct_1])])])])])])).
cnf(c_0_80,plain,(r1_tarski(X1,k1_xboole_0)|~r1_xboole_0(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73]), c_0_74]), c_0_63]), c_0_75])).
cnf(c_0_81,plain,(r1_xboole_0(k2_zfmisc_1(k1_xboole_0,X1),k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
fof(c_0_82, negated_conjecture, (~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>v5_funct_1(k1_xboole_0,X1)))), inference(assume_negation,[status(cth)],[t174_funct_1])).
fof(c_0_83, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(((r2_hidden(X4,X5)|~r2_hidden(X4,k4_xboole_0(X5,k1_tarski(X6))))&(X4!=X6|~r2_hidden(X4,k4_xboole_0(X5,k1_tarski(X6)))))&((~r2_hidden(X4,X5)|X4=X6)|r2_hidden(X4,k4_xboole_0(X5,k1_tarski(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_zfmisc_1])])])])])])).
fof(c_0_84, plain, (![X2]:k1_enumset1(X2,X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t76_enumset1])).
fof(c_0_85, plain, (![X4]:![X5]:![X6]:k2_xboole_0(k2_tarski(X5,X4),k2_tarski(X6,X4))=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t137_enumset1])).
cnf(c_0_86,plain,(k3_xboole_0(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_87,plain,(v1_relat_1(esk1_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_88,plain,(k1_relat_1(esk1_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_89,plain,(r1_tarski(k2_zfmisc_1(k1_xboole_0,X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
fof(c_0_90, plain, (![X3]:![X4]:(~v1_relat_1(X3)|v1_relat_1(k3_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_relat_1])])])])).
fof(c_0_91, negated_conjecture, (((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&~v5_funct_1(k1_xboole_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_82])])])).
cnf(c_0_92,plain,(~r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))|X1!=X3), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_93,plain,(k1_enumset1(X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_94,plain,(k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X2))=k1_enumset1(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_95,plain,(k3_xboole_0(esk1_2(X1,X2),k2_zfmisc_1(X1,k2_relat_1(esk1_2(X1,X2))))=esk1_2(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])).
cnf(c_0_96,plain,(k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_89]), c_0_30])).
cnf(c_0_97,plain,(v1_relat_1(k3_xboole_0(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_98,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_99,plain,(X1!=X3|~r2_hidden(X1,k4_xboole_0(X2,k2_xboole_0(k2_tarski(X3,X3),k2_tarski(X3,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93]), c_0_94])).
fof(c_0_100, plain, (![X4]:![X5]:![X6]:((((~v5_funct_1(X5,X4)|(~r2_hidden(X6,k1_relat_1(X5))|r2_hidden(k1_funct_1(X5,X6),k1_funct_1(X4,X6))))|(~v1_relat_1(X5)|~v1_funct_1(X5)))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&((((r2_hidden(esk3_2(X4,X5),k1_relat_1(X5))|v5_funct_1(X5,X4))|(~v1_relat_1(X5)|~v1_funct_1(X5)))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&(((~r2_hidden(k1_funct_1(X5,esk3_2(X4,X5)),k1_funct_1(X4,esk3_2(X4,X5)))|v5_funct_1(X5,X4))|(~v1_relat_1(X5)|~v1_funct_1(X5)))|(~v1_relat_1(X4)|~v1_funct_1(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_funct_1])])])])])])])).
cnf(c_0_101,plain,(v1_funct_1(esk1_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_102,plain,(esk1_2(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_30])).
cnf(c_0_103,negated_conjecture,(v1_relat_1(k3_xboole_0(esk2_0,X1))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_104,plain,(~r2_hidden(X1,k4_xboole_0(X2,k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X1,X1))))), inference(er,[status(thm)],[c_0_99])).
cnf(c_0_105,plain,(k4_xboole_0(k2_xboole_0(k1_xboole_0,X1),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_77])).
cnf(c_0_106,plain,(v5_funct_1(X2,X1)|r2_hidden(esk3_2(X1,X2),k1_relat_1(X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_107,plain,(v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_108,negated_conjecture,(v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_103, c_0_30])).
cnf(c_0_109,plain,(k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_102])).
cnf(c_0_110,plain,(~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_111,plain,(v5_funct_1(k1_xboole_0,X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_108])]), c_0_109]), c_0_110])).
cnf(c_0_112,negated_conjecture,(v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_113,negated_conjecture,(~v5_funct_1(k1_xboole_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_114,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_98])]), c_0_113]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 115
# Proof object clause steps            : 72
# Proof object formula steps           : 43
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 42
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 31
# Processed clauses                    : 218
# ...of these trivial                  : 25
# ...subsumed                          : 27
# ...remaining for further processing  : 166
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 71
# Generated clauses                    : 1243
# ...of the previous two non-trivial   : 836
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1242
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 93
#    Positive orientable unit clauses  : 49
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 294
# ...number of literals in the above   : 415
# Current number of archived formulas  : 0
# Current number of archived clauses   : 74
# Clause-clause subsumption calls (NU) : 434
# Rec. Clause-clause subsumption calls : 305
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 175
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 115
# BW rewrite match successes           : 45
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15007

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.008 s
# Total time               : 0.028 s
# Maximum resident set size: 45428 pages
