# No SInE strategy applied
# Auto-Mode selected heuristic H_____011_C18_F1_PI_SE_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>k7_lattice3(k7_lattice3(X1))=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))), file('out/t17_waybel_7__10_premises__rerun__rerun.E_input', t8_lattice3)).
fof(abstractness_v1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))), file('out/t17_waybel_7__10_premises__rerun__rerun.E_input', abstractness_v1_orders_2)).
fof(t17_waybel_7, conjecture, (![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:((v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&(v1_lattice3(X2)&(v2_lattice3(X2)&l1_orders_2(X2))))))=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:((~(v1_xboole_0(X3))&(v1_waybel_0(X3,X1)&(v12_waybel_0(X3,X1)&(v1_waybel_7(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))))=>(~(v1_xboole_0(X3))&(v1_waybel_0(X3,X2)&(v12_waybel_0(X3,X2)&(v1_waybel_7(X3,X2)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))))))))))), file('out/t17_waybel_7__10_premises__rerun__rerun.E_input', t17_waybel_7)).
fof(dt_k7_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(k7_lattice3(X1))&l1_orders_2(k7_lattice3(X1))))), file('out/t17_waybel_7__10_premises__rerun__rerun.E_input', dt_k7_lattice3)).
fof(t16_waybel_7, axiom, (![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&(v1_waybel_7(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))))<=>(~(v1_xboole_0(X2))&(v2_waybel_0(X2,k7_lattice3(X1))&(v13_waybel_0(X2,k7_lattice3(X1))&(v2_waybel_7(X2,k7_lattice3(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))))))))), file('out/t17_waybel_7__10_premises__rerun__rerun.E_input', t16_waybel_7)).
fof(c_0_5, plain, (![X2]:(~l1_orders_2(X2)|k7_lattice3(k7_lattice3(X2))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_lattice3])])).
fof(c_0_6, plain, (![X2]:(~l1_orders_2(X2)|(~v1_orders_2(X2)|X2=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
fof(c_0_7, plain, (![X1]:(epred1_1(X1)<=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&(v1_waybel_7(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))))<=>(~(v1_xboole_0(X2))&(v2_waybel_0(X2,k7_lattice3(X1))&(v13_waybel_0(X2,k7_lattice3(X1))&(v2_waybel_7(X2,k7_lattice3(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))))))))), introduced(definition)).
fof(c_0_8, negated_conjecture, (~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:((v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&(v1_lattice3(X2)&(v2_lattice3(X2)&l1_orders_2(X2))))))=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:((~(v1_xboole_0(X3))&(v1_waybel_0(X3,X1)&(v12_waybel_0(X3,X1)&(v1_waybel_7(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))))=>(~(v1_xboole_0(X3))&(v1_waybel_0(X3,X2)&(v12_waybel_0(X3,X2)&(v1_waybel_7(X3,X2)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))))))))))), inference(assume_negation,[status(cth)],[t17_waybel_7])).
cnf(c_0_9,plain,(k7_lattice3(k7_lattice3(X1))=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, plain, (![X2]:((v1_orders_2(k7_lattice3(X2))|~l1_orders_2(X2))&(l1_orders_2(k7_lattice3(X2))|~l1_orders_2(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_lattice3])])])).
fof(c_0_12, axiom, (![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>epred1_1(X1))), inference(apply_def,[status(thm)],[t16_waybel_7, c_0_7])).
fof(c_0_13, negated_conjecture, (((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0))))))&((v3_orders_2(esk2_0)&(v4_orders_2(esk2_0)&(v5_orders_2(esk2_0)&(v1_lattice3(esk2_0)&(v2_lattice3(esk2_0)&l1_orders_2(esk2_0))))))&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&((~v1_xboole_0(esk3_0)&(v1_waybel_0(esk3_0,esk1_0)&(v12_waybel_0(esk3_0,esk1_0)&(v1_waybel_7(esk3_0,esk1_0)&m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))))))&(v1_xboole_0(esk3_0)|(~v1_waybel_0(esk3_0,esk2_0)|(~v12_waybel_0(esk3_0,esk2_0)|(~v1_waybel_7(esk3_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
cnf(c_0_14,plain,(k7_lattice3(k7_lattice3(X1))=X1|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_15,plain,(v1_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, (![X1]:(epred1_1(X1)=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&(v1_waybel_7(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))))<=>(~(v1_xboole_0(X2))&(v2_waybel_0(X2,k7_lattice3(X1))&(v13_waybel_0(X2,k7_lattice3(X1))&(v2_waybel_7(X2,k7_lattice3(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))))))))), inference(split_equiv,[status(thm)],[c_0_7])).
fof(c_0_17, plain, (![X2]:((~v3_orders_2(X2)|(~v4_orders_2(X2)|(~v5_orders_2(X2)|(~v1_lattice3(X2)|(~v2_lattice3(X2)|~l1_orders_2(X2))))))|epred1_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
cnf(c_0_18,negated_conjecture,(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k7_lattice3(k7_lattice3(k7_lattice3(X1)))=k7_lattice3(X1)|~l1_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21,plain,(l1_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_22, plain, (![X3]:![X4]:![X4]:((((~v1_xboole_0(X4)|(v1_xboole_0(X4)|(~v1_waybel_0(X4,X3)|(~v12_waybel_0(X4,X3)|(~v1_waybel_7(X4,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))))))|~epred1_1(X3))&(((v2_waybel_0(X4,k7_lattice3(X3))|(v1_xboole_0(X4)|(~v1_waybel_0(X4,X3)|(~v12_waybel_0(X4,X3)|(~v1_waybel_7(X4,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))))))|~epred1_1(X3))&(((v13_waybel_0(X4,k7_lattice3(X3))|(v1_xboole_0(X4)|(~v1_waybel_0(X4,X3)|(~v12_waybel_0(X4,X3)|(~v1_waybel_7(X4,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))))))|~epred1_1(X3))&(((v2_waybel_7(X4,k7_lattice3(X3))|(v1_xboole_0(X4)|(~v1_waybel_0(X4,X3)|(~v12_waybel_0(X4,X3)|(~v1_waybel_7(X4,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))))))|~epred1_1(X3))&((m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k7_lattice3(X3))))|(v1_xboole_0(X4)|(~v1_waybel_0(X4,X3)|(~v12_waybel_0(X4,X3)|(~v1_waybel_7(X4,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))))))|~epred1_1(X3))))))&(((~v1_xboole_0(X4)|(v1_xboole_0(X4)|(~v2_waybel_0(X4,k7_lattice3(X3))|(~v13_waybel_0(X4,k7_lattice3(X3))|(~v2_waybel_7(X4,k7_lattice3(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k7_lattice3(X3)))))))))|~epred1_1(X3))&(((v1_waybel_0(X4,X3)|(v1_xboole_0(X4)|(~v2_waybel_0(X4,k7_lattice3(X3))|(~v13_waybel_0(X4,k7_lattice3(X3))|(~v2_waybel_7(X4,k7_lattice3(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k7_lattice3(X3)))))))))|~epred1_1(X3))&(((v12_waybel_0(X4,X3)|(v1_xboole_0(X4)|(~v2_waybel_0(X4,k7_lattice3(X3))|(~v13_waybel_0(X4,k7_lattice3(X3))|(~v2_waybel_7(X4,k7_lattice3(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k7_lattice3(X3)))))))))|~epred1_1(X3))&(((v1_waybel_7(X4,X3)|(v1_xboole_0(X4)|(~v2_waybel_0(X4,k7_lattice3(X3))|(~v13_waybel_0(X4,k7_lattice3(X3))|(~v2_waybel_7(X4,k7_lattice3(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k7_lattice3(X3)))))))))|~epred1_1(X3))&((m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|(v1_xboole_0(X4)|(~v2_waybel_0(X4,k7_lattice3(X3))|(~v13_waybel_0(X4,k7_lattice3(X3))|(~v2_waybel_7(X4,k7_lattice3(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k7_lattice3(X3)))))))))|~epred1_1(X3)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])])])).
cnf(c_0_23,plain,(epred1_1(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,negated_conjecture,(v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_26,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_27,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_28,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29,negated_conjecture,(g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))=k7_lattice3(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_9]), c_0_19])])).
cnf(c_0_30,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31,plain,(k7_lattice3(k7_lattice3(k7_lattice3(X1)))=k7_lattice3(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_32,plain,(v1_xboole_0(X2)|v2_waybel_7(X2,k7_lattice3(X1))|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_7(X2,X1)|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34,negated_conjecture,(v1_waybel_7(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35,negated_conjecture,(v12_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36,negated_conjecture,(v1_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_37,negated_conjecture,(~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38,negated_conjecture,(epred1_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_19])])).
cnf(c_0_39,negated_conjecture,(k7_lattice3(k7_lattice3(esk1_0))=k7_lattice3(k7_lattice3(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_29]), c_0_30])])).
cnf(c_0_40,negated_conjecture,(k7_lattice3(k7_lattice3(k7_lattice3(esk2_0)))=k7_lattice3(esk2_0)), inference(spm,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_41,plain,(v1_xboole_0(X2)|v13_waybel_0(X2,k7_lattice3(X1))|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_7(X2,X1)|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_42,plain,(v1_xboole_0(X2)|v2_waybel_0(X2,k7_lattice3(X1))|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_7(X2,X1)|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_43,plain,(v1_xboole_0(X2)|m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_7(X2,X1)|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_44,negated_conjecture,(v2_waybel_7(esk3_0,k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35]), c_0_36])]), c_0_37]), c_0_38])])).
cnf(c_0_45,negated_conjecture,(k7_lattice3(esk1_0)=k7_lattice3(esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_19]), c_0_39]), c_0_40])).
cnf(c_0_46,negated_conjecture,(v2_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_47,negated_conjecture,(v1_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_48,negated_conjecture,(v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_49,negated_conjecture,(v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_50,negated_conjecture,(v3_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_51,negated_conjecture,(v13_waybel_0(esk3_0,k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_33]), c_0_34]), c_0_35]), c_0_36])]), c_0_37]), c_0_38])])).
cnf(c_0_52,negated_conjecture,(v2_waybel_0(esk3_0,k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_33]), c_0_34]), c_0_35]), c_0_36])]), c_0_37]), c_0_38])])).
cnf(c_0_53,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k7_lattice3(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_33]), c_0_34]), c_0_35]), c_0_36])]), c_0_37]), c_0_38])])).
cnf(c_0_54,negated_conjecture,(v1_xboole_0(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_waybel_7(esk3_0,esk2_0)|~v12_waybel_0(esk3_0,esk2_0)|~v1_waybel_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_55,plain,(v1_xboole_0(X2)|v1_waybel_0(X2,X1)|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))|~v2_waybel_7(X2,k7_lattice3(X1))|~v13_waybel_0(X2,k7_lattice3(X1))|~v2_waybel_0(X2,k7_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_56,negated_conjecture,(v2_waybel_7(esk3_0,k7_lattice3(esk2_0))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_57,negated_conjecture,(epred1_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_46]), c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_30])])).
cnf(c_0_58,negated_conjecture,(v13_waybel_0(esk3_0,k7_lattice3(esk2_0))), inference(rw,[status(thm)],[c_0_51, c_0_45])).
cnf(c_0_59,negated_conjecture,(v2_waybel_0(esk3_0,k7_lattice3(esk2_0))), inference(rw,[status(thm)],[c_0_52, c_0_45])).
cnf(c_0_60,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k7_lattice3(esk2_0))))), inference(rw,[status(thm)],[c_0_53, c_0_45])).
cnf(c_0_61,plain,(v1_xboole_0(X2)|v12_waybel_0(X2,X1)|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))|~v2_waybel_7(X2,k7_lattice3(X1))|~v13_waybel_0(X2,k7_lattice3(X1))|~v2_waybel_0(X2,k7_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_62,negated_conjecture,(~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_waybel_7(esk3_0,esk2_0)|~v12_waybel_0(esk3_0,esk2_0)|~v1_waybel_0(esk3_0,esk2_0)), inference(sr,[status(thm)],[c_0_54, c_0_37])).
cnf(c_0_63,plain,(v1_waybel_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])]), c_0_37]), c_0_58]), c_0_59]), c_0_60])])).
cnf(c_0_64,plain,(v12_waybel_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_56]), c_0_57])]), c_0_37]), c_0_58]), c_0_59]), c_0_60])])).
cnf(c_0_65,plain,(v1_xboole_0(X2)|v1_waybel_7(X2,X1)|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))|~v2_waybel_7(X2,k7_lattice3(X1))|~v13_waybel_0(X2,k7_lattice3(X1))|~v2_waybel_0(X2,k7_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_66,negated_conjecture,(~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_waybel_7(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63])]), c_0_64])])).
cnf(c_0_67,plain,(v1_waybel_7(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_56]), c_0_57])]), c_0_37]), c_0_58]), c_0_59]), c_0_60])])).
cnf(c_0_68,plain,(v1_xboole_0(X2)|m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~epred1_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k7_lattice3(X1))))|~v2_waybel_7(X2,k7_lattice3(X1))|~v13_waybel_0(X2,k7_lattice3(X1))|~v2_waybel_0(X2,k7_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_69,negated_conjecture,(~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67])])).
cnf(c_0_70,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_56]), c_0_57])]), c_0_37]), c_0_58]), c_0_59]), c_0_60])]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 56
# Proof object formula steps           : 15
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 86
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 32
# Processed clauses                    : 76
# ...of these trivial                  : 5
# ...subsumed                          : 0
# ...remaining for further processing  : 70
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 12
# Generated clauses                    : 91
# ...of the previous two non-trivial   : 78
# Contextual simplify-reflections      : 0
# Paramodulations                      : 91
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 23
# ...number of literals in the above   : 125
# Current number of archived formulas  : 0
# Current number of archived clauses   : 13
# Clause-clause subsumption calls (NU) : 57
# Rec. Clause-clause subsumption calls : 26
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4174

# -------------------------------------------------
# User time                : 0.000 s
# System time              : 0.008 s
# Total time               : 0.008 s
# Maximum resident set size: 43208 pages
