# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t17_waybel_0__10_premises__rerun.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t17_waybel_0__10_premises__rerun.E_input', dt_l1_orders_2)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t17_waybel_0__10_premises__rerun.E_input', d1_tarski)).
fof(d17_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('out/t17_waybel_0__10_premises__rerun.E_input', d17_waybel_0)).
fof(redefinition_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2))), file('out/t17_waybel_0__10_premises__rerun.E_input', redefinition_k6_domain_1)).
fof(fraenkel_a_2_0_waybel_0, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_0_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(X5,X3)))))))), file('out/t17_waybel_0__10_premises__rerun.E_input', fraenkel_a_2_0_waybel_0)).
fof(t14_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2)))), file('out/t17_waybel_0__10_premises__rerun.E_input', t14_waybel_0)).
fof(dt_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1)))), file('out/t17_waybel_0__10_premises__rerun.E_input', dt_k6_domain_1)).
fof(t17_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), file('out/t17_waybel_0__10_premises__rerun.E_input', t17_waybel_0)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t17_waybel_0__10_premises__rerun.E_input', t2_subset)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t17_waybel_0__10_premises__rerun.E_input', d1_xboole_0)).
fof(c_0_11, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_12, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
fof(c_0_14, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_orders_2(X3))|(~m1_subset_1(X4,u1_struct_0(X3))|k5_waybel_0(X3,X4)=k3_waybel_0(X3,k6_domain_1(u1_struct_0(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_waybel_0])])])])])])).
fof(c_0_15, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|k6_domain_1(X3,X4)=k1_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_16,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, (![X6]:![X7]:![X8]:![X11]:![X12]:((((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X7))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((X6=esk2_3(X6,X7,X8)|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((m1_subset_1(esk3_3(X6,X7,X8),u1_struct_0(X7))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((r1_orders_2(X7,esk2_3(X6,X7,X8),esk3_3(X6,X7,X8))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&((r2_hidden(esk3_3(X6,X7,X8),X8)|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))))))&(((~m1_subset_1(X11,u1_struct_0(X7))|(X6!=X11|(~m1_subset_1(X12,u1_struct_0(X7))|(~r1_orders_2(X7,X11,X12)|~r2_hidden(X12,X8)))))|r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_waybel_0])])])])])])])])).
cnf(c_0_19,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_20, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_orders_2(X3))|(~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|k3_waybel_0(X3,X4)=a_2_0_waybel_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_waybel_0])])])])])])).
cnf(c_0_21,plain,(k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,plain,(k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(v2_struct_0(X1)|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_24, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|m1_subset_1(k6_domain_1(X3,X4),k1_zfmisc_1(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_25,plain,(v2_struct_0(X2)|r1_orders_2(X2,esk2_3(X3,X2,X1),esk3_3(X3,X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(v2_struct_0(X2)|X3=esk2_3(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(v2_struct_0(X2)|r2_hidden(esk3_3(X3,X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29,plain,(k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30,plain,(k3_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_31,plain,(m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(r1_orders_2(X1,X2,esk3_3(X2,X1,X3))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33,plain,(esk3_3(X1,X2,k1_tarski(X3))=X3|v2_struct_0(X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,a_2_0_waybel_0(X2,k1_tarski(X3)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34,plain,(a_2_0_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v2_struct_0(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,plain,(m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(spm,[status(thm)],[c_0_31, c_0_22])).
cnf(c_0_36,plain,(r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,a_2_0_waybel_0(X1,k1_tarski(X3)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37,plain,(a_2_0_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_23])).
fof(c_0_38, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))))), inference(assume_negation,[status(cth)],[t17_waybel_0])).
cnf(c_0_39,plain,(r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X2,k5_waybel_0(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_40, negated_conjecture, (((~v2_struct_0(esk4_0)&l1_orders_2(esk4_0))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&(m1_subset_1(esk6_0,u1_struct_0(esk4_0))&((~r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))|~r1_orders_2(esk4_0,esk6_0,esk5_0))&(r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))|r1_orders_2(esk4_0,esk6_0,esk5_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_38])])])])])])).
cnf(c_0_41,plain,(v2_struct_0(X2)|r2_hidden(X3,a_2_0_waybel_0(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X4,X1)|~r1_orders_2(X2,X5,X4)|~m1_subset_1(X4,u1_struct_0(X2))|X3!=X5|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_42,plain,(r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X2,k5_waybel_0(X1,X3))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_35]), c_0_23])).
cnf(c_0_43,negated_conjecture,(r1_orders_2(esk4_0,esk6_0,esk5_0)|r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45,negated_conjecture,(l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_47, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_48,plain,(v2_struct_0(X2)|m1_subset_1(esk3_3(X3,X2,X1),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_49,plain,(r2_hidden(X1,a_2_0_waybel_0(X2,X3))|v2_struct_0(X2)|~r1_orders_2(X2,X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X4,X3)|~l1_orders_2(X2)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_50,negated_conjecture,(r1_orders_2(esk4_0,esk6_0,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_45])]), c_0_46])).
cnf(c_0_51,negated_conjecture,(m1_subset_1(esk6_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53,plain,(m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X3,a_2_0_waybel_0(X2,k1_tarski(X1)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_48, c_0_33])).
cnf(c_0_54,negated_conjecture,(r2_hidden(esk6_0,a_2_0_waybel_0(esk4_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(esk5_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_44]), c_0_45])]), c_0_46])).
fof(c_0_55, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk7_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_56,negated_conjecture,(r2_hidden(esk6_0,u1_struct_0(esk4_0))|v1_xboole_0(u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_52, c_0_51])).
cnf(c_0_57,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk4_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(esk5_0,k1_tarski(X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_45])]), c_0_46])).
cnf(c_0_58,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_59,negated_conjecture,(r2_hidden(esk6_0,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_56]), c_0_45])]), c_0_46])).
cnf(c_0_60,negated_conjecture,(~r1_orders_2(esk4_0,esk6_0,esk5_0)|~r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_61,negated_conjecture,(r2_hidden(esk6_0,k5_waybel_0(esk4_0,X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(esk5_0,k1_tarski(X1))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_37]), c_0_45])]), c_0_46]), c_0_57])).
cnf(c_0_62,negated_conjecture,(~v1_xboole_0(u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_64,negated_conjecture,(~r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_50])])).
cnf(c_0_65,negated_conjecture,(r2_hidden(esk6_0,k5_waybel_0(esk4_0,X1))|~m1_subset_1(X1,u1_struct_0(esk4_0))|~r2_hidden(esk5_0,k1_tarski(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_35]), c_0_62])).
cnf(c_0_66,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_63])])).
cnf(c_0_67,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_44]), c_0_66])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 45
# Proof object formula steps           : 23
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 156
# ...of these trivial                  : 0
# ...subsumed                          : 32
# ...remaining for further processing  : 124
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 5
# Generated clauses                    : 183
# ...of the previous two non-trivial   : 157
# Contextual simplify-reflections      : 41
# Paramodulations                      : 174
# Factorizations                       : 2
# Equation resolutions                 : 8
# Current number of processed clauses  : 81
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 69
# Current number of unprocessed clauses: 45
# ...number of literals in the above   : 341
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 4566
# Rec. Clause-clause subsumption calls : 567
# Non-unit clause-clause subsumptions  : 72
# Unit Clause-clause subsumption calls : 31
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7597

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45300 pages
