# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t17_waybel_0__9_premises__rerun.E_input', d1_tarski)).
fof(fraenkel_a_2_0_waybel_0, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_0_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(X5,X3)))))))), file('out/t17_waybel_0__9_premises__rerun.E_input', fraenkel_a_2_0_waybel_0)).
fof(d17_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('out/t17_waybel_0__9_premises__rerun.E_input', d17_waybel_0)).
fof(redefinition_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2))), file('out/t17_waybel_0__9_premises__rerun.E_input', redefinition_k6_domain_1)).
fof(dt_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1)))), file('out/t17_waybel_0__9_premises__rerun.E_input', dt_k6_domain_1)).
fof(t14_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2)))), file('out/t17_waybel_0__9_premises__rerun.E_input', t14_waybel_0)).
fof(t17_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), file('out/t17_waybel_0__9_premises__rerun.E_input', t17_waybel_0)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t17_waybel_0__9_premises__rerun.E_input', fc2_struct_0)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t17_waybel_0__9_premises__rerun.E_input', t4_subset)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t17_waybel_0__9_premises__rerun.E_input', dt_l1_orders_2)).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk1_2(X4,X5),X5)|esk1_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
fof(c_0_11, plain, (![X6]:![X7]:![X8]:![X11]:![X12]:((((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X7))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((X6=esk2_3(X6,X7,X8)|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((m1_subset_1(esk3_3(X6,X7,X8),u1_struct_0(X7))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(((r1_orders_2(X7,esk2_3(X6,X7,X8),esk3_3(X6,X7,X8))|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&((r2_hidden(esk3_3(X6,X7,X8),X8)|~r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))))))&(((~m1_subset_1(X11,u1_struct_0(X7))|(X6!=X11|(~m1_subset_1(X12,u1_struct_0(X7))|(~r1_orders_2(X7,X11,X12)|~r2_hidden(X12,X8)))))|r2_hidden(X6,a_2_0_waybel_0(X7,X8)))|((v2_struct_0(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_waybel_0])])])])])])])])).
cnf(c_0_12,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_orders_2(X3))|(~m1_subset_1(X4,u1_struct_0(X3))|k5_waybel_0(X3,X4)=k3_waybel_0(X3,k6_domain_1(u1_struct_0(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_waybel_0])])])])])])).
fof(c_0_14, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|k6_domain_1(X3,X4)=k1_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_15, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|m1_subset_1(k6_domain_1(X3,X4),k1_zfmisc_1(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_16,plain,(v2_struct_0(X2)|r1_orders_2(X2,esk2_3(X3,X2,X1),esk3_3(X3,X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(v2_struct_0(X2)|X3=esk2_3(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(v2_struct_0(X2)|r2_hidden(esk3_3(X3,X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_20, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_orders_2(X3))|(~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|k3_waybel_0(X3,X4)=a_2_0_waybel_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_waybel_0])])])])])])).
cnf(c_0_21,plain,(k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,plain,(k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,plain,(m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24,plain,(r1_orders_2(X1,X2,esk3_3(X2,X1,X3))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_25,plain,(esk3_3(X1,X2,k1_tarski(X3))=X3|v2_struct_0(X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,k1_tarski(X3)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26,plain,(k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,plain,(k3_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28,plain,(v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_23, c_0_22])).
fof(c_0_29, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))))), inference(assume_negation,[status(cth)],[t17_waybel_0])).
cnf(c_0_30,plain,(r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,k1_tarski(X3)))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,plain,(a_2_0_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
fof(c_0_32, negated_conjecture, (((~v2_struct_0(esk4_0)&l1_orders_2(esk4_0))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&(m1_subset_1(esk6_0,u1_struct_0(esk4_0))&((~r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))|~r1_orders_2(esk4_0,esk6_0,esk5_0))&(r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))|r1_orders_2(esk4_0,esk6_0,esk5_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_29])])])])])])).
fof(c_0_33, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_34,plain,(v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k5_waybel_0(X1,X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_28])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk4_0,esk6_0,esk5_0)|r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37,negated_conjecture,(l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_39, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_40,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41,negated_conjecture,(v1_xboole_0(u1_struct_0(esk4_0))|r1_orders_2(esk4_0,esk6_0,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_42, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_43,plain,(v2_struct_0(X2)|r2_hidden(X3,a_2_0_waybel_0(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X4,X1)|~r1_orders_2(X2,X5,X4)|~m1_subset_1(X4,u1_struct_0(X2))|X3!=X5|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_44,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45,negated_conjecture,(r1_orders_2(esk4_0,esk6_0,esk5_0)|~l1_struct_0(esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_38])).
cnf(c_0_46,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_47,plain,(v2_struct_0(X2)|m1_subset_1(esk3_3(X3,X2,X1),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_48,plain,(v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X4,X3)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_43, c_0_44])])).
cnf(c_0_49,negated_conjecture,(r1_orders_2(esk4_0,esk6_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_37])])).
cnf(c_0_50,negated_conjecture,(m1_subset_1(esk6_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_51,plain,(m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,a_2_0_waybel_0(X2,k1_tarski(X1)))), inference(spm,[status(thm)],[c_0_47, c_0_25])).
cnf(c_0_52,negated_conjecture,(r2_hidden(esk6_0,a_2_0_waybel_0(esk4_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(esk5_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_37])]), c_0_38])).
cnf(c_0_53,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk4_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(esk5_0,k1_tarski(X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_37])]), c_0_38])).
cnf(c_0_54,negated_conjecture,(~r1_orders_2(esk4_0,esk6_0,esk5_0)|~r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_55,negated_conjecture,(v1_xboole_0(u1_struct_0(esk4_0))|r2_hidden(esk6_0,k5_waybel_0(esk4_0,X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(esk5_0,k1_tarski(X1))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_31]), c_0_37])]), c_0_38]), c_0_53])).
cnf(c_0_56,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_57,negated_conjecture,(~r2_hidden(esk6_0,k5_waybel_0(esk4_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_49])])).
cnf(c_0_58,negated_conjecture,(v1_xboole_0(u1_struct_0(esk4_0))|r2_hidden(esk6_0,k5_waybel_0(esk4_0,X1))|~m1_subset_1(X1,u1_struct_0(esk4_0))|~r2_hidden(esk5_0,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_55, c_0_28])).
cnf(c_0_59,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_56])])).
cnf(c_0_60,negated_conjecture,(v1_xboole_0(u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_36]), c_0_59])])).
cnf(c_0_61,negated_conjecture,(~l1_struct_0(esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_60]), c_0_38])).
cnf(c_0_62,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_46]), c_0_37])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 42
# Proof object formula steps           : 21
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 80
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 77
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 74
# ...of the previous two non-trivial   : 54
# Contextual simplify-reflections      : 6
# Paramodulations                      : 65
# Factorizations                       : 2
# Equation resolutions                 : 8
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 131
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 678
# Rec. Clause-clause subsumption calls : 88
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3856

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45304 pages
