# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C18_F1_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.016 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t18_lattice3, conjecture, (![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k12_lattice3(X1,X2,k13_lattice3(X1,X2,X3))=X2)))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', t18_lattice3)).
fof(reflexivity_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', reflexivity_r3_orders_2)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', redefinition_r3_orders_2)).
fof(d14_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v5_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))))))))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', d14_lattice3)).
fof(cc1_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', cc1_lattice3)).
fof(redefinition_k13_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', redefinition_k13_lattice3)).
fof(commutativity_k13_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k13_lattice3(X1,X3,X2))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', commutativity_k13_lattice3)).
fof(redefinition_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', redefinition_k12_lattice3)).
fof(commutativity_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', commutativity_k12_lattice3)).
fof(l26_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k10_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))))))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', l26_lattice3)).
fof(dt_k10_lattice3, axiom, (![X1]:![X2]:![X3]:((l1_orders_2(X1)&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k10_lattice3(X1,X2,X3),u1_struct_0(X1)))), file('out/t18_lattice3__52_premises__rerun__rerun.E_input', dt_k10_lattice3)).
fof(c_0_11, negated_conjecture, (~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k12_lattice3(X1,X2,k13_lattice3(X1,X2,X3))=X2))))), inference(assume_negation,[status(cth)],[t18_lattice3])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|r3_orders_2(X4,X5,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])])])).
fof(c_0_13, negated_conjecture, (((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&k12_lattice3(esk1_0,esk2_0,k13_lattice3(esk1_0,esk2_0,esk3_0))!=esk2_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_15,plain,(r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:![X12]:((((((((((r1_orders_2(X6,X11,X7)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((r1_orders_2(X6,X11,X8)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((~m1_subset_1(X12,u1_struct_0(X6))|((~r1_orders_2(X6,X12,X7)|~r1_orders_2(X6,X12,X8))|r1_orders_2(X6,X12,X11)))|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))))&((((((((((m1_subset_1(esk5_4(X6,X7,X8,X11),u1_struct_0(X6))|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X7)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X8)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6)))&(((((((((~r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X11)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6)))))&(((((((((((r1_orders_2(X6,X11,X7)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((r1_orders_2(X6,X11,X8)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((~m1_subset_1(X12,u1_struct_0(X6))|((~r1_orders_2(X6,X12,X7)|~r1_orders_2(X6,X12,X8))|r1_orders_2(X6,X12,X11)))|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))))&((((((((((m1_subset_1(esk5_4(X6,X7,X8,X11),u1_struct_0(X6))|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X7)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X8)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6)))&(((((((((~r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X11)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,X11,X7)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((r1_orders_2(X6,X11,X8)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((~m1_subset_1(X12,u1_struct_0(X6))|((~r1_orders_2(X6,X12,X7)|~r1_orders_2(X6,X12,X8))|r1_orders_2(X6,X12,X11)))|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))))&((((((((((m1_subset_1(esk5_4(X6,X7,X8,X11),u1_struct_0(X6))|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X7)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X8)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6)))&(((((((((~r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X11)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))))))&(((((((((r1_orders_2(X6,X11,X7)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((r1_orders_2(X6,X11,X8)|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&((((((((~m1_subset_1(X12,u1_struct_0(X6))|((~r1_orders_2(X6,X12,X7)|~r1_orders_2(X6,X12,X8))|r1_orders_2(X6,X12,X11)))|X11!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))))&((((((((((m1_subset_1(esk5_4(X6,X7,X8,X11),u1_struct_0(X6))|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X7)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6))&(((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X8)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6)))&(((((((((~r1_orders_2(X6,esk5_4(X6,X7,X8,X11),X11)|~r1_orders_2(X6,X11,X8))|~r1_orders_2(X6,X11,X7))|X11=k11_lattice3(X6,X7,X8))|~m1_subset_1(X11,u1_struct_0(X6)))|(((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|~m1_subset_1(X9,u1_struct_0(X6))))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|~v5_orders_2(X6))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_lattice3])])])])])])])).
cnf(c_0_20,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,negated_conjecture,(r3_orders_2(esk1_0,X1,X1)|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_22,plain,(r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X3)|X5=k11_lattice3(X1,X2,X3)|r1_orders_2(X1,esk5_4(X1,X2,X3,X5),X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~m1_subset_1(X5,u1_struct_0(X1))|~r1_orders_2(X1,X5,X2)|~r1_orders_2(X1,X5,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_25, plain, (![X2]:(~l1_orders_2(X2)|(~v1_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_26,negated_conjecture,(r1_orders_2(esk1_0,X1,X1)|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_17]), c_0_18])])).
cnf(c_0_27,negated_conjecture,(k11_lattice3(esk1_0,X1,X2)=esk2_0|r1_orders_2(esk1_0,esk5_4(esk1_0,X1,X2,esk2_0),X2)|r1_orders_2(esk1_0,esk4_4(esk1_0,X1,X2,X3),X2)|~r1_orders_2(esk1_0,esk2_0,X2)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X3,X2)|~r1_orders_2(esk1_0,X3,X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_17]), c_0_24])])).
cnf(c_0_28,plain,(~v2_struct_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk2_0)|v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_26, c_0_23])).
cnf(c_0_30,negated_conjecture,(v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_31, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v1_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k13_lattice3(X4,X5,X6)=k10_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
fof(c_0_32, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v1_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k13_lattice3(X4,X5,X6)=k13_lattice3(X4,X6,X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k13_lattice3])])).
cnf(c_0_33,plain,(r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X3)|X5=k11_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~m1_subset_1(X5,u1_struct_0(X1))|~r1_orders_2(X1,X5,X2)|~r1_orders_2(X1,X5,X3)|~r1_orders_2(X1,esk5_4(X1,X2,X3,X5),X5)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34,negated_conjecture,(k11_lattice3(esk1_0,X1,X2)=esk2_0|r1_orders_2(esk1_0,esk4_4(esk1_0,X1,X2,esk2_0),X2)|r1_orders_2(esk1_0,esk5_4(esk1_0,X1,X2,esk2_0),X2)|~r1_orders_2(esk1_0,esk2_0,X2)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_23])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_17]), c_0_30])])).
cnf(c_0_36,plain,(k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37,plain,(k13_lattice3(X1,X2,X3)=k13_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k12_lattice3(X4,X5,X6)=k11_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
fof(c_0_39, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k12_lattice3(X4,X5,X6)=k12_lattice3(X4,X6,X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k12_lattice3])])).
cnf(c_0_40,negated_conjecture,(k11_lattice3(esk1_0,X1,esk2_0)=esk2_0|r1_orders_2(esk1_0,esk4_4(esk1_0,X1,esk2_0,esk2_0),esk2_0)|r1_orders_2(esk1_0,esk4_4(esk1_0,X1,esk2_0,X2),esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X2,esk2_0)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_23]), c_0_17]), c_0_24])])).
cnf(c_0_41,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,k13_lattice3(esk1_0,esk2_0,esk3_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_42,plain,(k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,plain,(k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44,plain,(k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45,plain,(X5=k11_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X4)|~m1_subset_1(X5,u1_struct_0(X1))|~r1_orders_2(X1,X5,X2)|~r1_orders_2(X1,X5,X3)|~r1_orders_2(X1,esk5_4(X1,X2,X3,X5),X5)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_46,negated_conjecture,(k11_lattice3(esk1_0,X1,esk2_0)=esk2_0|r1_orders_2(esk1_0,esk4_4(esk1_0,X1,esk2_0,esk2_0),esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_23]), c_0_35])])).
cnf(c_0_47,plain,(X5=k11_lattice3(X1,X2,X3)|r1_orders_2(X1,esk5_4(X1,X2,X3,X5),X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X4)|~m1_subset_1(X5,u1_struct_0(X1))|~r1_orders_2(X1,X5,X2)|~r1_orders_2(X1,X5,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_48, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X7,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X8,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X7,X10)|~r1_orders_2(X6,X8,X10))|r1_orders_2(X6,X9,X10)))|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk6_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,X7,esk6_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,X8,esk6_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,X9,esk6_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l26_lattice3])])])])])])])])).
cnf(c_0_49,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,k10_lattice3(esk1_0,esk3_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_23]), c_0_16]), c_0_17]), c_0_30]), c_0_24])])).
cnf(c_0_50,plain,(k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_52,negated_conjecture,(k11_lattice3(esk1_0,X1,esk2_0)=esk2_0|X2=k11_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk5_4(esk1_0,X1,esk2_0,X2),X2)|~r1_orders_2(esk1_0,X2,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_35]), c_0_23]), c_0_17]), c_0_24])])).
cnf(c_0_53,negated_conjecture,(k11_lattice3(esk1_0,X1,esk2_0)=esk2_0|X2=k11_lattice3(esk1_0,X1,esk2_0)|r1_orders_2(esk1_0,esk5_4(esk1_0,X1,esk2_0,X2),esk2_0)|~r1_orders_2(esk1_0,X2,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_46]), c_0_35]), c_0_23]), c_0_17]), c_0_24])])).
cnf(c_0_54,plain,(v2_struct_0(X1)|r1_orders_2(X1,X3,X4)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X4!=k10_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_55, plain, (![X4]:![X5]:![X6]:((~l1_orders_2(X4)|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k10_lattice3(X4,X5,X6),u1_struct_0(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_lattice3])])).
cnf(c_0_56,negated_conjecture,(k11_lattice3(esk1_0,k10_lattice3(esk1_0,esk3_0,esk2_0),esk2_0)!=esk2_0|~m1_subset_1(k10_lattice3(esk1_0,esk3_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_23]), c_0_17]), c_0_51]), c_0_24])])).
cnf(c_0_57,negated_conjecture,(k11_lattice3(esk1_0,X1,esk2_0)=esk2_0|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_35]), c_0_23])])).
cnf(c_0_58,plain,(r1_orders_2(X1,X2,X3)|X3!=k10_lattice3(X1,X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[c_0_54, c_0_28])).
cnf(c_0_59,plain,(m1_subset_1(k10_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60,negated_conjecture,(~r1_orders_2(esk1_0,esk2_0,k10_lattice3(esk1_0,esk3_0,esk2_0))|~m1_subset_1(k10_lattice3(esk1_0,esk3_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61,plain,(r1_orders_2(X1,X2,k10_lattice3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_58]), c_0_59])).
cnf(c_0_62,negated_conjecture,(~m1_subset_1(k10_lattice3(esk1_0,esk3_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_23]), c_0_16]), c_0_17]), c_0_30]), c_0_24])])).
cnf(c_0_63,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_59]), c_0_23]), c_0_16]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 41
# Proof object formula steps           : 23
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 55
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 52
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 52
# Processed clauses                    : 1936
# ...of these trivial                  : 105
# ...subsumed                          : 1142
# ...remaining for further processing  : 689
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 148
# Backward-rewritten                   : 60
# Generated clauses                    : 7453
# ...of the previous two non-trivial   : 6978
# Contextual simplify-reflections      : 1231
# Paramodulations                      : 7383
# Factorizations                       : 0
# Equation resolutions                 : 70
# Current number of processed clauses  : 481
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 461
# Current number of unprocessed clauses: 3221
# ...number of literals in the above   : 40479
# Current number of archived formulas  : 0
# Current number of archived clauses   : 208
# Clause-clause subsumption calls (NU) : 221078
# Rec. Clause-clause subsumption calls : 8637
# Non-unit clause-clause subsumptions  : 2324
# Unit Clause-clause subsumption calls : 1106
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 374092

# -------------------------------------------------
# User time                : 0.392 s
# System time              : 0.004 s
# Total time               : 0.396 s
# Maximum resident set size: 45276 pages
