# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l26_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k10_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))))))), file('out/t18_lattice3__10_premises__rerun.E_input', l26_lattice3)).
fof(cc1_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t18_lattice3__10_premises__rerun.E_input', cc1_lattice3)).
fof(dt_k13_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1)))), file('out/t18_lattice3__10_premises__rerun.E_input', dt_k13_lattice3)).
fof(redefinition_k13_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3))), file('out/t18_lattice3__10_premises__rerun.E_input', redefinition_k13_lattice3)).
fof(t17_lattice3, axiom, (![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k13_lattice3(X1,k12_lattice3(X1,X2,X3),X3)=X3)))), file('out/t18_lattice3__10_premises__rerun.E_input', t17_lattice3)).
fof(dt_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1)))), file('out/t18_lattice3__10_premises__rerun.E_input', dt_k12_lattice3)).
fof(t18_lattice3, conjecture, (![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k12_lattice3(X1,X2,k13_lattice3(X1,X2,X3))=X2)))), file('out/t18_lattice3__10_premises__rerun.E_input', t18_lattice3)).
fof(l28_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))))))), file('out/t18_lattice3__10_premises__rerun.E_input', l28_lattice3)).
fof(cc2_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t18_lattice3__10_premises__rerun.E_input', cc2_lattice3)).
fof(redefinition_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3))), file('out/t18_lattice3__10_premises__rerun.E_input', redefinition_k12_lattice3)).
fof(t13_lattice3, axiom, (![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2))))), file('out/t18_lattice3__10_premises__rerun.E_input', t13_lattice3)).
fof(c_0_11, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X7,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X8,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X7,X10)|~r1_orders_2(X6,X8,X10))|r1_orders_2(X6,X9,X10)))|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,X7,esk1_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,X8,esk1_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,X9,esk1_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l26_lattice3])])])])])])])])).
fof(c_0_12, plain, (![X2]:(~l1_orders_2(X2)|(~v1_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v1_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k13_lattice3(X4,X5,X6),u1_struct_0(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k13_lattice3])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v1_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k13_lattice3(X4,X5,X6)=k10_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
cnf(c_0_15,plain,(v2_struct_0(X1)|r1_orders_2(X1,X3,X4)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X4!=k10_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(~v2_struct_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:((~v3_orders_2(X4)|(~v5_orders_2(X4)|(~v1_lattice3(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|k13_lattice3(X4,k12_lattice3(X4,X5,X6),X6)=X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_lattice3])])])])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k12_lattice3(X4,X5,X6),u1_struct_0(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k12_lattice3])])).
fof(c_0_21, negated_conjecture, (~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k12_lattice3(X1,X2,k13_lattice3(X1,X2,X3))=X2))))), inference(assume_negation,[status(cth)],[t18_lattice3])).
fof(c_0_22, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X9,X7)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X9,X8)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X10,X7)|~r1_orders_2(X6,X10,X8))|r1_orders_2(X6,X10,X9)))|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk5_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,esk5_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,esk5_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l28_lattice3])])])])])])])])).
fof(c_0_23, plain, (![X2]:(~l1_orders_2(X2)|(~v2_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
cnf(c_0_24,plain,(r1_orders_2(X1,X2,k10_lattice3(X1,X3,X2))|~m1_subset_1(k10_lattice3(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_15, c_0_16])])).
cnf(c_0_25,plain,(m1_subset_1(k10_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26,plain,(k13_lattice3(X1,k12_lattice3(X1,X2,X3),X3)=X3|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, negated_conjecture, (((v3_orders_2(esk2_0)&(v5_orders_2(esk2_0)&(v1_lattice3(esk2_0)&(v2_lattice3(esk2_0)&l1_orders_2(esk2_0)))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&k12_lattice3(esk2_0,esk3_0,k13_lattice3(esk2_0,esk3_0,esk4_0))!=esk3_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
cnf(c_0_29,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,esk5_4(X1,X2,X3,X4),X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,plain,(~v2_struct_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|r1_orders_2(X1,esk5_4(X1,X2,X3,X4),X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,plain,(r1_orders_2(X1,X2,k10_lattice3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,plain,(k10_lattice3(X1,k12_lattice3(X1,X2,X3),X3)=X3|~v2_lattice3(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_26]), c_0_27])).
cnf(c_0_34,negated_conjecture,(k12_lattice3(esk2_0,esk3_0,k13_lattice3(esk2_0,esk3_0,esk4_0))!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38,negated_conjecture,(v1_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39,negated_conjecture,(v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_40, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k12_lattice3(X4,X5,X6)=k11_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_41,plain,(X1=k11_lattice3(X2,X3,X4)|~v2_lattice3(X2)|~r1_orders_2(X2,esk5_4(X2,X3,X4,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_42,plain,(X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk5_4(X2,X3,X4,X1),X3)|~v2_lattice3(X2)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_43,plain,(r1_orders_2(X1,X2,X2)|~v2_lattice3(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_27])).
cnf(c_0_44,negated_conjecture,(v2_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_45,negated_conjecture,(v3_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_46,negated_conjecture,(k12_lattice3(esk2_0,esk3_0,k10_lattice3(esk2_0,esk3_0,esk4_0))!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_18]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])])).
cnf(c_0_47,plain,(k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48,plain,(k11_lattice3(X1,X2,X3)=X2|~v2_lattice3(X1)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49,negated_conjecture,(r1_orders_2(esk2_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_35]), c_0_44]), c_0_45]), c_0_37]), c_0_38]), c_0_39])])).
fof(c_0_50, plain, (![X4]:![X5]:![X6]:((~v5_orders_2(X4)|(~v1_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|k10_lattice3(X4,X5,X6)=k10_lattice3(X4,X6,X5))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_lattice3])])])])])).
cnf(c_0_51,negated_conjecture,(k11_lattice3(esk2_0,esk3_0,k10_lattice3(esk2_0,esk3_0,esk4_0))!=esk3_0|~m1_subset_1(k10_lattice3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_44]), c_0_36]), c_0_37]), c_0_39])])).
cnf(c_0_52,negated_conjecture,(k11_lattice3(esk2_0,X1,X2)=X1|~r1_orders_2(esk2_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_44]), c_0_37]), c_0_39])])).
cnf(c_0_53,plain,(k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54,negated_conjecture,(~r1_orders_2(esk2_0,esk3_0,k10_lattice3(esk2_0,esk3_0,esk4_0))|~m1_subset_1(k10_lattice3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_36])])).
cnf(c_0_55,plain,(r1_orders_2(X1,X2,k10_lattice3(X1,X2,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_53])).
cnf(c_0_56,negated_conjecture,(~m1_subset_1(k10_lattice3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_36]), c_0_35]), c_0_37]), c_0_38]), c_0_39])])).
cnf(c_0_57,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_25]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 35
# Proof object formula steps           : 23
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 42
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 159
# ...of these trivial                  : 0
# ...subsumed                          : 51
# ...remaining for further processing  : 108
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 179
# ...of the previous two non-trivial   : 153
# Contextual simplify-reflections      : 58
# Paramodulations                      : 173
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 59
# Current number of unprocessed clauses: 43
# ...number of literals in the above   : 324
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 2486
# Rec. Clause-clause subsumption calls : 209
# Non-unit clause-clause subsumptions  : 108
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9823

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 45256 pages
