# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t18_yellow_1, conjecture, (![X1]:k3_yellow_0(k3_yellow_1(X1))=k1_xboole_0), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', t18_yellow_1)).
fof(d2_yellow_1, axiom, (![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', d2_yellow_1)).
fof(d11_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', d11_yellow_0)).
fof(t29_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k1_yellow_0(k3_lattice3(X1),X2)&k16_lattice3(X1,X2)=k2_yellow_0(k3_lattice3(X1),X2)))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', t29_yellow_0)).
fof(dt_k3_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>(v1_orders_2(k3_lattice3(X1))&(v3_orders_2(k3_lattice3(X1))&(v4_orders_2(k3_lattice3(X1))&(v5_orders_2(k3_lattice3(X1))&l1_orders_2(k3_lattice3(X1)))))))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', dt_k3_lattice3)).
fof(dt_k1_lattice3, axiom, (![X1]:(v3_lattices(k1_lattice3(X1))&l3_lattices(k1_lattice3(X1)))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', dt_k1_lattice3)).
fof(fc7_lattice3, axiom, (![X1]:(v3_lattices(k1_lattice3(X1))&v4_lattice3(k1_lattice3(X1)))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', fc7_lattice3)).
fof(fc2_lattice3, axiom, (![X1]:(v3_lattices(k1_lattice3(X1))&v10_lattices(k1_lattice3(X1)))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', fc2_lattice3)).
fof(fc1_lattice3, axiom, (![X1]:(~(v2_struct_0(k1_lattice3(X1)))&v3_lattices(k1_lattice3(X1)))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', fc1_lattice3)).
fof(t49_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', t49_lattice3)).
fof(t3_lattice3, axiom, (![X1]:(v13_lattices(k1_lattice3(X1))&k5_lattices(k1_lattice3(X1))=k1_xboole_0)), file('out/t18_yellow_1__213_premises__rerun__rerun.E_input', t3_lattice3)).
fof(c_0_11, negated_conjecture, (~(![X1]:k3_yellow_0(k3_yellow_1(X1))=k1_xboole_0)), inference(assume_negation,[status(cth)],[t18_yellow_1])).
fof(c_0_12, negated_conjecture, (k3_yellow_0(k3_yellow_1(esk1_0))!=k1_xboole_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_13, plain, (![X2]:k3_yellow_1(X2)=k3_lattice3(k1_lattice3(X2))), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_14, plain, (![X2]:(~l1_orders_2(X2)|k3_yellow_0(X2)=k1_yellow_0(X2,k1_xboole_0))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
fof(c_0_15, plain, (![X3]:![X4]:![X4]:((k15_lattice3(X3,X4)=k1_yellow_0(k3_lattice3(X3),X4)|(v2_struct_0(X3)|(~v10_lattices(X3)|(~v4_lattice3(X3)|~l3_lattices(X3)))))&(k16_lattice3(X3,X4)=k2_yellow_0(k3_lattice3(X3),X4)|(v2_struct_0(X3)|(~v10_lattices(X3)|(~v4_lattice3(X3)|~l3_lattices(X3))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_yellow_0])])])])])])])).
fof(c_0_16, plain, (![X2]:((v1_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v3_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v4_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v5_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&(l1_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_lattice3])])])])).
cnf(c_0_17,negated_conjecture,(k3_yellow_0(k3_yellow_1(esk1_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(v2_struct_0(X1)|k15_lattice3(X1,X2)=k1_yellow_0(k3_lattice3(X1),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(v2_struct_0(X1)|l1_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, (![X2]:![X2]:(v3_lattices(k1_lattice3(X2))&l3_lattices(k1_lattice3(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k1_lattice3])])])).
fof(c_0_23, plain, (![X2]:![X2]:(v3_lattices(k1_lattice3(X2))&v4_lattice3(k1_lattice3(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc7_lattice3])])])).
fof(c_0_24, plain, (![X2]:![X2]:(v3_lattices(k1_lattice3(X2))&v10_lattices(k1_lattice3(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc2_lattice3])])])).
fof(c_0_25, plain, (![X2]:![X2]:(~v2_struct_0(k1_lattice3(X2))&v3_lattices(k1_lattice3(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_lattice3])])])])).
cnf(c_0_26,negated_conjecture,(k3_yellow_0(k3_lattice3(k1_lattice3(esk1_0)))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_27,plain,(k3_yellow_0(k3_lattice3(X1))=k15_lattice3(X1,k1_xboole_0)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_28,plain,(l3_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(v4_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(v10_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(~v2_struct_0(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|(~v10_lattices(X2)|(~v4_lattice3(X2)|~l3_lattices(X2)))))&((v10_lattices(X2)|(v2_struct_0(X2)|(~v10_lattices(X2)|(~v4_lattice3(X2)|~l3_lattices(X2)))))&((v13_lattices(X2)|(v2_struct_0(X2)|(~v10_lattices(X2)|(~v4_lattice3(X2)|~l3_lattices(X2)))))&(l3_lattices(X2)|(v2_struct_0(X2)|(~v10_lattices(X2)|(~v4_lattice3(X2)|~l3_lattices(X2))))))))&(k5_lattices(X2)=k15_lattice3(X2,k1_xboole_0)|(v2_struct_0(X2)|(~v10_lattices(X2)|(~v4_lattice3(X2)|~l3_lattices(X2))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_lattice3])])])])).
fof(c_0_33, plain, (![X2]:![X2]:(v13_lattices(k1_lattice3(X2))&k5_lattices(k1_lattice3(X2))=k1_xboole_0)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t3_lattice3])])])).
cnf(c_0_34,negated_conjecture,(k15_lattice3(k1_lattice3(esk1_0),k1_xboole_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_35,plain,(v2_struct_0(X1)|k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_36,plain,(k5_lattices(k1_lattice3(X1))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_37,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_28]), c_0_29]), c_0_30])]), c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 15
# Proof object formula steps           : 23
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 21
# Processed clauses                    : 41
# ...of these trivial                  : 3
# ...subsumed                          : 0
# ...remaining for further processing  : 38
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 3
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 1
# Paramodulations                      : 3
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 20
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 11
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 52
# Rec. Clause-clause subsumption calls : 25
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1443

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 43180 pages
