# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t18_lattice3, conjecture, (![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k12_lattice3(X1,X2,k13_lattice3(X1,X2,X3))=X2)))), file('out/t18_lattice3__8_premises__rerun.E_input', t18_lattice3)).
fof(cc2_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t18_lattice3__8_premises__rerun.E_input', cc2_lattice3)).
fof(l28_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))))))), file('out/t18_lattice3__8_premises__rerun.E_input', l28_lattice3)).
fof(reflexivity_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2))), file('out/t18_lattice3__8_premises__rerun.E_input', reflexivity_r3_orders_2)).
fof(redefinition_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3))), file('out/t18_lattice3__8_premises__rerun.E_input', redefinition_k12_lattice3)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t18_lattice3__8_premises__rerun.E_input', redefinition_r3_orders_2)).
fof(redefinition_k13_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3))), file('out/t18_lattice3__8_premises__rerun.E_input', redefinition_k13_lattice3)).
fof(l26_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k10_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))))))), file('out/t18_lattice3__8_premises__rerun.E_input', l26_lattice3)).
fof(dt_k13_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1)))), file('out/t18_lattice3__8_premises__rerun.E_input', dt_k13_lattice3)).
fof(c_0_9, negated_conjecture, (~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k12_lattice3(X1,X2,k13_lattice3(X1,X2,X3))=X2))))), inference(assume_negation,[status(cth)],[t18_lattice3])).
fof(c_0_10, plain, (![X2]:(~l1_orders_2(X2)|(~v2_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_11, negated_conjecture, (((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&k12_lattice3(esk1_0,esk2_0,k13_lattice3(esk1_0,esk2_0,esk3_0))!=esk2_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
fof(c_0_12, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X9,X7)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X9,X8)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X10,X7)|~r1_orders_2(X6,X10,X8))|r1_orders_2(X6,X10,X9)))|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l28_lattice3])])])])])])])])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|r3_orders_2(X4,X5,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])])])).
cnf(c_0_14,plain,(~v2_struct_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k12_lattice3(X4,X5,X6)=k11_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_18,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_21,plain,(r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24,negated_conjecture,(~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
cnf(c_0_25,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,k13_lattice3(esk1_0,esk2_0,esk3_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_26,plain,(k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_29, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v1_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k13_lattice3(X4,X5,X6)=k10_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
cnf(c_0_30,plain,(X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk4_4(X2,X3,X4,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_18, c_0_14])).
cnf(c_0_31,plain,(X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk4_4(X2,X3,X4,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_19, c_0_14])).
cnf(c_0_32,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33,negated_conjecture,(r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_16]), c_0_23])]), c_0_24])).
fof(c_0_34, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X7,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X8,X9)|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X7,X10)|~r1_orders_2(X6,X8,X10))|r1_orders_2(X6,X9,X10)))|X9!=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk5_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,X7,esk5_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,X8,esk5_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,X9,esk5_4(X6,X7,X8,X9))|~r1_orders_2(X6,X8,X9))|~r1_orders_2(X6,X7,X9))|X9=k10_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l26_lattice3])])])])])])])])).
fof(c_0_35, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v1_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k13_lattice3(X4,X5,X6),u1_struct_0(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k13_lattice3])])).
cnf(c_0_36,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,k13_lattice3(esk1_0,esk2_0,esk3_0))!=esk2_0|~m1_subset_1(k13_lattice3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_16]), c_0_15]), c_0_28])])).
cnf(c_0_37,plain,(k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38,negated_conjecture,(v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39,plain,(k11_lattice3(X1,X2,X3)=X2|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40,negated_conjecture,(r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_16]), c_0_23])]), c_0_24])).
cnf(c_0_41,plain,(v2_struct_0(X1)|r1_orders_2(X1,X2,X4)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X4!=k10_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42,plain,(m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,k10_lattice3(esk1_0,esk2_0,esk3_0))!=esk2_0|~m1_subset_1(k10_lattice3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_22]), c_0_27]), c_0_16]), c_0_38]), c_0_28])])).
cnf(c_0_44,negated_conjecture,(k11_lattice3(esk1_0,X1,X2)=X1|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_16]), c_0_15]), c_0_28])])).
cnf(c_0_45,plain,(r1_orders_2(X1,X2,k10_lattice3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(k10_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_46,plain,(m1_subset_1(k10_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_42, c_0_37])).
cnf(c_0_47,negated_conjecture,(~r1_orders_2(esk1_0,esk2_0,k10_lattice3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k10_lattice3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_27])])).
cnf(c_0_48,plain,(r1_orders_2(X1,X2,k10_lattice3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_49,negated_conjecture,(~m1_subset_1(k10_lattice3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_22]), c_0_27]), c_0_16]), c_0_38]), c_0_28])]), c_0_24])).
cnf(c_0_50,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_46]), c_0_22]), c_0_27]), c_0_16]), c_0_38]), c_0_28])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 32
# Proof object formula steps           : 19
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 100
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 64
# ...of the previous two non-trivial   : 55
# Contextual simplify-reflections      : 35
# Paramodulations                      : 58
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 13
# ...number of literals in the above   : 123
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 1538
# Rec. Clause-clause subsumption calls : 82
# Non-unit clause-clause subsumptions  : 48
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6017

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 45256 pages
