# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d16_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k4_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X5,X4)&r2_hidden(X5,X2)))))))))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', d16_waybel_0)).
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', t7_boole)).
fof(d9_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))))))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', d9_orders_2)).
fof(cc3_relset_1, axiom, (![X1]:![X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3)))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', cc3_relset_1)).
fof(dt_u1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', dt_u1_orders_2)).
fof(d18_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k6_waybel_0(X1,X2)=k4_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', d18_waybel_0)).
fof(dt_k6_waybel_0, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', dt_k6_waybel_0)).
fof(dt_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1)))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', dt_k6_domain_1)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', d1_tarski)).
fof(redefinition_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', redefinition_k6_domain_1)).
fof(t18_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3)))))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', t18_waybel_0)).
fof(t5_subset, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3))))), file('out/t18_waybel_0__12_premises__rerun__rerun.E_input', t5_subset)).
fof(c_0_12, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:![X13]:(((((((((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((((r1_orders_2(X6,esk1_4(X6,X7,X8,X9),X9)|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((r2_hidden(esk1_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))))&(((((((~m1_subset_1(X11,u1_struct_0(X6))|(~r1_orders_2(X6,X11,X9)|~r2_hidden(X11,X7)))|r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))&(((((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X6))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((~r2_hidden(esk2_3(X6,X7,X8),X8)|(~m1_subset_1(X13,u1_struct_0(X6))|(~r1_orders_2(X6,X13,esk2_3(X6,X7,X8))|~r2_hidden(X13,X7))))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((m1_subset_1(esk3_3(X6,X7,X8),u1_struct_0(X6))|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((r1_orders_2(X6,esk3_3(X6,X7,X8),esk2_3(X6,X7,X8))|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((r2_hidden(esk3_3(X6,X7,X8),X7)|r2_hidden(esk2_3(X6,X7,X8),X8))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_waybel_0])])])])])])])).
fof(c_0_13, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(((((~r1_orders_2(X4,X5,X6)|r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&((((~r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4))|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(~v1_xboole_0(X4)|(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])])])).
fof(c_0_16, plain, (![X2]:(~l1_orders_2(X2)|m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_17,plain,(r2_hidden(X4,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X5,X2)|~r1_orders_2(X1,X5,X4)|~m1_subset_1(X5,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_orders_2(X3))|(~m1_subset_1(X4,u1_struct_0(X3))|k6_waybel_0(X3,X4)=k4_waybel_0(X3,k6_domain_1(u1_struct_0(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_waybel_0])])])])])])).
fof(c_0_19, plain, (![X3]:![X4]:(((v2_struct_0(X3)|~l1_orders_2(X3))|~m1_subset_1(X4,u1_struct_0(X3)))|m1_subset_1(k6_waybel_0(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_waybel_0])])])).
cnf(c_0_20,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,plain,(r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,plain,(v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,plain,(r2_hidden(X1,k4_waybel_0(X2,X3))|~r1_orders_2(X2,X4,X1)|~r2_hidden(X4,X3)|~m1_subset_1(k4_waybel_0(X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(k6_waybel_0(X1,X2)=k4_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|m1_subset_1(k6_domain_1(X3,X4),k1_zfmisc_1(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_28,plain,(~v1_xboole_0(u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29,plain,(v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_30, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_31,plain,(v2_struct_0(X1)|r2_hidden(X2,k6_waybel_0(X1,X3))|~r1_orders_2(X1,X4,X2)|~r2_hidden(X4,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_32,plain,(m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,plain,(~v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_34, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|k6_domain_1(X3,X4)=k1_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_35,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36,plain,(r2_hidden(esk1_4(X1,X2,X3,X4),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_37, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))))), inference(assume_negation,[status(cth)],[t18_waybel_0])).
cnf(c_0_38,plain,(v2_struct_0(X1)|r2_hidden(X2,k6_waybel_0(X1,X3))|~r1_orders_2(X1,X4,X2)|~r2_hidden(X4,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_39,plain,(k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41,plain,(r1_orders_2(X1,esk1_4(X1,X2,X3,X4),X4)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_43,plain,(r2_hidden(esk1_4(X1,X2,k4_waybel_0(X1,X2),X3),X2)|~r2_hidden(X3,k4_waybel_0(X1,X2))|~m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_36])).
fof(c_0_44, plain, (![X4]:![X5]:![X6]:(~r2_hidden(X4,X5)|(~m1_subset_1(X5,k1_zfmisc_1(X6))|~v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])])])).
fof(c_0_45, negated_conjecture, (((~v2_struct_0(esk5_0)&l1_orders_2(esk5_0))&(m1_subset_1(esk6_0,u1_struct_0(esk5_0))&(m1_subset_1(esk7_0,u1_struct_0(esk5_0))&((~r2_hidden(esk7_0,k6_waybel_0(esk5_0,esk6_0))|~r1_orders_2(esk5_0,esk6_0,esk7_0))&(r2_hidden(esk7_0,k6_waybel_0(esk5_0,esk6_0))|r1_orders_2(esk5_0,esk6_0,esk7_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_37])])])])])])).
cnf(c_0_46,plain,(v2_struct_0(X1)|r2_hidden(X2,k6_waybel_0(X1,X3))|~r1_orders_2(X1,X4,X2)|~r2_hidden(X4,k1_tarski(X3))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_33])).
cnf(c_0_47,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_40])])).
cnf(c_0_48,plain,(r1_orders_2(X1,esk1_4(X1,X2,k4_waybel_0(X1,X2),X3),X3)|~r2_hidden(X3,k4_waybel_0(X1,X2))|~m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_49,plain,(esk1_4(X1,k1_tarski(X2),k4_waybel_0(X1,k1_tarski(X2)),X3)=X2|~r2_hidden(X3,k4_waybel_0(X1,k1_tarski(X2)))|~m1_subset_1(k4_waybel_0(X1,k1_tarski(X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50,plain,(~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51,negated_conjecture,(~r1_orders_2(esk5_0,esk6_0,esk7_0)|~r2_hidden(esk7_0,k6_waybel_0(esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52,plain,(v2_struct_0(X1)|r2_hidden(X2,k6_waybel_0(X1,X3))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53,negated_conjecture,(m1_subset_1(esk6_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54,negated_conjecture,(m1_subset_1(esk7_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55,negated_conjecture,(l1_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56,negated_conjecture,(~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57,plain,(r1_orders_2(X1,X2,X3)|~r2_hidden(X3,k4_waybel_0(X1,k1_tarski(X2)))|~m1_subset_1(k4_waybel_0(X1,k1_tarski(X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_58,plain,(k4_waybel_0(X1,k1_tarski(X2))=k6_waybel_0(X1,X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_39])).
cnf(c_0_59,plain,(v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_32, c_0_39])).
cnf(c_0_60,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,k6_waybel_0(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_50, c_0_26])).
cnf(c_0_61,negated_conjecture,(r1_orders_2(esk5_0,esk6_0,esk7_0)|r2_hidden(esk7_0,k6_waybel_0(esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_62,negated_conjecture,(~r1_orders_2(esk5_0,esk6_0,esk7_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_54]), c_0_55])]), c_0_56])).
cnf(c_0_63,plain,(v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~r2_hidden(X3,k6_waybel_0(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_26]), c_0_60])).
cnf(c_0_64,negated_conjecture,(r2_hidden(esk7_0,k6_waybel_0(esk5_0,esk6_0))), inference(sr,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_65,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_54]), c_0_53]), c_0_55])]), c_0_56]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 41
# Proof object formula steps           : 25
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 113
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 95
# ...of the previous two non-trivial   : 74
# Contextual simplify-reflections      : 29
# Paramodulations                      : 82
# Factorizations                       : 2
# Equation resolutions                 : 11
# Current number of processed clauses  : 60
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 17
# ...number of literals in the above   : 148
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 3715
# Rec. Clause-clause subsumption calls : 210
# Non-unit clause-clause subsumptions  : 41
# Unit Clause-clause subsumption calls : 38
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6094

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45380 pages
