# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t18_funct_1, conjecture, (![X1]:![X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k1_relat_1(X3)&r1_tarski(k2_relat_1(X3),X1))))))), file('out/t18_funct_1__11_premises__rerun.E_input', t18_funct_1)).
fof(t64_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0))), file('out/t18_funct_1__11_premises__rerun.E_input', t64_relat_1)).
fof(s3_funct_1__e2_25__funct_1, axiom, (![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=k1_xboole_0)))), file('out/t18_funct_1__11_premises__rerun.E_input', s3_funct_1__e2_25__funct_1)).
fof(t15_funct_1, axiom, (![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2))))), file('out/t18_funct_1__11_premises__rerun.E_input', t15_funct_1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t18_funct_1__11_premises__rerun.E_input', t2_xboole_1)).
fof(cc1_relat_1, axiom, (![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1))), file('out/t18_funct_1__11_premises__rerun.E_input', cc1_relat_1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t18_funct_1__11_premises__rerun.E_input', t69_enumset1)).
fof(l1_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t18_funct_1__11_premises__rerun.E_input', l1_zfmisc_1)).
fof(t60_relat_1, axiom, ((k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0)), file('out/t18_funct_1__11_premises__rerun.E_input', t60_relat_1)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t18_funct_1__11_premises__rerun.E_input', fc1_xboole_0)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t18_funct_1__11_premises__rerun.E_input', d1_xboole_0)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t18_funct_1__11_premises__rerun.E_input', l13_xboole_0)).
fof(c_0_12, negated_conjecture, (~(![X1]:![X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k1_relat_1(X3)&r1_tarski(k2_relat_1(X3),X1)))))))), inference(assume_negation,[status(cth)],[t18_funct_1])).
fof(c_0_13, plain, (![X2]:(((k1_relat_1(X2)!=k1_xboole_0|X2=k1_xboole_0)|~v1_relat_1(X2))&((k2_relat_1(X2)!=k1_xboole_0|X2=k1_xboole_0)|~v1_relat_1(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
fof(c_0_14, plain, (![X4]:![X6]:((v1_relat_1(esk3_1(X4))&v1_funct_1(esk3_1(X4)))&(k1_relat_1(esk3_1(X4))=X4&(~r2_hidden(X6,X4)|k1_funct_1(esk3_1(X4),X6)=k1_xboole_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_25__funct_1])])])])])])).
fof(c_0_15, negated_conjecture, (![X6]:((esk1_0!=k1_xboole_0|esk2_0=k1_xboole_0)&((~v1_relat_1(X6)|~v1_funct_1(X6))|(esk2_0!=k1_relat_1(X6)|~r1_tarski(k2_relat_1(X6),esk1_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])])).
fof(c_0_16, plain, (![X4]:![X5]:(((v1_relat_1(esk4_2(X4,X5))|X4=k1_xboole_0)&(v1_funct_1(esk4_2(X4,X5))|X4=k1_xboole_0))&((k1_relat_1(esk4_2(X4,X5))=X4|X4=k1_xboole_0)&(k2_relat_1(esk4_2(X4,X5))=k1_tarski(X5)|X4=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_funct_1])])])])])])])])).
fof(c_0_17, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_18, plain, (![X2]:(~v1_xboole_0(X2)|v1_relat_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
cnf(c_0_19,plain,(X1=k1_xboole_0|~v1_relat_1(X1)|k1_relat_1(X1)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k1_relat_1(esk3_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(v1_relat_1(esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,negated_conjecture,(~r1_tarski(k2_relat_1(X1),esk1_0)|esk2_0!=k1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(X1=k1_xboole_0|k1_relat_1(esk4_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,plain,(X1=k1_xboole_0|v1_relat_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,plain,(X1=k1_xboole_0|v1_funct_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_27, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])])])).
cnf(c_0_28,plain,(k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_29,plain,(k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_30,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31,plain,(v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_33,plain,(esk3_1(X1)=k1_xboole_0|X1!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_34,negated_conjecture,(X1=k1_xboole_0|X1!=esk2_0|~r1_tarski(k2_relat_1(esk4_2(X1,X2)),esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_35,plain,(X1=k1_xboole_0|k2_relat_1(esk4_2(X1,X2))=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_36,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_38, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk5_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_39,negated_conjecture,(esk2_0!=k1_xboole_0|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_40,plain,(v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41,plain,(v1_funct_1(esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_42,plain,(esk3_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_43,negated_conjecture,(esk2_0=k1_xboole_0|~r1_tarski(k2_relat_1(esk4_2(esk2_0,X1)),esk1_0)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_44,plain,(X1=k1_xboole_0|k2_relat_1(esk4_2(X1,X2))=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45,plain,(r1_tarski(k2_tarski(X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_46,plain,(v1_xboole_0(X1)|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47,negated_conjecture,(esk2_0!=k1_xboole_0|~v1_funct_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40])])).
cnf(c_0_48,plain,(v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_49, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_50,negated_conjecture,(esk2_0=k1_xboole_0|~r1_tarski(k2_tarski(X1,X1),esk1_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51,plain,(v1_xboole_0(X1)|r1_tarski(k2_tarski(esk5_1(X1),esk5_1(X1)),X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52,negated_conjecture,(esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_53,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54,negated_conjecture,(v1_xboole_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_55,negated_conjecture,(esk2_0=k1_xboole_0|esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_56,negated_conjecture,(esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_57,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56])]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 35
# Proof object formula steps           : 23
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 42
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 40
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 11
# Generated clauses                    : 28
# ...of the previous two non-trivial   : 32
# Contextual simplify-reflections      : 4
# Paramodulations                      : 25
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 11
# ...number of literals in the above   : 26
# Current number of archived formulas  : 0
# Current number of archived clauses   : 12
# Clause-clause subsumption calls (NU) : 79
# Rec. Clause-clause subsumption calls : 60
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 31
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1294

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45400 pages
