# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t18_zfmisc_1, conjecture, (![X1]:![X2]:(k3_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t18_zfmisc_1)).
fof(t62_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t62_enumset1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t70_enumset1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t100_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t75_enumset1)).
fof(t92_xboole_1, axiom, (![X1]:k5_xboole_0(X1,X1)=k1_xboole_0), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t92_xboole_1)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t5_boole)).
fof(t134_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t134_enumset1)).
fof(t68_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t68_enumset1)).
fof(t125_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t125_enumset1)).
fof(t133_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k2_tarski(X8,X9))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t133_enumset1)).
fof(t51_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t51_enumset1)).
fof(l142_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k3_enumset1(X5,X6,X7,X8,X9))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', l142_enumset1)).
fof(t132_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t132_enumset1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', commutativity_k3_xboole_0)).
fof(commutativity_k5_xboole_0, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), file('out/t18_zfmisc_1__29_premises__rerun.E_input', commutativity_k5_xboole_0)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t91_xboole_1)).
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t4_xboole_0)).
fof(t7_xboole_0, axiom, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1))))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t7_xboole_0)).
fof(t82_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t82_xboole_1)).
fof(t136_enumset1, axiom, (![X1]:![X2]:![X3]:~((~(X1=X2)&(~(X1=X3)&~(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)))))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t136_enumset1)).
fof(t12_zfmisc_1, axiom, (![X1]:![X2]:r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t12_zfmisc_1)).
fof(fc3_xboole_0, axiom, (![X1]:![X2]:~(v1_xboole_0(k2_tarski(X1,X2)))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', fc3_xboole_0)).
fof(t69_xboole_1, axiom, (![X1]:![X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1))))), file('out/t18_zfmisc_1__29_premises__rerun.E_input', t69_xboole_1)).
fof(c_0_30, negated_conjecture, (~(![X1]:![X2]:(k3_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2))), inference(assume_negation,[status(cth)],[t18_zfmisc_1])).
fof(c_0_31, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:k6_enumset1(X9,X10,X11,X12,X13,X14,X15,X16)=k2_xboole_0(k1_tarski(X9),k5_enumset1(X10,X11,X12,X13,X14,X15,X16))), inference(variable_rename,[status(thm)],[t62_enumset1])).
fof(c_0_32, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_33, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_34, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_35, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_36, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_37, plain, (![X5]:![X6]:![X7]:![X8]:k3_enumset1(X5,X5,X6,X7,X8)=k2_enumset1(X5,X6,X7,X8)), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_38, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k4_enumset1(X6,X6,X7,X8,X9,X10)=k3_enumset1(X6,X7,X8,X9,X10)), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_39, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:k5_enumset1(X7,X7,X8,X9,X10,X11,X12)=k4_enumset1(X7,X8,X9,X10,X11,X12)), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_40, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:k6_enumset1(X8,X8,X9,X10,X11,X12,X13,X14)=k5_enumset1(X8,X9,X10,X11,X12,X13,X14)), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_41, negated_conjecture, ((k3_xboole_0(k1_tarski(esk1_0),k1_tarski(esk2_0))=k1_tarski(esk1_0)&esk1_0!=esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
cnf(c_0_42,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48,plain,(k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49,plain,(k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50,plain,(k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51,plain,(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52,negated_conjecture,(k3_xboole_0(k1_tarski(esk1_0),k1_tarski(esk2_0))=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_53, plain, (![X2]:k5_xboole_0(X2,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_54, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
fof(c_0_55, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:k7_enumset1(X10,X11,X12,X13,X14,X15,X16,X17,X18)=k2_xboole_0(k6_enumset1(X10,X11,X12,X13,X14,X15,X16,X17),k1_tarski(X18))), inference(variable_rename,[status(thm)],[t134_enumset1])).
fof(c_0_56, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:k6_enumset1(X9,X10,X11,X12,X13,X14,X15,X16)=k2_xboole_0(k5_enumset1(X9,X10,X11,X12,X13,X14,X15),k1_tarski(X16))), inference(variable_rename,[status(thm)],[t68_enumset1])).
cnf(c_0_57,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k5_xboole_0(k6_enumset1(X2,X2,X3,X4,X5,X6,X7,X8),k3_xboole_0(k6_enumset1(X2,X2,X3,X4,X5,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_58,negated_conjecture,(k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_44]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_59,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63,negated_conjecture,(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60])).
cnf(c_0_64,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k5_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k3_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
cnf(c_0_65,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k3_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_66,negated_conjecture,(k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_58, c_0_63])).
cnf(c_0_67,negated_conjecture,(k7_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1)=k7_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_63]), c_0_64])).
cnf(c_0_68,plain,(k7_enumset1(X1,X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[c_0_65, c_0_64])).
fof(c_0_69, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X8,X7,X6,X5)), inference(variable_rename,[status(thm)],[t125_enumset1])).
cnf(c_0_70,negated_conjecture,(k7_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_66]), c_0_59]), c_0_60])).
cnf(c_0_71,negated_conjecture,(k7_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_72,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
fof(c_0_73, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:k7_enumset1(X10,X11,X12,X13,X14,X15,X16,X17,X18)=k2_xboole_0(k5_enumset1(X10,X11,X12,X13,X14,X15,X16),k2_tarski(X17,X18))), inference(variable_rename,[status(thm)],[t133_enumset1])).
cnf(c_0_74,negated_conjecture,(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk1_0)=k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_75,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X4,X4,X4,X4,X4,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
cnf(c_0_76,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k2_tarski(X8,X9))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_77,negated_conjecture,(k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_78,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X9),k3_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X9),k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_79,negated_conjecture,(k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_66, c_0_77])).
fof(c_0_80, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:k4_enumset1(X7,X8,X9,X10,X11,X12)=k2_xboole_0(k1_tarski(X7),k3_enumset1(X8,X9,X10,X11,X12))), inference(variable_rename,[status(thm)],[t51_enumset1])).
cnf(c_0_81,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_59]), c_0_60]), c_0_68])).
cnf(c_0_82,plain,(k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_83,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_81]), c_0_57])).
cnf(c_0_84,plain,(k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k5_xboole_0(k6_enumset1(X2,X2,X2,X2,X3,X4,X5,X6),k3_xboole_0(k6_enumset1(X2,X2,X2,X2,X3,X4,X5,X6),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_85,negated_conjecture,(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_63, c_0_77])).
cnf(c_0_86,negated_conjecture,(k7_enumset1(X1,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X2)=k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_83]), c_0_64])).
cnf(c_0_87,plain,(k6_enumset1(X1,X2,X2,X2,X3,X4,X5,X6)=k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)), inference(rw,[status(thm)],[c_0_84, c_0_57])).
cnf(c_0_88,negated_conjecture,(k6_enumset1(X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_85]), c_0_57])).
cnf(c_0_89,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_86]), c_0_68])).
cnf(c_0_90,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,X1,X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_91,negated_conjecture,(k6_enumset1(X1,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_89]), c_0_57])).
fof(c_0_92, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:k7_enumset1(X10,X11,X12,X13,X14,X15,X16,X17,X18)=k2_xboole_0(k2_enumset1(X10,X11,X12,X13),k3_enumset1(X14,X15,X16,X17,X18))), inference(variable_rename,[status(thm)],[l142_enumset1])).
cnf(c_0_93,negated_conjecture,(k6_enumset1(X1,X1,X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,X1,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_87, c_0_90])).
cnf(c_0_94,negated_conjecture,(k7_enumset1(X1,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X2,X3)=k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_91]), c_0_64])).
cnf(c_0_95,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k3_enumset1(X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_92])).
fof(c_0_96, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:k7_enumset1(X10,X11,X12,X13,X14,X15,X16,X17,X18)=k2_xboole_0(k4_enumset1(X10,X11,X12,X13,X14,X15),k1_enumset1(X16,X17,X18))), inference(variable_rename,[status(thm)],[t132_enumset1])).
cnf(c_0_97,negated_conjecture,(k6_enumset1(X1,X2,X2,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,X2,X2,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_93]), c_0_57])).
cnf(c_0_98,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1,X2)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_94]), c_0_68]), c_0_87])).
cnf(c_0_99,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k5_xboole_0(k6_enumset1(X5,X5,X5,X5,X6,X7,X8,X9),k3_xboole_0(k6_enumset1(X5,X5,X5,X5,X6,X7,X8,X9),k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_45]), c_0_46]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_100,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_101,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_102,negated_conjecture,(k7_enumset1(esk1_0,esk2_0,esk2_0,esk2_0,X1,X1,X1,X1,X1)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_77]), c_0_99]), c_0_71])).
cnf(c_0_103,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6),k5_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X8,X9),k3_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X8,X9),k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_45]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_104,negated_conjecture,(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_101]), c_0_99]), c_0_102]), c_0_68])).
cnf(c_0_105,negated_conjecture,(k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_85, c_0_88])).
cnf(c_0_106,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_101]), c_0_57])).
cnf(c_0_107,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_79]), c_0_59]), c_0_60]), c_0_68])).
cnf(c_0_108,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_104]), c_0_87]), c_0_105])).
cnf(c_0_109,negated_conjecture,(k7_enumset1(X1,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X2)=k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_106]), c_0_64])).
cnf(c_0_110,negated_conjecture,(k6_enumset1(X1,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,esk1_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_107]), c_0_57])).
cnf(c_0_111,negated_conjecture,(k6_enumset1(X1,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_108]), c_0_57])).
cnf(c_0_112,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_109]), c_0_68])).
cnf(c_0_113,negated_conjecture,(k6_enumset1(X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X2)=k6_enumset1(X1,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_104]), c_0_57])).
cnf(c_0_114,negated_conjecture,(k7_enumset1(X1,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,esk1_0,X2)=k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_110]), c_0_64])).
cnf(c_0_115,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_116,negated_conjecture,(k6_enumset1(esk2_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_104, c_0_113])).
cnf(c_0_117,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2,X3)=k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_98]), c_0_57])).
cnf(c_0_118,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_114]), c_0_68]), c_0_87])).
cnf(c_0_119,negated_conjecture,(k7_enumset1(esk1_0,esk2_0,esk2_0,esk2_0,X1,X1,X1,X1,X1)=k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_102, c_0_104])).
cnf(c_0_120,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_115]), c_0_99]), c_0_102]), c_0_113]), c_0_116]), c_0_112]), c_0_68]), c_0_117])).
cnf(c_0_121,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_118]), c_0_57])).
cnf(c_0_122,negated_conjecture,(k7_enumset1(esk1_0,esk2_0,esk2_0,esk2_0,X1,X1,X1,X1,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_119, c_0_112])).
cnf(c_0_123,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_78]), c_0_68])).
cnf(c_0_124,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_103]), c_0_68])).
cnf(c_0_125,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_120]), c_0_57])).
cnf(c_0_126,negated_conjecture,(k7_enumset1(X1,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X2,X3)=k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_121]), c_0_64])).
cnf(c_0_127,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_79]), c_0_59]), c_0_60]), c_0_68])).
cnf(c_0_128,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,X1,X1,X1,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123]), c_0_87])).
cnf(c_0_129,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_124]), c_0_124]), c_0_125])).
cnf(c_0_130,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X1,X2)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_126]), c_0_68])).
cnf(c_0_131,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_127]), c_0_57])).
cnf(c_0_132,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,X1,X2,X2,X3)=k6_enumset1(esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X1,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_121]), c_0_103]), c_0_68]), c_0_68]), c_0_87])).
cnf(c_0_133,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,X1,X1,X1,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_134,negated_conjecture,(k6_enumset1(X1,esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X2,X3)=k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_130]), c_0_57])).
cnf(c_0_135,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X7)=k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_103]), c_0_123]), c_0_68])).
cnf(c_0_136,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,X1,X1,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_87, c_0_131])).
cnf(c_0_137,negated_conjecture,(k6_enumset1(esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X1,X1,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_138,negated_conjecture,(k6_enumset1(X1,esk1_0,esk2_0,esk1_0,esk1_0,esk1_0,X2,X3)=k6_enumset1(X1,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,X2,X3)), inference(rw,[status(thm)],[c_0_134, c_0_117])).
cnf(c_0_139,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_135]), c_0_135]), c_0_135])).
cnf(c_0_140,negated_conjecture,(k6_enumset1(X1,X1,X1,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0)=k6_enumset1(X1,X1,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_87, c_0_136])).
cnf(c_0_141,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_135]), c_0_138]), c_0_135]), c_0_139])).
fof(c_0_142, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_143,negated_conjecture,(k6_enumset1(X1,X2,X2,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0)=k6_enumset1(X1,X2,X2,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_140]), c_0_57])).
fof(c_0_144, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k5_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_145,plain,(k6_enumset1(X1,X2,X3,X4,X5,X5,X5,X5)=k6_enumset1(X4,X4,X4,X4,X3,X2,X1,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_75]), c_0_99]), c_0_68]), c_0_124])).
cnf(c_0_146,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_141]), c_0_57])).
cnf(c_0_147,plain,(k6_enumset1(X1,X2,X3,X4,X5,X5,X5,X6)=k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_99]), c_0_124]), c_0_68])).
cnf(c_0_148,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_149,negated_conjecture,(k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_123]), c_0_117]), c_0_121]), c_0_143])).
fof(c_0_150, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_151,plain,(k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_152,negated_conjecture,(k6_enumset1(esk2_0,esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_145]), c_0_135]), c_0_138]), c_0_135]), c_0_146]), c_0_147])).
cnf(c_0_153,negated_conjecture,(k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,X2)=k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[c_0_114, c_0_94])).
cnf(c_0_154,plain,(k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X9),k3_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X9))))=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_78, c_0_148])).
cnf(c_0_155,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_149, c_0_124])).
cnf(c_0_156,negated_conjecture,(k6_enumset1(X1,esk2_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_88, c_0_97])).
fof(c_0_157, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk3_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
fof(c_0_158, plain, (![X3]:(X3=k1_xboole_0|r2_hidden(esk4_1(X3),X3))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])])])).
fof(c_0_159, plain, (![X3]:![X4]:r1_xboole_0(k4_xboole_0(X3,X4),k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t82_xboole_1])).
cnf(c_0_160,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_161,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_60, c_0_151])).
fof(c_0_162, plain, (![X4]:![X5]:![X6]:(X4=X5|(X4=X6|k4_xboole_0(k1_enumset1(X4,X5,X6),k1_tarski(X4))=k2_tarski(X5,X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_enumset1])])])])])).
cnf(c_0_163,plain,(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k5_xboole_0(k6_enumset1(X2,X3,X3,X3,X4,X5,X6,X7),k3_xboole_0(k6_enumset1(X2,X3,X3,X3,X4,X5,X6,X7),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))))=k6_enumset1(X1,X2,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_57, c_0_87])).
cnf(c_0_164,negated_conjecture,(k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_135, c_0_152])).
cnf(c_0_165,negated_conjecture,(k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[c_0_153, c_0_124])).
cnf(c_0_166,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_139]), c_0_57])).
cnf(c_0_167,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk1_0,X1,X2)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk1_0,X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_154]), c_0_68]), c_0_68]), c_0_117])).
cnf(c_0_168,negated_conjecture,(k7_enumset1(X1,esk2_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)=k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_156]), c_0_64])).
cnf(c_0_169,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_170,plain,(r2_hidden(esk4_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_158])).
cnf(c_0_171,plain,(r1_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_172,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_59]), c_0_161])).
cnf(c_0_173,plain,(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)|X1=X3|X1=X2), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_174,negated_conjecture,(k6_enumset1(X1,esk2_0,esk2_0,esk1_0,esk1_0,esk1_0,X2,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_57])).
cnf(c_0_175,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X7)=k7_enumset1(X1,X2,X2,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_87]), c_0_103]), c_0_123])).
cnf(c_0_176,negated_conjecture,(k7_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X2)), inference(rw,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_177,plain,(k7_enumset1(X1,X2,X3,X4,X5,X5,X5,X6,X7)=k6_enumset1(X4,X4,X3,X2,X1,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_75]), c_0_99]), c_0_68])).
cnf(c_0_178,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk1_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_135]), c_0_135]), c_0_141])).
cnf(c_0_179,plain,(k5_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k5_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k3_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9))))=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_64, c_0_148])).
cnf(c_0_180,negated_conjecture,(k7_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)=k6_enumset1(esk2_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_68, c_0_168])).
cnf(c_0_181,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_169, c_0_170])).
cnf(c_0_182,plain,(r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_171, c_0_46]), c_0_46])).
cnf(c_0_183,plain,(k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_172, c_0_151])).
cnf(c_0_184,negated_conjecture,(k6_enumset1(esk2_0,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_116, c_0_112])).
cnf(c_0_185,negated_conjecture,(k6_enumset1(X1,esk1_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_112]), c_0_57])).
fof(c_0_186, plain, (![X3]:![X4]:r1_tarski(k1_tarski(X3),k2_tarski(X3,X4))), inference(variable_rename,[status(thm)],[t12_zfmisc_1])).
fof(c_0_187, plain, (![X3]:![X4]:~v1_xboole_0(k2_tarski(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
cnf(c_0_188,plain,(X1=X3|X1=X2|k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_173, c_0_43]), c_0_44]), c_0_44]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_189,plain,(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))=k5_xboole_0(k6_enumset1(X2,X2,X3,X4,X5,X6,X7,X8),k3_xboole_0(k6_enumset1(X2,X2,X3,X4,X5,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))), inference(spm,[status(thm)],[c_0_172, c_0_57])).
cnf(c_0_190,negated_conjecture,(k6_enumset1(X1,esk2_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175]), c_0_124])).
cnf(c_0_191,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X2)), inference(rw,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_192,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk1_0,X2)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_178]), c_0_57])).
cnf(c_0_193,plain,(k5_xboole_0(X1,k5_xboole_0(X2,X3))=k5_xboole_0(X3,k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_151, c_0_160])).
cnf(c_0_194,negated_conjecture,(k6_enumset1(esk2_0,esk2_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_105]), c_0_179]), c_0_68]), c_0_180])).
cnf(c_0_195,plain,(k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_181, c_0_182])).
cnf(c_0_196,plain,(k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X1,X2,X3,X4,X5,X6,X7),k3_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8))))=k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_57]), c_0_160]), c_0_151])).
cnf(c_0_197,negated_conjecture,(k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_184, c_0_185])).
fof(c_0_198, plain, (![X3]:![X4]:(v1_xboole_0(X4)|(~r1_tarski(X4,X3)|~r1_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
cnf(c_0_199,plain,(r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_200,plain,(~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_201,plain,(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)|X1=X2|X1=X3), inference(rw,[status(thm)],[c_0_188, c_0_189])).
cnf(c_0_202,negated_conjecture,(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X1,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_190]), c_0_191]), c_0_192])).
cnf(c_0_203,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X8,X1,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_193]), c_0_151]), c_0_154]), c_0_123])).
cnf(c_0_204,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_139]), c_0_139]), c_0_135])).
cnf(c_0_205,negated_conjecture,(k7_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[c_0_180, c_0_194])).
cnf(c_0_206,plain,(r2_hidden(esk3_2(X1,X2),k3_xboole_0(X1,X2))|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_207,plain,(k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X1,X2)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_195, c_0_148])).
cnf(c_0_208,plain,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_195]), c_0_182])])).
cnf(c_0_209,negated_conjecture,(k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,X1))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_88]), c_0_105]), c_0_172]), c_0_88]), c_0_105]), c_0_148])).
cnf(c_0_210,plain,(v1_xboole_0(X1)|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_211,plain,(r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_199, c_0_43]), c_0_44]), c_0_44]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
cnf(c_0_212,plain,(~v1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_44]), c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_213,plain,(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3))=k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)|X1=X3|X1=X2), inference(spm,[status(thm)],[c_0_183, c_0_201])).
cnf(c_0_214,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,X1,esk2_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_202, c_0_203]), c_0_87]), c_0_203])).
cnf(c_0_215,negated_conjecture,(k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk1_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_204]), c_0_57])).
cnf(c_0_216,negated_conjecture,(k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_205, c_0_124]), c_0_129])).
cnf(c_0_217,negated_conjecture,(esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_218,plain,(r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X1,X2)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207]), c_0_208])).
cnf(c_0_219,negated_conjecture,(k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_209, c_0_143])).
cnf(c_0_220,plain,(~r1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_212])).
cnf(c_0_221,negated_conjecture,(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214]), c_0_215]), c_0_88]), c_0_155]), c_0_216]), c_0_139]), c_0_215]), c_0_59]), c_0_217])).
cnf(c_0_222,negated_conjecture,(r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219]), c_0_59]), c_0_59])).
cnf(c_0_223,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221]), c_0_221]), c_0_222])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 224
# Proof object clause steps            : 163
# Proof object formula steps           : 61
# Proof object conjectures             : 97
# Proof object clause conjectures      : 94
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 30
# Proof object generating inferences   : 88
# Proof object simplifying inferences  : 359
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 30
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 23
# Processed clauses                    : 9621
# ...of these trivial                  : 814
# ...subsumed                          : 7678
# ...remaining for further processing  : 1129
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 467
# Generated clauses                    : 202599
# ...of the previous two non-trivial   : 164439
# Contextual simplify-reflections      : 0
# Paramodulations                      : 202599
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 636
#    Positive orientable unit clauses  : 352
#    Positive unorientable unit clauses: 77
#    Negative unit clauses             : 145
#    Non-unit-clauses                  : 62
# Current number of unprocessed clauses: 116434
# ...number of literals in the above   : 161332
# Current number of archived formulas  : 0
# Current number of archived clauses   : 502
# Clause-clause subsumption calls (NU) : 1170
# Rec. Clause-clause subsumption calls : 904
# Non-unit clause-clause subsumptions  : 145
# Unit Clause-clause subsumption calls : 11945
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26123
# BW rewrite match successes           : 2333
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4615330

# -------------------------------------------------
# User time                : 2.896 s
# System time              : 0.088 s
# Total time               : 2.984 s
# Maximum resident set size: 112136 pages
