# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t193_funct_2, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>![X5]:((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))=>![X6]:(m1_subset_1(X6,X2)=>(v1_partfun1(X5,X1)=>k3_funct_2(X2,X3,k8_funct_2(X2,X1,X3,X4,X5),X6)=k7_partfun1(X3,X5,k3_funct_2(X2,X1,X4,X6))))))))), file('out/t193_funct_2__15_premises__rerun.E_input', t193_funct_2)).
fof(cc2_partfun1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>~(v1_partfun1(X3,X1))))), file('out/t193_funct_2__15_premises__rerun.E_input', cc2_partfun1)).
fof(t176_funct_2, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:(m1_subset_1(X4,X1)=>k7_partfun1(X2,X3,X4)=k3_funct_2(X1,X2,X3,X4)))))), file('out/t193_funct_2__15_premises__rerun.E_input', t176_funct_2)).
fof(redefinition_k3_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4))), file('out/t193_funct_2__15_premises__rerun.E_input', redefinition_k3_funct_2)).
fof(dt_k3_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>m1_subset_1(k3_funct_2(X1,X2,X3,X4),X2))), file('out/t193_funct_2__15_premises__rerun.E_input', dt_k3_funct_2)).
fof(cc2_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2)))), file('out/t193_funct_2__15_premises__rerun.E_input', cc2_relset_1)).
fof(cc2_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2)))), file('out/t193_funct_2__15_premises__rerun.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, (![X1]:![X2]:v1_relat_1(k2_zfmisc_1(X1,X2))), file('out/t193_funct_2__15_premises__rerun.E_input', fc6_relat_1)).
fof(t185_funct_2, axiom, (![X1]:![X2]:![X3]:(~(v1_xboole_0(X3))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>![X5]:((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))=>![X6]:(m1_subset_1(X6,X2)=>(r1_tarski(k2_relset_1(X2,X3,X4),k1_relset_1(X3,X1,X5))=>(X2=k1_xboole_0|k1_funct_1(k8_funct_2(X2,X3,X1,X4,X5),X6)=k1_funct_1(X5,k1_funct_1(X4,X6))))))))), file('out/t193_funct_2__15_premises__rerun.E_input', t185_funct_2)).
fof(redefinition_k1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3))), file('out/t193_funct_2__15_premises__rerun.E_input', redefinition_k1_relset_1)).
fof(d4_partfun1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>(v1_partfun1(X2,X1)<=>k1_relat_1(X2)=X1))), file('out/t193_funct_2__15_premises__rerun.E_input', d4_partfun1)).
fof(dt_k8_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>(v1_funct_1(k8_funct_2(X1,X2,X3,X4,X5))&(v1_funct_2(k8_funct_2(X1,X2,X3,X4,X5),X1,X3)&m1_subset_1(k8_funct_2(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X3))))))), file('out/t193_funct_2__15_premises__rerun.E_input', dt_k8_funct_2)).
fof(redefinition_k2_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3))), file('out/t193_funct_2__15_premises__rerun.E_input', redefinition_k2_relset_1)).
fof(d19_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1)))), file('out/t193_funct_2__15_premises__rerun.E_input', d19_relat_1)).
fof(cc1_funct_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2))))), file('out/t193_funct_2__15_premises__rerun.E_input', cc1_funct_2)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t193_funct_2__15_premises__rerun.E_input', fc1_xboole_0)).
fof(c_0_16, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>![X5]:((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))=>![X6]:(m1_subset_1(X6,X2)=>(v1_partfun1(X5,X1)=>k3_funct_2(X2,X3,k8_funct_2(X2,X1,X3,X4,X5),X6)=k7_partfun1(X3,X5,k3_funct_2(X2,X1,X4,X6)))))))))), inference(assume_negation,[status(cth)],[t193_funct_2])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:((v1_xboole_0(X4)|~v1_xboole_0(X5))|(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_partfun1(X6,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_partfun1])])])])])])).
fof(c_0_18, negated_conjecture, ((~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk2_0,esk1_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))))&((v1_funct_1(esk5_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0))))&(m1_subset_1(esk6_0,esk2_0)&(v1_partfun1(esk5_0,esk1_0)&k3_funct_2(esk2_0,esk3_0,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)!=k7_partfun1(esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])])).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:(v1_xboole_0(X5)|(v1_xboole_0(X6)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|(~m1_subset_1(X8,X5)|k7_partfun1(X6,X7,X8)=k3_funct_2(X5,X6,X7,X8)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t176_funct_2])])])])])])).
cnf(c_0_20,plain,(v1_xboole_0(X2)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22,negated_conjecture,(v1_partfun1(esk5_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X8,X5)))|k3_funct_2(X5,X6,X7,X8)=k1_funct_1(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
fof(c_0_25, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X8,X5)))|m1_subset_1(k3_funct_2(X5,X6,X7,X8),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_funct_2])])])).
cnf(c_0_26,negated_conjecture,(k3_funct_2(esk2_0,esk3_0,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)!=k7_partfun1(esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(k7_partfun1(X1,X2,X3)=k3_funct_2(X4,X1,X2,X3)|v1_xboole_0(X1)|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~v1_funct_2(X2,X4,X1)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,negated_conjecture,(v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29,negated_conjecture,(~v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_30,plain,(k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(m1_subset_1(k3_funct_2(X1,X2,X3,X4),X2)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,negated_conjecture,(v1_xboole_0(X1)|k3_funct_2(X1,esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))!=k3_funct_2(esk2_0,esk3_0,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0),X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X1,esk3_0)))|~v1_funct_2(esk5_0,X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])]), c_0_29])).
cnf(c_0_33,plain,(k3_funct_2(X1,X2,X3,X4)=k3_funct_2(X5,X6,X3,X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X5)|~v1_funct_2(X3,X1,X2)|~v1_funct_2(X3,X5,X6)|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_30, c_0_30])).
cnf(c_0_34,negated_conjecture,(m1_subset_1(esk6_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_35,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_36,plain,(m1_subset_1(k1_funct_1(X1,X2),X3)|v1_xboole_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~m1_subset_1(X2,X4)|~v1_funct_2(X1,X4,X3)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_37,negated_conjecture,(v1_funct_2(esk4_0,esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_38,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_39,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_40, plain, (![X4]:![X5]:![X6]:((v4_relat_1(X6,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(v5_relat_1(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_41, plain, (![X3]:![X4]:(~v1_relat_1(X3)|(~m1_subset_1(X4,k1_zfmisc_1(X3))|v1_relat_1(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])])])).
fof(c_0_42, plain, (![X3]:![X4]:v1_relat_1(k2_zfmisc_1(X3,X4))), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_43,negated_conjecture,(v1_xboole_0(X1)|v1_xboole_0(X2)|k3_funct_2(X2,esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))!=k3_funct_2(X1,X3,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0),X2)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X2,esk3_0)))|~m1_subset_1(esk6_0,X1)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),X1,X3)|~v1_funct_2(esk5_0,X2,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(k1_funct_1(esk4_0,X1),esk1_0)|~m1_subset_1(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39])]), c_0_35])).
fof(c_0_45, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(v1_xboole_0(X9)|((~v1_funct_1(X10)|(~v1_funct_2(X10,X8,X9)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))))|((~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X7))))|(~m1_subset_1(X12,X8)|(~r1_tarski(k2_relset_1(X8,X9,X10),k1_relset_1(X9,X7,X11))|(X8=k1_xboole_0|k1_funct_1(k8_funct_2(X8,X9,X7,X10,X11),X12)=k1_funct_1(X11,k1_funct_1(X10,X12))))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t185_funct_2])])])])])])).
fof(c_0_46, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|k1_relset_1(X4,X5,X6)=k1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_47, plain, (![X3]:![X4]:(((~v1_partfun1(X4,X3)|k1_relat_1(X4)=X3)|(~v1_relat_1(X4)|~v4_relat_1(X4,X3)))&((k1_relat_1(X4)!=X3|v1_partfun1(X4,X3))|(~v1_relat_1(X4)|~v4_relat_1(X4,X3))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_partfun1])])])).
cnf(c_0_48,plain,(v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50,plain,(v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51,negated_conjecture,(v1_xboole_0(X1)|v1_xboole_0(X2)|k7_partfun1(esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))!=k3_funct_2(X2,X3,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0),X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X1,esk3_0)))|~m1_subset_1(esk6_0,X2)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),X2,X3)|~v1_funct_2(esk5_0,X1,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_27]), c_0_28])]), c_0_29])).
cnf(c_0_52,negated_conjecture,(m1_subset_1(k3_funct_2(X1,X2,esk4_0,X3),esk1_0)|v1_xboole_0(X1)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X3,X1)|~v1_funct_2(esk4_0,X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_30]), c_0_39])])).
cnf(c_0_53,plain,(k1_funct_1(k8_funct_2(X1,X2,X3,X4,X5),X6)=k1_funct_1(X5,k1_funct_1(X4,X6))|X1=k1_xboole_0|v1_xboole_0(X2)|~r1_tarski(k2_relset_1(X1,X2,X4),k1_relset_1(X2,X3,X5))|~m1_subset_1(X6,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X4,X1,X2)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54,plain,(k1_relset_1(X1,X2,X3)=k1_relat_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55,plain,(k1_relat_1(X1)=X2|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_partfun1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56,negated_conjecture,(v4_relat_1(esk5_0,esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_21])).
cnf(c_0_57,negated_conjecture,(v1_relat_1(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_21]), c_0_50])])).
cnf(c_0_58,negated_conjecture,(v1_xboole_0(X1)|k7_partfun1(esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))!=k3_funct_2(X1,X2,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk6_0,X1)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),X1,X2)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_21]), c_0_38]), c_0_34]), c_0_37])]), c_0_23]), c_0_35])).
fof(c_0_59, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((v1_funct_1(k8_funct_2(X6,X7,X8,X9,X10))|((~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))|(~v1_funct_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&((v1_funct_2(k8_funct_2(X6,X7,X8,X9,X10),X6,X8)|((~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))|(~v1_funct_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))&(m1_subset_1(k8_funct_2(X6,X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X6,X8)))|((~v1_funct_1(X9)|(~v1_funct_2(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))|(~v1_funct_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_funct_2])])])).
cnf(c_0_60,plain,(k1_funct_1(k8_funct_2(X1,X2,X3,X4,X5),X6)=k1_funct_1(X5,k1_funct_1(X4,X6))|X1=k1_xboole_0|v1_xboole_0(X2)|~r1_tarski(k2_relset_1(X1,X2,X4),k1_relat_1(X5))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,X1)|~v1_funct_2(X4,X1,X2)|~v1_funct_1(X5)|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61,negated_conjecture,(k1_relat_1(esk5_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), c_0_22])])).
fof(c_0_62, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|k2_relset_1(X4,X5,X6)=k2_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_63,negated_conjecture,(v1_xboole_0(X1)|v1_xboole_0(X2)|k7_partfun1(esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))!=k7_partfun1(X1,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(esk6_0,X2)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),X2,X1)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_58, c_0_27])).
cnf(c_0_64,plain,(v1_funct_2(k8_funct_2(X5,X2,X3,X4,X1),X5,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~v1_funct_2(X4,X5,X2)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65,plain,(k7_partfun1(X1,X2,X3)=k1_funct_1(X2,X3)|v1_xboole_0(X1)|v1_xboole_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(X3,X4)|~v1_funct_2(X2,X4,X1)|~v1_funct_1(X2)), inference(spm,[status(thm)],[c_0_30, c_0_27])).
cnf(c_0_66,plain,(v1_funct_1(k8_funct_2(X5,X2,X3,X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~v1_funct_2(X4,X5,X2)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67,plain,(m1_subset_1(k8_funct_2(X5,X2,X3,X4,X1),k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~v1_funct_2(X4,X5,X2)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68,negated_conjecture,(k1_funct_1(k8_funct_2(X1,X2,X3,X4,esk5_0),X5)=k1_funct_1(esk5_0,k1_funct_1(X4,X5))|X1=k1_xboole_0|v1_xboole_0(X2)|~r1_tarski(k2_relset_1(X1,X2,X4),esk1_0)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,X2)|~v1_funct_1(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_28])])).
cnf(c_0_69,plain,(k2_relset_1(X1,X2,X3)=k2_relat_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_70, plain, (![X3]:![X4]:(((~v5_relat_1(X4,X3)|r1_tarski(k2_relat_1(X4),X3))|~v1_relat_1(X4))&((~r1_tarski(k2_relat_1(X4),X3)|v5_relat_1(X4,X3))|~v1_relat_1(X4)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
cnf(c_0_71,negated_conjecture,(k7_partfun1(esk3_0,k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)!=k7_partfun1(esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_34]), c_0_38]), c_0_21]), c_0_37]), c_0_39]), c_0_28])]), c_0_29]), c_0_35])).
cnf(c_0_72,plain,(k7_partfun1(X1,k8_funct_2(X2,X3,X1,X4,X5),X6)=k1_funct_1(k8_funct_2(X2,X3,X1,X4,X5),X6)|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~m1_subset_1(X6,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4)|~v1_funct_1(X5)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_64]), c_0_66]), c_0_67])).
cnf(c_0_73,negated_conjecture,(k1_funct_1(k8_funct_2(X1,X2,X3,X4,esk5_0),X5)=k1_funct_1(esk5_0,k1_funct_1(X4,X5))|X1=k1_xboole_0|v1_xboole_0(X2)|~r1_tarski(k2_relat_1(X4),esk1_0)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,X2)|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_74,plain,(r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_75,plain,(v5_relat_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_76,negated_conjecture,(k7_partfun1(esk3_0,esk5_0,k3_funct_2(esk2_0,esk1_0,esk4_0,esk6_0))!=k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_38]), c_0_21]), c_0_34]), c_0_37]), c_0_39]), c_0_28])]), c_0_35]), c_0_29])).
cnf(c_0_77,negated_conjecture,(k1_funct_1(k8_funct_2(X1,X2,X3,X4,esk5_0),X5)=k1_funct_1(esk5_0,k1_funct_1(X4,X5))|X1=k1_xboole_0|v1_xboole_0(X2)|~v1_relat_1(X4)|~v5_relat_1(X4,esk1_0)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,X2)|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_78,negated_conjecture,(v5_relat_1(esk4_0,esk1_0)), inference(spm,[status(thm)],[c_0_75, c_0_38])).
cnf(c_0_79,negated_conjecture,(v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_38]), c_0_50])])).
cnf(c_0_80,negated_conjecture,(v1_xboole_0(X1)|k7_partfun1(esk3_0,esk5_0,k3_funct_2(X1,X2,esk4_0,esk6_0))!=k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk6_0,X1)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_2(esk4_0,X1,X2)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_33]), c_0_38]), c_0_34]), c_0_37]), c_0_39])]), c_0_35])).
fof(c_0_81, plain, (![X4]:![X5]:![X6]:(((v1_funct_1(X6)|(~v1_funct_1(X6)|~v1_partfun1(X6,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&((v1_funct_2(X6,X4,X5)|(~v1_funct_1(X6)|~v1_partfun1(X6,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_82,negated_conjecture,(k1_funct_1(k8_funct_2(X1,X2,X3,esk4_0,esk5_0),X4)=k1_funct_1(esk5_0,k1_funct_1(esk4_0,X4))|X1=k1_xboole_0|v1_xboole_0(X2)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(esk4_0,X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79]), c_0_39])])).
cnf(c_0_83,negated_conjecture,(v1_xboole_0(X1)|v1_xboole_0(X2)|k7_partfun1(esk3_0,esk5_0,k7_partfun1(X1,esk4_0,esk6_0))!=k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(esk6_0,X2)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_2(esk4_0,X2,X1)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_27]), c_0_39])])).
cnf(c_0_84,plain,(v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_partfun1(X1,X2)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_85,negated_conjecture,(k1_funct_1(k8_funct_2(X1,esk1_0,esk3_0,esk4_0,esk5_0),X2)=k1_funct_1(esk5_0,k1_funct_1(esk4_0,X2))|X1=k1_xboole_0|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,esk1_0)))|~m1_subset_1(X2,X1)|~v1_funct_2(esk4_0,X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_21]), c_0_23])).
cnf(c_0_86,negated_conjecture,(k7_partfun1(esk3_0,esk5_0,k7_partfun1(esk1_0,esk4_0,esk6_0))!=k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_37]), c_0_38]), c_0_34])]), c_0_23]), c_0_35])).
cnf(c_0_87,negated_conjecture,(k7_partfun1(esk1_0,esk4_0,X1)=k1_funct_1(esk4_0,X1)|~m1_subset_1(X1,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_37]), c_0_38]), c_0_39])]), c_0_23]), c_0_35])).
cnf(c_0_88,negated_conjecture,(v1_funct_2(esk5_0,esk1_0,X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_22]), c_0_28])])).
cnf(c_0_89,negated_conjecture,(k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),X1)=k1_funct_1(esk5_0,k1_funct_1(esk4_0,X1))|k1_xboole_0=esk2_0|~m1_subset_1(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_38]), c_0_37])])).
cnf(c_0_90,negated_conjecture,(k7_partfun1(esk3_0,esk5_0,k1_funct_1(esk4_0,esk6_0))!=k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_34])])).
cnf(c_0_91,negated_conjecture,(k7_partfun1(X1,esk5_0,X2)=k1_funct_1(esk5_0,X2)|v1_xboole_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|~m1_subset_1(X2,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_88]), c_0_28])]), c_0_23])).
cnf(c_0_92,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_93,negated_conjecture,(k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)=k1_funct_1(esk5_0,k1_funct_1(esk4_0,esk6_0))|k1_xboole_0=esk2_0), inference(spm,[status(thm)],[c_0_89, c_0_34])).
cnf(c_0_94,negated_conjecture,(k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)!=k1_funct_1(esk5_0,k1_funct_1(esk4_0,esk6_0))|~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(k1_funct_1(esk4_0,esk6_0),esk1_0)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_21])]), c_0_29])).
cnf(c_0_95,negated_conjecture,(k1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk6_0)=k1_funct_1(esk5_0,k1_funct_1(esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_35])).
cnf(c_0_96,negated_conjecture,(~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~m1_subset_1(k1_funct_1(esk4_0,esk6_0),esk1_0)|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_97,negated_conjecture,(~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk3_0)|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_44]), c_0_34])])).
cnf(c_0_98,negated_conjecture,(~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(esk5_0,esk1_0,esk3_0)|~v1_funct_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_64]), c_0_38]), c_0_21]), c_0_37]), c_0_39]), c_0_28])])).
cnf(c_0_99,negated_conjecture,(~m1_subset_1(k8_funct_2(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(esk5_0,esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_66]), c_0_38]), c_0_21]), c_0_37]), c_0_39]), c_0_28])])).
cnf(c_0_100,negated_conjecture,(~v1_funct_2(esk5_0,esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_67]), c_0_38]), c_0_21]), c_0_37]), c_0_39]), c_0_28])])).
cnf(c_0_101,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_88]), c_0_21])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 102
# Proof object clause steps            : 70
# Proof object formula steps           : 32
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 113
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 30
# Processed clauses                    : 833
# ...of these trivial                  : 0
# ...subsumed                          : 485
# ...remaining for further processing  : 348
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 46
# Backward-rewritten                   : 25
# Generated clauses                    : 1479
# ...of the previous two non-trivial   : 1495
# Contextual simplify-reflections      : 948
# Paramodulations                      : 1479
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 247
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 224
# Current number of unprocessed clauses: 438
# ...number of literals in the above   : 5654
# Current number of archived formulas  : 0
# Current number of archived clauses   : 101
# Clause-clause subsumption calls (NU) : 160491
# Rec. Clause-clause subsumption calls : 4014
# Non-unit clause-clause subsumptions  : 1477
# Unit Clause-clause subsumption calls : 112
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 91915

# -------------------------------------------------
# User time                : 0.204 s
# System time              : 0.004 s
# Total time               : 0.208 s
# Maximum resident set size: 44988 pages
