# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), file('out/t19_ordinal1__20_premises__rerun.E_input', t22_xboole_1)).
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t19_ordinal1__20_premises__rerun.E_input', l51_zfmisc_1)).
fof(t21_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), file('out/t19_ordinal1__20_premises__rerun.E_input', t21_xboole_1)).
fof(t14_zfmisc_1, axiom, (![X1]:![X2]:k2_xboole_0(k1_tarski(X1),k2_tarski(X1,X2))=k2_tarski(X1,X2)), file('out/t19_ordinal1__20_premises__rerun.E_input', t14_zfmisc_1)).
fof(t87_enumset1, axiom, (![X1]:k3_enumset1(X1,X1,X1,X1,X1)=k1_tarski(X1)), file('out/t19_ordinal1__20_premises__rerun.E_input', t87_enumset1)).
fof(t80_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), file('out/t19_ordinal1__20_premises__rerun.E_input', t80_enumset1)).
fof(t60_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_tarski(X6,X7))), file('out/t19_ordinal1__20_premises__rerun.E_input', t60_enumset1)).
fof(d5_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:(X8=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)<=>![X9]:(r2_hidden(X9,X8)<=>~((~(X9=X1)&(~(X9=X2)&(~(X9=X3)&(~(X9=X4)&(~(X9=X5)&(~(X9=X6)&~(X9=X7))))))))))), file('out/t19_ordinal1__20_premises__rerun.E_input', d5_enumset1)).
fof(l49_zfmisc_1, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2)))), file('out/t19_ordinal1__20_premises__rerun.E_input', l49_zfmisc_1)).
fof(d6_xboole_0, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), file('out/t19_ordinal1__20_premises__rerun.E_input', d6_xboole_0)).
fof(t40_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), file('out/t19_ordinal1__20_premises__rerun.E_input', t40_xboole_1)).
fof(t45_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1)))), file('out/t19_ordinal1__20_premises__rerun.E_input', t45_xboole_1)).
fof(t51_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), file('out/t19_ordinal1__20_premises__rerun.E_input', t51_xboole_1)).
fof(t107_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k5_xboole_0(X2,X3))<=>(r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,k3_xboole_0(X2,X3))))), file('out/t19_ordinal1__20_premises__rerun.E_input', t107_xboole_1)).
fof(l32_xboole_1, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('out/t19_ordinal1__20_premises__rerun.E_input', l32_xboole_1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t19_ordinal1__20_premises__rerun.E_input', t2_xboole_1)).
fof(t19_ordinal1, conjecture, (![X1]:![X2]:![X3]:(v1_ordinal1(X3)=>((r2_hidden(X1,X2)&r2_hidden(X2,X3))=>r2_hidden(X1,X3)))), file('out/t19_ordinal1__20_premises__rerun.E_input', t19_ordinal1)).
fof(d2_ordinal1, axiom, (![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1)))), file('out/t19_ordinal1__20_premises__rerun.E_input', d2_ordinal1)).
fof(l38_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)<=>(~(r2_hidden(X1,X3))&(r2_hidden(X2,X3)|X1=X2)))), file('out/t19_ordinal1__20_premises__rerun.E_input', l38_zfmisc_1)).
fof(t81_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_xboole_0(X1,k4_xboole_0(X2,X3))=>r1_xboole_0(X2,k4_xboole_0(X1,X3)))), file('out/t19_ordinal1__20_premises__rerun.E_input', t81_xboole_1)).
fof(t55_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((r1_xboole_0(k2_tarski(X1,X2),X3)&r2_hidden(X1,X3)))), file('out/t19_ordinal1__20_premises__rerun.E_input', t55_zfmisc_1)).
fof(c_0_21, plain, (![X3]:![X4]:k2_xboole_0(X3,k3_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t22_xboole_1])).
fof(c_0_22, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_23, plain, (![X3]:![X4]:k3_xboole_0(X3,k2_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_24, plain, (![X3]:![X4]:k2_xboole_0(k1_tarski(X3),k2_tarski(X3,X4))=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t14_zfmisc_1])).
fof(c_0_25, plain, (![X2]:k3_enumset1(X2,X2,X2,X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t87_enumset1])).
fof(c_0_26, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k5_enumset1(X6,X6,X6,X7,X8,X9,X10)=k3_enumset1(X6,X7,X8,X9,X10)), inference(variable_rename,[status(thm)],[t80_enumset1])).
fof(c_0_27, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:k5_enumset1(X8,X9,X10,X11,X12,X13,X14)=k2_xboole_0(k3_enumset1(X8,X9,X10,X11,X12),k2_tarski(X13,X14))), inference(variable_rename,[status(thm)],[t60_enumset1])).
cnf(c_0_28,plain,(k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,plain,(k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X18]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:((((~r2_hidden(X18,X17)|(X18=X10|(X18=X11|(X18=X12|(X18=X13|(X18=X14|(X18=X15|X18=X16)))))))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((X18!=X10|r2_hidden(X18,X17))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((X18!=X11|r2_hidden(X18,X17))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((X18!=X12|r2_hidden(X18,X17))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((X18!=X13|r2_hidden(X18,X17))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((X18!=X14|r2_hidden(X18,X17))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((X18!=X15|r2_hidden(X18,X17))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&((X18!=X16|r2_hidden(X18,X17))|X17!=k5_enumset1(X10,X11,X12,X13,X14,X15,X16)))))))))&((((esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)!=X10|~r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)!=X11|~r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)!=X12|~r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)!=X13|~r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)!=X14|~r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&(((esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)!=X15|~r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))&((esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)!=X16|~r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))))))))&((r2_hidden(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17),X17)|(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)=X10|(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)=X11|(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)=X12|(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)=X13|(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)=X14|(esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)=X15|esk1_8(X10,X11,X12,X13,X14,X15,X16,X17)=X16)))))))|X17=k5_enumset1(X10,X11,X12,X13,X14,X15,X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_enumset1])])])])])])])])).
cnf(c_0_32,plain,(k2_xboole_0(k1_tarski(X1),k2_tarski(X1,X2))=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,plain,(k3_enumset1(X1,X1,X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35,plain,(k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_tarski(X6,X7))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_36, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|r1_tarski(X3,k3_tarski(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
cnf(c_0_37,plain,(k3_tarski(k2_tarski(X1,k3_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38,plain,(k3_xboole_0(X1,k3_tarski(k2_tarski(X1,X2)))=X1), inference(rw,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_39,plain,(r2_hidden(X9,X1)|X1!=k5_enumset1(X2,X3,X4,X5,X6,X7,X8)|X9!=X2), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40,plain,(k3_tarski(k2_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X1),k2_tarski(X1,X2)))=k2_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_29]), c_0_33]), c_0_34])).
cnf(c_0_41,plain,(k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k3_tarski(k2_tarski(k5_enumset1(X1,X1,X1,X2,X3,X4,X5),k2_tarski(X6,X7)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_29]), c_0_34])).
fof(c_0_42, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k2_xboole_0(k4_xboole_0(X3,X4),k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[d6_xboole_0])).
fof(c_0_43, plain, (![X3]:![X4]:k4_xboole_0(k2_xboole_0(X3,X4),X4)=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t40_xboole_1])).
fof(c_0_44, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|X4=k2_xboole_0(X3,k4_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
cnf(c_0_45,plain,(r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46,plain,(k3_tarski(k2_tarski(X1,X1))=X1), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47,plain,(r2_hidden(X1,k5_enumset1(X1,X2,X3,X4,X5,X6,X7))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_39])])).
cnf(c_0_48,plain,(k5_enumset1(X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_49, plain, (![X3]:![X4]:k2_xboole_0(k3_xboole_0(X3,X4),k4_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t51_xboole_1])).
fof(c_0_50, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(((r1_tarski(X4,k2_xboole_0(X5,X6))|~r1_tarski(X4,k5_xboole_0(X5,X6)))&(r1_xboole_0(X4,k3_xboole_0(X5,X6))|~r1_tarski(X4,k5_xboole_0(X5,X6))))&((~r1_tarski(X4,k2_xboole_0(X5,X6))|~r1_xboole_0(X4,k3_xboole_0(X5,X6)))|r1_tarski(X4,k5_xboole_0(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t107_xboole_1])])])])])).
cnf(c_0_51,plain,(k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_52, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|k4_xboole_0(X3,X4)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])])])).
fof(c_0_53, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_54,plain,(k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55,plain,(X1=k2_xboole_0(X2,k4_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56,plain,(r1_tarski(X1,X2)|~r2_hidden(X1,k2_tarski(X2,X2))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57,plain,(r2_hidden(X1,k2_tarski(X1,X2))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_58,plain,(k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59,plain,(r2_hidden(X9,X1)|X1!=k5_enumset1(X2,X3,X4,X5,X6,X7,X8)|X9!=X8), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_60, negated_conjecture, (~(![X1]:![X2]:![X3]:(v1_ordinal1(X3)=>((r2_hidden(X1,X2)&r2_hidden(X2,X3))=>r2_hidden(X1,X3))))), inference(assume_negation,[status(cth)],[t19_ordinal1])).
cnf(c_0_61,plain,(r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_62,plain,(k5_xboole_0(X1,X2)=k3_tarski(k2_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1)))), inference(rw,[status(thm)],[c_0_51, c_0_29])).
cnf(c_0_63,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65,plain,(k4_xboole_0(k3_tarski(k2_tarski(X1,X2)),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_54, c_0_29])).
cnf(c_0_66,plain,(X1=k3_tarski(k2_tarski(X2,k4_xboole_0(X1,X2)))|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[c_0_55, c_0_29])).
cnf(c_0_67,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68,plain,(k3_tarski(k2_tarski(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[c_0_58, c_0_29])).
cnf(c_0_69,plain,(r2_hidden(X1,k5_enumset1(X2,X3,X4,X5,X6,X7,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_59])])).
fof(c_0_70, plain, (![X3]:![X4]:![X3]:((~v1_ordinal1(X3)|(~r2_hidden(X4,X3)|r1_tarski(X4,X3)))&((r2_hidden(esk5_1(X3),X3)|v1_ordinal1(X3))&(~r1_tarski(esk5_1(X3),X3)|v1_ordinal1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])])).
fof(c_0_71, negated_conjecture, ((v1_ordinal1(esk4_0)&((r2_hidden(esk2_0,esk3_0)&r2_hidden(esk3_0,esk4_0))&~r2_hidden(esk2_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_60])])])).
cnf(c_0_72,plain,(r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,k3_tarski(k2_tarski(k4_xboole_0(X2,X3),k4_xboole_0(X3,X2))))), inference(rw,[status(thm)],[c_0_61, c_0_62])).
fof(c_0_73, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(((~r2_hidden(X4,X6)|k4_xboole_0(k2_tarski(X4,X5),X6)!=k1_tarski(X4))&((r2_hidden(X5,X6)|X4=X5)|k4_xboole_0(k2_tarski(X4,X5),X6)!=k1_tarski(X4)))&(((~r2_hidden(X5,X6)|r2_hidden(X4,X6))|k4_xboole_0(k2_tarski(X4,X5),X6)=k1_tarski(X4))&((X4!=X5|r2_hidden(X4,X6))|k4_xboole_0(k2_tarski(X4,X5),X6)=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l38_zfmisc_1])])])])])])).
cnf(c_0_74,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_75,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_76,plain,(k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_67])).
cnf(c_0_77,plain,(k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_38, c_0_37])).
cnf(c_0_78,plain,(r1_tarski(X1,X2)|~r2_hidden(X1,k2_tarski(k3_xboole_0(X2,X3),k4_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_45, c_0_68])).
cnf(c_0_79,plain,(r2_hidden(X1,k2_tarski(X2,X1))), inference(spm,[status(thm)],[c_0_69, c_0_48])).
cnf(c_0_80,plain,(r1_tarski(X1,X2)|~r2_hidden(X1,X2)|~v1_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_81,negated_conjecture,(v1_ordinal1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_82, plain, (![X4]:![X5]:![X6]:(~r1_xboole_0(X4,k4_xboole_0(X5,X6))|r1_xboole_0(X5,k4_xboole_0(X4,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t81_xboole_1])])).
cnf(c_0_83,plain,(r1_xboole_0(k1_xboole_0,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_72, c_0_64])).
cnf(c_0_84,plain,(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)|r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_85,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_64])])).
cnf(c_0_86,plain,(k3_tarski(k2_tarski(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_76]), c_0_77])).
cnf(c_0_87,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_88,negated_conjecture,(r1_tarski(X1,esk4_0)|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_89,plain,(r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X2,k4_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_90,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_83, c_0_38])).
fof(c_0_91, plain, (![X4]:![X5]:![X6]:(~r1_xboole_0(k2_tarski(X4,X5),X6)|~r2_hidden(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_zfmisc_1])])).
cnf(c_0_92,plain,(k4_xboole_0(k2_tarski(X1,X2),X3)=k5_enumset1(X1,X1,X1,X1,X1,X1,X1)|r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_33]), c_0_34])).
cnf(c_0_93,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_85]), c_0_86]), c_0_87])])).
cnf(c_0_94,negated_conjecture,(k4_xboole_0(X1,esk4_0)=k1_xboole_0|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_63, c_0_88])).
cnf(c_0_95,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_74])).
cnf(c_0_96,plain,(~r2_hidden(X1,X2)|~r1_xboole_0(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_97,plain,(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X1)|r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(rw,[status(thm)],[c_0_92, c_0_48])).
cnf(c_0_98,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_93]), c_0_93])).
cnf(c_0_99,negated_conjecture,(r1_xboole_0(X1,k4_xboole_0(X2,esk4_0))|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_94]), c_0_95])])).
cnf(c_0_100,plain,(r2_hidden(X1,X2)|~r1_xboole_0(k4_xboole_0(k2_tarski(X1,X3),X2),X4)|~r2_hidden(X1,X4)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_101,negated_conjecture,(r1_xboole_0(k4_xboole_0(X1,esk4_0),X2)|~r2_hidden(X2,esk4_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_102,negated_conjecture,(r2_hidden(X1,esk4_0)|~r2_hidden(X2,esk4_0)|~r2_hidden(X3,esk4_0)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_103,negated_conjecture,(r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_104,negated_conjecture,(r2_hidden(X1,esk4_0)|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_105,negated_conjecture,(r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_104, c_0_103])).
cnf(c_0_106,negated_conjecture,(r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_107,negated_conjecture,(~r2_hidden(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_108,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 66
# Proof object formula steps           : 43
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 43
# Processed clauses                    : 436
# ...of these trivial                  : 10
# ...subsumed                          : 166
# ...remaining for further processing  : 260
# Other redundant clauses eliminated   : 104
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 21
# Generated clauses                    : 1511
# ...of the previous two non-trivial   : 1125
# Contextual simplify-reflections      : 34
# Paramodulations                      : 1371
# Factorizations                       : 42
# Equation resolutions                 : 105
# Current number of processed clauses  : 178
#    Positive orientable unit clauses  : 44
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 123
# Current number of unprocessed clauses: 714
# ...number of literals in the above   : 2468
# Current number of archived formulas  : 0
# Current number of archived clauses   : 77
# Clause-clause subsumption calls (NU) : 3937
# Rec. Clause-clause subsumption calls : 3465
# Non-unit clause-clause subsumptions  : 135
# Unit Clause-clause subsumption calls : 517
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 113
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 19473

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.008 s
# Total time               : 0.040 s
# Maximum resident set size: 45484 pages
