# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t19_mcart_1, conjecture, (![X1]:![X2]:![X3]:![X4]:![X5]:(r2_hidden(X1,k2_zfmisc_1(k2_tarski(X2,X3),k2_tarski(X4,X5)))=>((k1_xtuple_0(X1)=X2|k1_xtuple_0(X1)=X3)&(k2_xtuple_0(X1)=X4|k2_xtuple_0(X1)=X5)))), file('out/t19_mcart_1__5_premises__rerun.E_input', t19_mcart_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t19_mcart_1__5_premises__rerun.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t19_mcart_1__5_premises__rerun.E_input', t71_enumset1)).
fof(t146_zfmisc_1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X4))=k2_enumset1(k4_tarski(X1,X3),k4_tarski(X1,X4),k4_tarski(X2,X3),k4_tarski(X2,X4))), file('out/t19_mcart_1__5_premises__rerun.E_input', t146_zfmisc_1)).
fof(d2_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(X5=k2_enumset1(X1,X2,X3,X4)<=>![X6]:(r2_hidden(X6,X5)<=>~((~(X6=X1)&(~(X6=X2)&(~(X6=X3)&~(X6=X4)))))))), file('out/t19_mcart_1__5_premises__rerun.E_input', d2_enumset1)).
fof(t7_mcart_1, axiom, (![X1]:![X2]:(k1_xtuple_0(k4_tarski(X1,X2))=X1&k2_xtuple_0(k4_tarski(X1,X2))=X2)), file('out/t19_mcart_1__5_premises__rerun.E_input', t7_mcart_1)).
fof(c_0_6, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:![X5]:(r2_hidden(X1,k2_zfmisc_1(k2_tarski(X2,X3),k2_tarski(X4,X5)))=>((k1_xtuple_0(X1)=X2|k1_xtuple_0(X1)=X3)&(k2_xtuple_0(X1)=X4|k2_xtuple_0(X1)=X5))))), inference(assume_negation,[status(cth)],[t19_mcart_1])).
fof(c_0_7, negated_conjecture, ((r2_hidden(esk1_0,k2_zfmisc_1(k2_tarski(esk2_0,esk3_0),k2_tarski(esk4_0,esk5_0)))&(((k2_xtuple_0(esk1_0)!=esk4_0|k1_xtuple_0(esk1_0)!=esk2_0)&(k2_xtuple_0(esk1_0)!=esk5_0|k1_xtuple_0(esk1_0)!=esk2_0))&((k2_xtuple_0(esk1_0)!=esk4_0|k1_xtuple_0(esk1_0)!=esk3_0)&(k2_xtuple_0(esk1_0)!=esk5_0|k1_xtuple_0(esk1_0)!=esk3_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
fof(c_0_8, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X8]:k2_zfmisc_1(k2_tarski(X5,X6),k2_tarski(X7,X8))=k2_enumset1(k4_tarski(X5,X7),k4_tarski(X5,X8),k4_tarski(X6,X7),k4_tarski(X6,X8))), inference(variable_rename,[status(thm)],[t146_zfmisc_1])).
fof(c_0_11, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:![X12]:![X7]:![X8]:![X9]:![X10]:![X11]:((((~r2_hidden(X12,X11)|(X12=X7|(X12=X8|(X12=X9|X12=X10))))|X11!=k2_enumset1(X7,X8,X9,X10))&(((X12!=X7|r2_hidden(X12,X11))|X11!=k2_enumset1(X7,X8,X9,X10))&(((X12!=X8|r2_hidden(X12,X11))|X11!=k2_enumset1(X7,X8,X9,X10))&(((X12!=X9|r2_hidden(X12,X11))|X11!=k2_enumset1(X7,X8,X9,X10))&((X12!=X10|r2_hidden(X12,X11))|X11!=k2_enumset1(X7,X8,X9,X10))))))&((((esk6_5(X7,X8,X9,X10,X11)!=X7|~r2_hidden(esk6_5(X7,X8,X9,X10,X11),X11))|X11=k2_enumset1(X7,X8,X9,X10))&(((esk6_5(X7,X8,X9,X10,X11)!=X8|~r2_hidden(esk6_5(X7,X8,X9,X10,X11),X11))|X11=k2_enumset1(X7,X8,X9,X10))&(((esk6_5(X7,X8,X9,X10,X11)!=X9|~r2_hidden(esk6_5(X7,X8,X9,X10,X11),X11))|X11=k2_enumset1(X7,X8,X9,X10))&((esk6_5(X7,X8,X9,X10,X11)!=X10|~r2_hidden(esk6_5(X7,X8,X9,X10,X11),X11))|X11=k2_enumset1(X7,X8,X9,X10)))))&((r2_hidden(esk6_5(X7,X8,X9,X10,X11),X11)|(esk6_5(X7,X8,X9,X10,X11)=X7|(esk6_5(X7,X8,X9,X10,X11)=X8|(esk6_5(X7,X8,X9,X10,X11)=X9|esk6_5(X7,X8,X9,X10,X11)=X10))))|X11=k2_enumset1(X7,X8,X9,X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_enumset1])])])])])])])])).
cnf(c_0_12,negated_conjecture,(r2_hidden(esk1_0,k2_zfmisc_1(k2_tarski(esk2_0,esk3_0),k2_tarski(esk4_0,esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X4))=k2_enumset1(k4_tarski(X1,X3),k4_tarski(X1,X4),k4_tarski(X2,X3),k4_tarski(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(X6=X5|X6=X4|X6=X3|X6=X2|X1!=k2_enumset1(X2,X3,X4,X5)|~r2_hidden(X6,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(r2_hidden(esk1_0,k2_zfmisc_1(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k2_enumset1(esk4_0,esk4_0,esk4_0,esk5_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_13]), c_0_14]), c_0_14])).
cnf(c_0_18,plain,(k2_zfmisc_1(k2_enumset1(X1,X1,X1,X2),k2_enumset1(X3,X3,X3,X4))=k2_enumset1(k4_tarski(X1,X3),k4_tarski(X1,X4),k4_tarski(X2,X3),k4_tarski(X2,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_13]), c_0_13]), c_0_14]), c_0_14])).
fof(c_0_19, plain, (![X3]:![X4]:![X3]:![X4]:(k1_xtuple_0(k4_tarski(X3,X4))=X3&k2_xtuple_0(k4_tarski(X3,X4))=X4)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t7_mcart_1])])])).
cnf(c_0_20,plain,(X1=X2|X3=X2|X4=X2|X5=X2|~r2_hidden(X2,k2_enumset1(X1,X3,X4,X5))), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_21,negated_conjecture,(r2_hidden(esk1_0,k2_enumset1(k4_tarski(esk2_0,esk4_0),k4_tarski(esk2_0,esk5_0),k4_tarski(esk3_0,esk4_0),k4_tarski(esk3_0,esk5_0)))), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,plain,(k1_xtuple_0(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(k4_tarski(esk3_0,esk5_0)=esk1_0|k4_tarski(esk3_0,esk4_0)=esk1_0|k4_tarski(esk2_0,esk5_0)=esk1_0|k4_tarski(esk2_0,esk4_0)=esk1_0), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24,plain,(k2_xtuple_0(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,negated_conjecture,(k4_tarski(esk2_0,esk4_0)=esk1_0|k4_tarski(esk2_0,esk5_0)=esk1_0|k4_tarski(esk3_0,esk4_0)=esk1_0|esk3_0=k1_xtuple_0(esk1_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26,negated_conjecture,(k1_xtuple_0(esk1_0)!=esk2_0|k2_xtuple_0(esk1_0)!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27,negated_conjecture,(k4_tarski(esk3_0,esk4_0)=esk1_0|k4_tarski(esk2_0,esk4_0)=esk1_0|esk3_0=k1_xtuple_0(esk1_0)|esk5_0=k2_xtuple_0(esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28,negated_conjecture,(k4_tarski(esk3_0,esk4_0)=esk1_0|k4_tarski(esk2_0,esk4_0)=esk1_0|esk3_0=k1_xtuple_0(esk1_0)|k1_xtuple_0(esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_22, c_0_25])).
cnf(c_0_29,negated_conjecture,(k4_tarski(esk2_0,esk4_0)=esk1_0|k4_tarski(esk3_0,esk4_0)=esk1_0|esk3_0=k1_xtuple_0(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_30,negated_conjecture,(k4_tarski(esk2_0,esk4_0)=esk1_0|esk3_0=k1_xtuple_0(esk1_0)), inference(spm,[status(thm)],[c_0_22, c_0_29])).
cnf(c_0_31,negated_conjecture,(k4_tarski(k1_xtuple_0(esk1_0),esk4_0)=esk1_0|k4_tarski(k1_xtuple_0(esk1_0),esk5_0)=esk1_0|k4_tarski(esk2_0,esk4_0)=esk1_0|k4_tarski(esk2_0,esk5_0)=esk1_0), inference(spm,[status(thm)],[c_0_23, c_0_30])).
cnf(c_0_32,negated_conjecture,(k4_tarski(k1_xtuple_0(esk1_0),esk4_0)=esk1_0|k4_tarski(esk2_0,esk5_0)=esk1_0|k4_tarski(esk2_0,esk4_0)=esk1_0|esk5_0=k2_xtuple_0(esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_31])).
cnf(c_0_33,negated_conjecture,(k1_xtuple_0(esk1_0)!=esk3_0|k2_xtuple_0(esk1_0)!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34,negated_conjecture,(k4_tarski(k1_xtuple_0(esk1_0),esk4_0)=esk1_0|k4_tarski(esk2_0,esk4_0)=esk1_0|esk5_0=k2_xtuple_0(esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_32])).
cnf(c_0_35,negated_conjecture,(k1_xtuple_0(esk1_0)!=esk3_0|k2_xtuple_0(esk1_0)!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_36,negated_conjecture,(k4_tarski(k1_xtuple_0(esk1_0),esk4_0)=esk1_0|k4_tarski(esk2_0,esk4_0)=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_30])).
cnf(c_0_37,negated_conjecture,(k4_tarski(esk2_0,esk4_0)=esk1_0|k2_xtuple_0(esk1_0)!=esk4_0), inference(spm,[status(thm)],[c_0_35, c_0_30])).
cnf(c_0_38,negated_conjecture,(k4_tarski(esk2_0,esk4_0)=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_36]), c_0_37])).
cnf(c_0_39,negated_conjecture,(k1_xtuple_0(esk1_0)!=esk2_0|k2_xtuple_0(esk1_0)!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_40,negated_conjecture,(k1_xtuple_0(esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_22, c_0_38])).
cnf(c_0_41,negated_conjecture,(k2_xtuple_0(esk1_0)!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40])])).
cnf(c_0_42,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_38]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 30
# Proof object formula steps           : 13
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 18
# Processed clauses                    : 62
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 61
# Other redundant clauses eliminated   : 88
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 10
# Generated clauses                    : 223
# ...of the previous two non-trivial   : 130
# Contextual simplify-reflections      : 3
# Paramodulations                      : 74
# Factorizations                       : 56
# Equation resolutions                 : 93
# Current number of processed clauses  : 22
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 91
# ...number of literals in the above   : 725
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 136
# Rec. Clause-clause subsumption calls : 107
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4881

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45276 pages
