# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t14_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)))), file('out/t19_yellow_1__11_premises__rerun.E_input', t14_yellow_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t19_yellow_1__11_premises__rerun.E_input', d1_xboole_0)).
fof(t99_zfmisc_1, axiom, (![X1]:k3_tarski(k1_zfmisc_1(X1))=X1), file('out/t19_yellow_1__11_premises__rerun.E_input', t99_zfmisc_1)).
fof(redefinition_k9_setfam_1, axiom, (![X1]:k9_setfam_1(X1)=k1_zfmisc_1(X1)), file('out/t19_yellow_1__11_premises__rerun.E_input', redefinition_k9_setfam_1)).
fof(l1_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t19_yellow_1__11_premises__rerun.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t19_yellow_1__11_premises__rerun.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t19_yellow_1__11_premises__rerun.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t19_yellow_1__11_premises__rerun.E_input', t71_enumset1)).
fof(t100_zfmisc_1, axiom, (![X1]:r1_tarski(X1,k1_zfmisc_1(k3_tarski(X1)))), file('out/t19_yellow_1__11_premises__rerun.E_input', t100_zfmisc_1)).
fof(t31_zfmisc_1, axiom, (![X1]:k3_tarski(k1_tarski(X1))=X1), file('out/t19_yellow_1__11_premises__rerun.E_input', t31_zfmisc_1)).
fof(t19_yellow_1, conjecture, (![X1]:k4_yellow_0(k3_yellow_1(X1))=X1), file('out/t19_yellow_1__11_premises__rerun.E_input', t19_yellow_1)).
fof(t4_yellow_1, axiom, (![X1]:k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1))), file('out/t19_yellow_1__11_premises__rerun.E_input', t4_yellow_1)).
fof(c_0_12, plain, (![X2]:(v1_xboole_0(X2)|(~r2_hidden(k3_tarski(X2),X2)|k4_yellow_0(k2_yellow_1(X2))=k3_tarski(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
fof(c_0_13, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk1_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
fof(c_0_14, plain, (![X2]:k3_tarski(k1_zfmisc_1(X2))=X2), inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
fof(c_0_15, plain, (![X2]:k9_setfam_1(X2)=k1_zfmisc_1(X2)), inference(variable_rename,[status(thm)],[redefinition_k9_setfam_1])).
fof(c_0_16, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])])])).
fof(c_0_17, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_18, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_20, plain, (![X2]:r1_tarski(X2,k1_zfmisc_1(k3_tarski(X2)))), inference(variable_rename,[status(thm)],[t100_zfmisc_1])).
fof(c_0_21, plain, (![X2]:k3_tarski(k1_tarski(X2))=X2), inference(variable_rename,[status(thm)],[t31_zfmisc_1])).
fof(c_0_22, negated_conjecture, (~(![X1]:k4_yellow_0(k3_yellow_1(X1))=X1)), inference(assume_negation,[status(cth)],[t19_yellow_1])).
cnf(c_0_23,plain,(k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|v1_xboole_0(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,plain,(k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26,plain,(k9_setfam_1(X1)=k1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27,plain,(r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31,plain,(r1_tarski(X1,k1_zfmisc_1(k3_tarski(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32,plain,(k3_tarski(k1_tarski(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_33, negated_conjecture, (k4_yellow_0(k3_yellow_1(esk2_0))!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
fof(c_0_34, plain, (![X2]:k3_yellow_1(X2)=k2_yellow_1(k9_setfam_1(X2))), inference(variable_rename,[status(thm)],[t4_yellow_1])).
cnf(c_0_35,plain,(k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_36,plain,(k3_tarski(k9_setfam_1(X1))=X1), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_37,plain,(r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_38,plain,(r1_tarski(X1,k9_setfam_1(k3_tarski(X1)))), inference(rw,[status(thm)],[c_0_31, c_0_26])).
cnf(c_0_39,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_40,negated_conjecture,(k4_yellow_0(k3_yellow_1(esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41,plain,(k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42,plain,(k4_yellow_0(k2_yellow_1(k9_setfam_1(X1)))=X1|~r2_hidden(X1,k9_setfam_1(X1))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43,plain,(r2_hidden(X1,k9_setfam_1(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_44,negated_conjecture,(k4_yellow_0(k2_yellow_1(k9_setfam_1(esk2_0)))!=esk2_0), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45,plain,(k4_yellow_0(k2_yellow_1(k9_setfam_1(X1)))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43])])).
cnf(c_0_46,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 22
# Proof object formula steps           : 25
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 9
# Processed clauses                    : 23
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 23
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 8
# ...of the previous two non-trivial   : 7
# Contextual simplify-reflections      : 1
# Paramodulations                      : 8
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 12
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 6
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 2
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 1
# Rec. Clause-clause subsumption calls : 1
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 715

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45512 pages
