# No SInE strategy applied
# Auto-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t19_yellow_6, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(r1_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X3))))), file('out/t19_yellow_6__4_premises__rerun.E_input', t19_yellow_6)).
fof(d11_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))))), file('out/t19_yellow_6__4_premises__rerun.E_input', d11_waybel_0)).
fof(dt_l1_waybel_0, axiom, (![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2)))), file('out/t19_yellow_6__4_premises__rerun.E_input', dt_l1_waybel_0)).
fof(d3_yellow_6, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(v7_waybel_0(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r1_orders_2(X1,X2,X4)&r1_orders_2(X1,X3,X4)))))))), file('out/t19_yellow_6__4_premises__rerun.E_input', d3_yellow_6)).
fof(d12_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))))), file('out/t19_yellow_6__4_premises__rerun.E_input', d12_waybel_0)).
fof(c_0_5, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(r1_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t19_yellow_6])).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X10]:![X8]:![X11]:(((((m1_subset_1(esk4_3(X6,X7,X8),u1_struct_0(X7))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~m1_subset_1(X10,u1_struct_0(X7))|(~r1_orders_2(X7,esk4_3(X6,X7,X8),X10)|r2_hidden(k2_waybel_0(X6,X7,X10),X8)))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))&(((((m1_subset_1(esk5_4(X6,X7,X8,X11),u1_struct_0(X7))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X11,esk5_4(X6,X7,X8,X11))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~r2_hidden(k2_waybel_0(X6,X7,esk5_4(X6,X7,X8,X11)),X8)|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_waybel_0])])])])])])])])).
fof(c_0_7, negated_conjecture, (((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&(r1_waybel_0(esk1_0,esk2_0,esk3_0)&~r2_waybel_0(esk1_0,esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])).
fof(c_0_8, plain, (![X3]:![X4]:(~l1_struct_0(X3)|(~l1_waybel_0(X4,X3)|l1_orders_2(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])])])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X11]:((((((m1_subset_1(esk6_3(X5,X6,X7),u1_struct_0(X5))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|~v7_waybel_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))&(((((r1_orders_2(X5,X6,esk6_3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|~v7_waybel_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))&((((r1_orders_2(X5,X7,esk6_3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|~v7_waybel_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))))&(((m1_subset_1(esk7_1(X5),u1_struct_0(X5))|v7_waybel_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))&(((m1_subset_1(esk8_1(X5),u1_struct_0(X5))|v7_waybel_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))&(((~m1_subset_1(X11,u1_struct_0(X5))|(~r1_orders_2(X5,esk7_1(X5),X11)|~r1_orders_2(X5,esk8_1(X5),X11)))|v7_waybel_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_yellow_6])])])])])])])])).
cnf(c_0_10,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,negated_conjecture,(l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16,plain,(l1_orders_2(X1)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_17, plain, (![X6]:![X7]:![X8]:![X9]:![X8]:![X12]:((((((m1_subset_1(esk9_4(X6,X7,X8,X9),u1_struct_0(X7))|~m1_subset_1(X9,u1_struct_0(X7)))|~r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X9,esk9_4(X6,X7,X8,X9))|~m1_subset_1(X9,u1_struct_0(X7)))|~r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((r2_hidden(k2_waybel_0(X6,X7,esk9_4(X6,X7,X8,X9)),X8)|~m1_subset_1(X9,u1_struct_0(X7)))|~r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))))&((((m1_subset_1(esk10_3(X6,X7,X8),u1_struct_0(X7))|r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~m1_subset_1(X12,u1_struct_0(X7))|(~r1_orders_2(X7,esk10_3(X6,X7,X8),X12)|~r2_hidden(k2_waybel_0(X6,X7,X12),X8)))|r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])])).
cnf(c_0_18,plain,(v2_struct_0(X1)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v7_waybel_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19,negated_conjecture,(m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])]), c_0_14]), c_0_15])).
cnf(c_0_20,negated_conjecture,(l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_12]), c_0_13])])).
cnf(c_0_21,negated_conjecture,(v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_waybel_0(X1,X2,X3)|m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk6_3(esk2_0,X1,esk4_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])]), c_0_14])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk10_3(esk1_0,esk2_0,X1),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_12]), c_0_13])]), c_0_14]), c_0_15])).
cnf(c_0_25,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(k2_waybel_0(X1,X2,X4),X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)|~r1_orders_2(X2,esk4_3(X1,X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26,negated_conjecture,(m1_subset_1(esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27,negated_conjecture,(r2_hidden(k2_waybel_0(X1,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X2),esk4_3(esk1_0,esk2_0,esk3_0))),X3)|r2_waybel_0(esk1_0,esk2_0,X2)|v2_struct_0(X1)|~r1_orders_2(esk2_0,esk4_3(X1,esk2_0,X3),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X2),esk4_3(esk1_0,esk2_0,esk3_0)))|~r1_waybel_0(X1,esk2_0,X3)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_14])).
cnf(c_0_28,plain,(v2_struct_0(X1)|r1_orders_2(X1,X3,esk6_3(X1,X2,X3))|~l1_orders_2(X1)|~v7_waybel_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_hidden(k2_waybel_0(X1,X2,X4),X3)|~r1_orders_2(X2,esk10_3(X1,X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),X2)|r2_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk4_3(esk1_0,esk2_0,X2),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|~r1_waybel_0(esk1_0,esk2_0,X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_12]), c_0_13])]), c_0_15])).
cnf(c_0_31,negated_conjecture,(r1_orders_2(esk2_0,esk4_3(esk1_0,esk2_0,esk3_0),esk6_3(esk2_0,X1,esk4_3(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_19]), c_0_20]), c_0_21])]), c_0_14])).
cnf(c_0_32,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(X2,esk2_0,X3)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X2,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),X3)|~r1_orders_2(esk2_0,esk10_3(X2,esk2_0,X3),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|~l1_waybel_0(esk2_0,X2)|~l1_struct_0(X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_26]), c_0_14])).
cnf(c_0_33,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),esk3_0)|r2_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk4_3(esk1_0,esk2_0,esk3_0),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_30, c_0_11])).
cnf(c_0_34,negated_conjecture,(r1_orders_2(esk2_0,esk4_3(esk1_0,esk2_0,esk3_0),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_31, c_0_24])).
cnf(c_0_35,plain,(v2_struct_0(X1)|r1_orders_2(X1,X2,esk6_3(X1,X2,X3))|~l1_orders_2(X1)|~v7_waybel_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(esk1_0,esk2_0,X2)|~r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),X2)|~r1_orders_2(esk2_0,esk10_3(esk1_0,esk2_0,X2),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_12]), c_0_13])]), c_0_15])).
cnf(c_0_37,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),esk3_0)|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38,negated_conjecture,(~r2_waybel_0(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_39,negated_conjecture,(r1_orders_2(esk2_0,X1,esk6_3(esk2_0,X1,esk4_3(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_19]), c_0_20]), c_0_21])]), c_0_14])).
cnf(c_0_40,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk10_3(esk1_0,esk2_0,esk3_0),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_41,negated_conjecture,(r1_orders_2(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_39, c_0_24])).
cnf(c_0_42,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 32
# Proof object formula steps           : 11
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 982
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 982
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 47582
# ...of the previous two non-trivial   : 47126
# Contextual simplify-reflections      : 0
# Paramodulations                      : 47582
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 980
#    Positive orientable unit clauses  : 43
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 933
# Current number of unprocessed clauses: 46169
# ...number of literals in the above   : 182920
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 54870
# Rec. Clause-clause subsumption calls : 8047
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 1442
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 743
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2814461

# -------------------------------------------------
# User time                : 0.568 s
# System time              : 0.024 s
# Total time               : 0.592 s
# Maximum resident set size: 47440 pages
