# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_lattice3, conjecture, (![X1]:![X2]:(m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))=>(k1_lattices(k1_lattice3(X1),X2,X3)=k2_xboole_0(X2,X3)&k2_lattices(k1_lattice3(X1),X2,X3)=k3_xboole_0(X2,X3))))), file('out/t1_lattice3__2_premises__rerun.E_input', t1_lattice3)).
fof(ie2_lattice3, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:((m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))&m1_subset_1(X3,u1_struct_0(k1_lattice3(X1))))=>((X2=X4&X3=X5)=>k2_lattices(k1_lattice3(X1),X2,X3)=k3_xboole_0(X4,X5)))), file('out/t1_lattice3__2_premises__rerun.E_input', ie2_lattice3)).
fof(ie1_lattice3, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:((m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))&m1_subset_1(X3,u1_struct_0(k1_lattice3(X1))))=>((X2=X4&X3=X5)=>k1_lattices(k1_lattice3(X1),X2,X3)=k2_xboole_0(X4,X5)))), file('out/t1_lattice3__2_premises__rerun.E_input', ie1_lattice3)).
fof(c_0_3, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))=>(k1_lattices(k1_lattice3(X1),X2,X3)=k2_xboole_0(X2,X3)&k2_lattices(k1_lattice3(X1),X2,X3)=k3_xboole_0(X2,X3)))))), inference(assume_negation,[status(cth)],[t1_lattice3])).
fof(c_0_4, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((~m1_subset_1(X7,u1_struct_0(k1_lattice3(X6)))|~m1_subset_1(X8,u1_struct_0(k1_lattice3(X6))))|((X7!=X9|X8!=X10)|k2_lattices(k1_lattice3(X6),X7,X8)=k3_xboole_0(X9,X10)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ie2_lattice3])])])])).
fof(c_0_5, negated_conjecture, ((m1_subset_1(esk2_0,u1_struct_0(k1_lattice3(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k1_lattice3(esk1_0)))&(k1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k2_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_6,plain,(k2_lattices(k1_lattice3(X1),X2,X3)=k3_xboole_0(X4,X5)|X3!=X5|X2!=X4|~m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((~m1_subset_1(X7,u1_struct_0(k1_lattice3(X6)))|~m1_subset_1(X8,u1_struct_0(k1_lattice3(X6))))|((X7!=X9|X8!=X10)|k1_lattices(k1_lattice3(X6),X7,X8)=k2_xboole_0(X9,X10)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ie1_lattice3])])])])).
cnf(c_0_8,negated_conjecture,(k2_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)|k1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(k2_lattices(k1_lattice3(X1),X2,X3)=k3_xboole_0(X2,X3)|~m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_6])])).
cnf(c_0_10,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(k1_lattice3(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(k1_lattice3(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,plain,(k1_lattices(k1_lattice3(X1),X2,X3)=k2_xboole_0(X4,X5)|X3!=X5|X2!=X4|~m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,negated_conjecture,(k1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])])).
cnf(c_0_14,plain,(k1_lattices(k1_lattice3(X1),X2,X3)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_12])])).
cnf(c_0_15,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_10]), c_0_11])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 16
# Proof object clause steps            : 9
# Proof object formula steps           : 7
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 13
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 13
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 4
# ...of the previous two non-trivial   : 3
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 6
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 643

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45272 pages
