# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_tarski, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3)))))))), file('out/t1_mcart_1__8_premises__rerun.E_input', t7_tarski)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t1_mcart_1__8_premises__rerun.E_input', t3_xboole_0)).
fof(t1_mcart_1, conjecture, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&r1_xboole_0(X2,X1)))))), file('out/t1_mcart_1__8_premises__rerun.E_input', t1_mcart_1)).
fof(t7_ordinal1, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1)))), file('out/t1_mcart_1__8_premises__rerun.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t1_mcart_1__8_premises__rerun.E_input', t2_xboole_1)).
fof(dt_k6_relat_1, axiom, (![X1]:v1_relat_1(k6_relat_1(X1))), file('out/t1_mcart_1__8_premises__rerun.E_input', dt_k6_relat_1)).
fof(d13_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:![X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X5,X4),X1)&r2_hidden(X5,X2)))))), file('out/t1_mcart_1__8_premises__rerun.E_input', d13_relat_1)).
fof(t150_relat_1, axiom, (![X1]:k9_relat_1(k1_xboole_0,X1)=k1_xboole_0), file('out/t1_mcart_1__8_premises__rerun.E_input', t150_relat_1)).
fof(t81_relat_1, axiom, (k6_relat_1(k1_xboole_0)=k1_xboole_0), file('out/t1_mcart_1__8_premises__rerun.E_input', t81_relat_1)).
fof(c_0_9, plain, (![X5]:![X6]:![X8]:((r2_hidden(esk1_2(X5,X6),X6)|~r2_hidden(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,esk1_2(X5,X6)))|~r2_hidden(X5,X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])])])).
fof(c_0_10, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk2_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk2_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_11, negated_conjecture, (~(![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&r1_xboole_0(X2,X1))))))), inference(assume_negation,[status(cth)],[t1_mcart_1])).
cnf(c_0_12,plain,(~r2_hidden(X1,X2)|~r2_hidden(X3,esk1_2(X1,X2))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~r1_tarski(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_15, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_16, plain, (![X2]:v1_relat_1(k6_relat_1(X2))), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
fof(c_0_17, negated_conjecture, (![X4]:(esk3_0!=k1_xboole_0&(~r2_hidden(X4,esk3_0)|~r1_xboole_0(X4,esk3_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])])).
cnf(c_0_18,plain,(r1_xboole_0(esk1_2(X1,X2),X3)|~r2_hidden(esk2_2(esk1_2(X1,X2),X3),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20,plain,(~r1_tarski(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_22, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:![X7]:![X8]:![X13]:((((((r2_hidden(k4_tarski(esk4_4(X6,X7,X8,X9),X9),X6)|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk4_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(k4_tarski(X11,X9),X6)|~r2_hidden(X11,X7))|r2_hidden(X9,X8))|X8!=k9_relat_1(X6,X7))|~v1_relat_1(X6)))&((((~r2_hidden(esk5_3(X6,X7,X8),X8)|(~r2_hidden(k4_tarski(X13,esk5_3(X6,X7,X8)),X6)|~r2_hidden(X13,X7)))|X8=k9_relat_1(X6,X7))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk6_3(X6,X7,X8),esk5_3(X6,X7,X8)),X6)|r2_hidden(esk5_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|~v1_relat_1(X6))&(((r2_hidden(esk6_3(X6,X7,X8),X7)|r2_hidden(esk5_3(X6,X7,X8),X8))|X8=k9_relat_1(X6,X7))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_relat_1])])])])])])])).
fof(c_0_23, plain, (![X2]:k9_relat_1(k1_xboole_0,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t150_relat_1])).
cnf(c_0_24,plain,(v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,plain,(k6_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t81_relat_1])).
cnf(c_0_26,negated_conjecture,(~r1_xboole_0(X1,esk3_0)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,plain,(r1_xboole_0(esk1_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_28,plain,(r2_hidden(esk1_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29,plain,(~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30,plain,(X2=k9_relat_1(X1,X3)|r2_hidden(esk5_3(X1,X3,X2),X2)|r2_hidden(k4_tarski(esk6_3(X1,X3,X2),esk5_3(X1,X3,X2)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31,plain,(k9_relat_1(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,negated_conjecture,(~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_34,plain,(X1=k1_xboole_0|r2_hidden(esk5_3(k1_xboole_0,X2,X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_35,negated_conjecture,(esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_36,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 19
# Proof object formula steps           : 18
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 95
# ...of these trivial                  : 0
# ...subsumed                          : 30
# ...remaining for further processing  : 65
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 97
# ...of the previous two non-trivial   : 85
# Contextual simplify-reflections      : 7
# Paramodulations                      : 94
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 99
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 353
# Rec. Clause-clause subsumption calls : 243
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2453

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 47184 pages
