# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', d3_tarski)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t7_xboole_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', d1_xboole_0)).
fof(t45_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1)))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t45_xboole_1)).
fof(t36_xboole_1, axiom, (![X1]:![X2]:r1_tarski(k4_xboole_0(X1,X2),X1)), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t36_xboole_1)).
fof(l32_xboole_1, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', l32_xboole_1)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', l13_xboole_0)).
fof(symmetry_r1_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', symmetry_r1_xboole_0)).
fof(t79_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2)), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t79_xboole_1)).
fof(t1_boole, axiom, (![X1]:k2_xboole_0(X1,k1_xboole_0)=X1), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t1_boole)).
fof(t1_tex_2, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2)))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t1_tex_2)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', d10_xboole_0)).
fof(t83_xboole_1, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1)), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t83_xboole_1)).
fof(t63_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t63_xboole_1)).
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t1_xboole_1)).
fof(t43_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t43_xboole_1)).
fof(t69_xboole_1, axiom, (![X1]:![X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1))))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t69_xboole_1)).
fof(redefinition_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', redefinition_k6_domain_1)).
fof(d1_tex_2, axiom, (![X1]:(~(v1_xboole_0(X1))=>(v1_zfmisc_1(X1)<=>?[X2]:(m1_subset_1(X2,X1)&X1=k6_domain_1(X1,X2))))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', d1_tex_2)).
fof(t44_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~(X2=X3)&(~(X2=k1_xboole_0)&~(X3=k1_xboole_0)))))), file('out/t1_tex_2__20_premises__rerun__rerun.E_input', t44_zfmisc_1)).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk1_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk1_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_21, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_22, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk2_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_23,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,plain,(r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_27, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|X4=k2_xboole_0(X3,k4_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
cnf(c_0_28,plain,(~v1_xboole_0(k2_xboole_0(X1,X2))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29,plain,(X1=k2_xboole_0(X2,k4_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_31,plain,(v1_xboole_0(X1)|r2_hidden(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_32, plain, (![X3]:![X4]:r1_tarski(k4_xboole_0(X3,X4),X3)), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_33, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|k4_xboole_0(X3,X4)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])])])).
fof(c_0_34, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_35,plain,(v1_xboole_0(X1)|~v1_xboole_0(X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_38, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|r1_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_39, plain, (![X3]:![X4]:r1_xboole_0(k4_xboole_0(X3,X4),X4)), inference(variable_rename,[status(thm)],[t79_xboole_1])).
cnf(c_0_40,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41,plain,(v1_xboole_0(k4_xboole_0(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_42, plain, (![X2]:k2_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_43, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))))), inference(assume_negation,[status(cth)],[t1_tex_2])).
fof(c_0_44, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_45,plain,(r1_tarski(k1_xboole_0,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_46, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k4_xboole_0(X3,X4)=X3)&(k4_xboole_0(X3,X4)!=X3|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])])])).
cnf(c_0_47,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48,plain,(r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_51, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_xboole_0(X5,X6))|r1_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
fof(c_0_52, negated_conjecture, ((~v1_xboole_0(esk3_0)&((~v1_xboole_0(esk4_0)&v1_zfmisc_1(esk4_0))&(r1_tarski(esk3_0,esk4_0)&esk3_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_43])])])])])])).
cnf(c_0_53,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_45, c_0_24])).
cnf(c_0_55,plain,(k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56,plain,(r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_57,plain,(X1=X2|~v1_xboole_0(X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_49]), c_0_50])).
fof(c_0_58, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_59, plain, (![X4]:![X5]:![X6]:(~r1_tarski(X4,k2_xboole_0(X5,X6))|r1_tarski(k4_xboole_0(X4,X5),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_60,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_61, plain, (![X3]:![X4]:(v1_xboole_0(X4)|(~r1_tarski(X4,X3)|~r1_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
cnf(c_0_62,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_63,negated_conjecture,(r1_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64,plain,(X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
fof(c_0_65, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|k6_domain_1(X3,X4)=k1_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_66, plain, (![X3]:![X5]:((((m1_subset_1(esk5_1(X3),X3)|~v1_zfmisc_1(X3))|v1_xboole_0(X3))&((X3=k6_domain_1(X3,esk5_1(X3))|~v1_zfmisc_1(X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X5,X3)|X3!=k6_domain_1(X3,X5))|v1_zfmisc_1(X3))|v1_xboole_0(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_2])])])])])])])])).
cnf(c_0_67,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_68,plain,(k4_xboole_0(X1,X2)=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_57, c_0_36])).
cnf(c_0_69,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_70,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_71,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_72,plain,(v1_xboole_0(X1)|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_73,negated_conjecture,(r1_xboole_0(esk3_0,X1)|~r1_xboole_0(esk4_0,X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_74,negated_conjecture,(~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_75,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_36])).
fof(c_0_76, plain, (![X4]:![X5]:![X6]:(k1_tarski(X4)!=k2_xboole_0(X5,X6)|(X5=X6|(X5=k1_xboole_0|X6=k1_xboole_0)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t44_zfmisc_1])])])).
cnf(c_0_77,plain,(k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_78,plain,(v1_xboole_0(X1)|m1_subset_1(esk5_1(X1),X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_79,plain,(k4_xboole_0(X1,X2)=X1|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_80,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_55, c_0_48])).
cnf(c_0_81,negated_conjecture,(r1_tarski(X1,esk4_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_69, c_0_63])).
cnf(c_0_82,plain,(r1_tarski(k4_xboole_0(k2_xboole_0(X1,X2),X1),X2)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_83,negated_conjecture,(~r1_xboole_0(esk4_0,X1)|~r1_tarski(esk3_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74])).
cnf(c_0_84,plain,(r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_85,negated_conjecture,(k4_xboole_0(esk3_0,X1)=esk3_0|~r1_xboole_0(esk4_0,X1)), inference(spm,[status(thm)],[c_0_55, c_0_73])).
cnf(c_0_86,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_48, c_0_75])).
cnf(c_0_87,plain,(X1=k1_xboole_0|X2=k1_xboole_0|X2=X1|k1_tarski(X3)!=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_88,plain,(k6_domain_1(X1,esk5_1(X1))=k1_tarski(esk5_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_89,plain,(X1=k1_xboole_0|~v1_xboole_0(X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_37, c_0_79])).
cnf(c_0_90,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_37, c_0_80])).
cnf(c_0_91,negated_conjecture,(r1_tarski(k4_xboole_0(k2_xboole_0(X1,esk3_0),X1),esk4_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_92,negated_conjecture,(k4_xboole_0(esk4_0,X1)!=esk4_0|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_93,plain,(r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_94,negated_conjecture,(k4_xboole_0(esk3_0,X1)=esk3_0|k4_xboole_0(esk4_0,X1)!=esk4_0), inference(spm,[status(thm)],[c_0_85, c_0_84])).
cnf(c_0_95,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_47, c_0_86])).
cnf(c_0_96,plain,(X1=k1_xboole_0|X2=k1_xboole_0|X1=X2|v1_xboole_0(X3)|k6_domain_1(X3,esk5_1(X3))!=k2_xboole_0(X2,X1)|~v1_zfmisc_1(X3)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_97,plain,(v1_xboole_0(X1)|X1=k6_domain_1(X1,esk5_1(X1))|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_98,plain,(X1=k1_xboole_0|~v1_xboole_0(k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_89, c_0_24])).
cnf(c_0_99,negated_conjecture,(k4_xboole_0(k2_xboole_0(esk4_0,esk3_0),esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_100,negated_conjecture,(esk4_0!=k1_xboole_0|~r1_tarski(esk3_0,X1)|~r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_92, c_0_37])).
cnf(c_0_101,negated_conjecture,(k4_xboole_0(esk4_0,X1)!=esk4_0|esk3_0!=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_92])).
cnf(c_0_102,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_55, c_0_95])).
cnf(c_0_103,plain,(X1=k1_xboole_0|X2=k1_xboole_0|X2=X1|~v1_zfmisc_1(k2_xboole_0(X1,X2))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97])]), c_0_98])).
cnf(c_0_104,negated_conjecture,(k2_xboole_0(esk4_0,esk3_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_99]), c_0_50]), c_0_24])])).
cnf(c_0_105,negated_conjecture,(v1_zfmisc_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_106,negated_conjecture,(esk4_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_63]), c_0_71])])).
cnf(c_0_107,negated_conjecture,(esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_108,negated_conjecture,(esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_109,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105])]), c_0_106]), c_0_107]), c_0_108]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 69
# Proof object formula steps           : 41
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 33
# Processed clauses                    : 18849
# ...of these trivial                  : 75
# ...subsumed                          : 17330
# ...remaining for further processing  : 1444
# Other redundant clauses eliminated   : 1807
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 85
# Backward-rewritten                   : 77
# Generated clauses                    : 135897
# ...of the previous two non-trivial   : 111278
# Contextual simplify-reflections      : 5445
# Paramodulations                      : 134070
# Factorizations                       : 20
# Equation resolutions                 : 1807
# Current number of processed clauses  : 1248
#    Positive orientable unit clauses  : 117
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 54
#    Non-unit-clauses                  : 1077
# Current number of unprocessed clauses: 84087
# ...number of literals in the above   : 305244
# Current number of archived formulas  : 0
# Current number of archived clauses   : 194
# Clause-clause subsumption calls (NU) : 1314904
# Rec. Clause-clause subsumption calls : 803277
# Non-unit clause-clause subsumptions  : 12883
# Unit Clause-clause subsumption calls : 18639
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1469
# BW rewrite match successes           : 39
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1793459

# -------------------------------------------------
# User time                : 1.996 s
# System time              : 0.064 s
# Total time               : 2.060 s
# Maximum resident set size: 70876 pages
