# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_tex_2, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(v3_card_1(X2,1)=>(r1_tarski(X1,X2)=>X1=X2)))), file('out/t1_tex_2__111_premises__rerun__rerun.E_input', t1_tex_2)).
fof(t33_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2)))), file('out/t1_tex_2__111_premises__rerun__rerun.E_input', t33_zfmisc_1)).
fof(t46_subset_1, axiom, (![X1]:((~(v1_xboole_0(X1))&v1_zfmisc_1(X1))=>?[X2]:(m1_subset_1(X2,X1)&X1=k1_tarski(X2)))), file('out/t1_tex_2__111_premises__rerun__rerun.E_input', t46_subset_1)).
fof(cc9_card_1, axiom, (![X1]:(v3_card_1(X1,1)=>(~(v1_xboole_0(X1))&v1_zfmisc_1(X1)))), file('out/t1_tex_2__111_premises__rerun__rerun.E_input', cc9_card_1)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t1_tex_2__111_premises__rerun__rerun.E_input', fc1_xboole_0)).
fof(c_0_5, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(v3_card_1(X2,1)=>(r1_tarski(X1,X2)=>X1=X2))))), inference(assume_negation,[status(cth)],[t1_tex_2])).
fof(c_0_6, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(X3,k1_tarski(X4))|(X3=k1_xboole_0|X3=k1_tarski(X4)))&((X3!=k1_xboole_0|r1_tarski(X3,k1_tarski(X4)))&(X3!=k1_tarski(X4)|r1_tarski(X3,k1_tarski(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_zfmisc_1])])])])])).
fof(c_0_7, plain, (![X3]:((m1_subset_1(esk3_1(X3),X3)|(v1_xboole_0(X3)|~v1_zfmisc_1(X3)))&(X3=k1_tarski(esk3_1(X3))|(v1_xboole_0(X3)|~v1_zfmisc_1(X3))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_subset_1])])])])])])])).
fof(c_0_8, plain, (![X2]:((~v1_xboole_0(X2)|~v3_card_1(X2,1))&(v1_zfmisc_1(X2)|~v3_card_1(X2,1)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc9_card_1])])])])).
fof(c_0_9, negated_conjecture, ((~v1_xboole_0(esk1_0)&(v3_card_1(esk2_0,1)&(r1_tarski(esk1_0,esk2_0)&esk1_0!=esk2_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_10,plain,(X1=k1_tarski(X2)|X1=k1_xboole_0|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(v1_xboole_0(X1)|X1=k1_tarski(esk3_1(X1))|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(v1_zfmisc_1(X1)|~v3_card_1(X1,1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,negated_conjecture,(v3_card_1(esk2_0,1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(~v3_card_1(X1,1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,plain,(X1=k1_xboole_0|X1=X2|v1_xboole_0(X2)|~v1_zfmisc_1(X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_16,negated_conjecture,(r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,negated_conjecture,(v1_zfmisc_1(esk2_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18,negated_conjecture,(esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_14, c_0_13])).
cnf(c_0_20,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_21,negated_conjecture,(k1_xboole_0=esk1_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])]), c_0_18]), c_0_19])).
cnf(c_0_22,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23,plain,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 24
# Proof object clause steps            : 14
# Proof object formula steps           : 10
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 34
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 34
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 10
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 0
# Paramodulations                      : 8
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 14
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 12
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 3
# Rec. Clause-clause subsumption calls : 1
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 720

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45252 pages
