# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t201_funct_2, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>(![X3]:(m1_subset_1(X3,X1)=>k3_funct_2(X1,X1,X2,X3)=X3)=>r2_funct_2(X1,X1,X2,k6_partfun1(X1)))))), file('out/t201_funct_2__9_premises__rerun.E_input', t201_funct_2)).
fof(t34_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k6_relat_1(X1)<=>(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=X3))))), file('out/t201_funct_2__9_premises__rerun.E_input', t34_funct_1)).
fof(redefinition_k6_partfun1, axiom, (![X1]:k6_partfun1(X1)=k6_relat_1(X1)), file('out/t201_funct_2__9_premises__rerun.E_input', redefinition_k6_partfun1)).
fof(cc5_funct_2, axiom, (![X1]:![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1)))))), file('out/t201_funct_2__9_premises__rerun.E_input', cc5_funct_2)).
fof(cc2_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2)))), file('out/t201_funct_2__9_premises__rerun.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('out/t201_funct_2__9_premises__rerun.E_input', cc1_relset_1)).
fof(d4_partfun1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>(v1_partfun1(X2,X1)<=>k1_relat_1(X2)=X1))), file('out/t201_funct_2__9_premises__rerun.E_input', d4_partfun1)).
fof(redefinition_k3_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4))), file('out/t201_funct_2__9_premises__rerun.E_input', redefinition_k3_funct_2)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t201_funct_2__9_premises__rerun.E_input', d2_subset_1)).
fof(reflexivity_r2_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>r2_funct_2(X1,X2,X3,X3))), file('out/t201_funct_2__9_premises__rerun.E_input', reflexivity_r2_funct_2)).
fof(c_0_10, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>(![X3]:(m1_subset_1(X3,X1)=>k3_funct_2(X1,X1,X2,X3)=X3)=>r2_funct_2(X1,X1,X2,k6_partfun1(X1))))))), inference(assume_negation,[status(cth)],[t201_funct_2])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:((((k1_relat_1(X5)=X4|X5!=k6_relat_1(X4))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((~r2_hidden(X6,X4)|k1_funct_1(X5,X6)=X6)|X5!=k6_relat_1(X4))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((r2_hidden(esk3_2(X4,X5),X4)|k1_relat_1(X5)!=X4)|X5=k6_relat_1(X4))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((k1_funct_1(X5,esk3_2(X4,X5))!=esk3_2(X4,X5)|k1_relat_1(X5)!=X4)|X5=k6_relat_1(X4))|(~v1_relat_1(X5)|~v1_funct_1(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_funct_1])])])])])])])).
fof(c_0_12, plain, (![X2]:k6_partfun1(X2)=k6_relat_1(X2)), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:((((v1_funct_1(X6)|(~v1_funct_1(X6)|~v1_funct_2(X6,X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))|v1_xboole_0(X5))&(((v1_partfun1(X6,X4)|(~v1_funct_1(X6)|~v1_funct_2(X6,X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))|v1_xboole_0(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])])])).
fof(c_0_14, negated_conjecture, (![X6]:(~v1_xboole_0(esk1_0)&((v1_funct_1(esk2_0)&(v1_funct_2(esk2_0,esk1_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((~m1_subset_1(X6,esk1_0)|k3_funct_2(esk1_0,esk1_0,esk2_0,X6)=X6)&~r2_funct_2(esk1_0,esk1_0,esk2_0,k6_partfun1(esk1_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:((v4_relat_1(X6,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))&(v5_relat_1(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_17,plain,(X1=k6_relat_1(X2)|r2_hidden(esk3_2(X2,X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|k1_relat_1(X1)!=X2), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,plain,(k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_19, plain, (![X3]:![X4]:(((~v1_partfun1(X4,X3)|k1_relat_1(X4)=X3)|(~v1_relat_1(X4)|~v4_relat_1(X4,X3)))&((k1_relat_1(X4)!=X3|v1_partfun1(X4,X3))|(~v1_relat_1(X4)|~v4_relat_1(X4,X3))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_partfun1])])])).
cnf(c_0_20,plain,(v1_xboole_0(X1)|v1_partfun1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X2,X3,X1)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,negated_conjecture,(v1_funct_2(esk2_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,negated_conjecture,(v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25,plain,(v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_27, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X8,X5)))|k3_funct_2(X5,X6,X7,X8)=k1_funct_1(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_28,plain,(X1=k6_partfun1(X2)|r2_hidden(esk3_2(X2,X1),X2)|k1_relat_1(X1)!=X2|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_29,plain,(k1_relat_1(X1)=X2|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_partfun1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,negated_conjecture,(v1_partfun1(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_31,negated_conjecture,(v4_relat_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_25, c_0_21])).
cnf(c_0_32,negated_conjecture,(v1_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_26, c_0_21])).
cnf(c_0_33,plain,(X1=k6_relat_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|k1_relat_1(X1)!=X2|k1_funct_1(X1,esk3_2(X2,X1))!=esk3_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_34,plain,(k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
cnf(c_0_36,plain,(k6_partfun1(k1_relat_1(X1))=X1|r2_hidden(esk3_2(k1_relat_1(X1),X1),k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_37,negated_conjecture,(k1_relat_1(esk2_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_38,plain,(X1=k6_partfun1(X2)|k1_relat_1(X1)!=X2|k1_funct_1(X1,esk3_2(X2,X1))!=esk3_2(X2,X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_33, c_0_18])).
cnf(c_0_39,negated_conjecture,(k3_funct_2(esk1_0,esk1_0,esk2_0,X1)=X1|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_40,negated_conjecture,(k3_funct_2(esk1_0,esk1_0,esk2_0,X1)=k1_funct_1(esk2_0,X1)|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_21]), c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_41,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42,negated_conjecture,(k6_partfun1(esk1_0)=esk2_0|r2_hidden(esk3_2(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_32]), c_0_23])])).
fof(c_0_43, plain, (![X5]:![X6]:![X7]:![X8]:(((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|(~v1_funct_1(X8)|(~v1_funct_2(X8,X5,X6)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))))|r2_funct_2(X5,X6,X7,X7))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r2_funct_2])])])])).
cnf(c_0_44,plain,(k6_partfun1(k1_relat_1(X1))=X1|k1_funct_1(X1,esk3_2(k1_relat_1(X1),X1))!=esk3_2(k1_relat_1(X1),X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_45,negated_conjecture,(k1_funct_1(esk2_0,X1)=X1|~m1_subset_1(X1,esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46,negated_conjecture,(k6_partfun1(esk1_0)=esk2_0|m1_subset_1(esk3_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_24])).
cnf(c_0_47,plain,(r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X4,X1,X2)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48,negated_conjecture,(~r2_funct_2(esk1_0,esk1_0,esk2_0,k6_partfun1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49,negated_conjecture,(k6_partfun1(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_37]), c_0_32]), c_0_23]), c_0_37])]), c_0_46])).
cnf(c_0_50,negated_conjecture,(r2_funct_2(esk1_0,esk1_0,X1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_2(X1,esk1_0,esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_21]), c_0_22]), c_0_23])])).
cnf(c_0_51,negated_conjecture,(~r2_funct_2(esk1_0,esk1_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_21]), c_0_22]), c_0_23])]), c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 32
# Proof object formula steps           : 21
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 22
# Processed clauses                    : 46
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 43
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 37
# ...of the previous two non-trivial   : 31
# Contextual simplify-reflections      : 1
# Paramodulations                      : 32
# Factorizations                       : 0
# Equation resolutions                 : 5
# Current number of processed clauses  : 35
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 37
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 138
# Rec. Clause-clause subsumption calls : 49
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2454

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45560 pages
