# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_zfmisc_1, axiom, (![X1]:![X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), file('out/t20_mcart_1__15_premises__rerun.E_input', t113_zfmisc_1)).
fof(t12_mcart_1, axiom, (![X1]:![X2]:![X3]:(r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))=>(k1_mcart_1(X1)=X2&r2_hidden(k2_mcart_1(X1),X3)))), file('out/t20_mcart_1__15_premises__rerun.E_input', t12_mcart_1)).
fof(t20_mcart_1, conjecture, (![X1]:(?[X2]:?[X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1))))), file('out/t20_mcart_1__15_premises__rerun.E_input', t20_mcart_1)).
fof(d1_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3))))))), file('out/t20_mcart_1__15_premises__rerun.E_input', d1_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t20_mcart_1__15_premises__rerun.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), file('out/t20_mcart_1__15_premises__rerun.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), file('out/t20_mcart_1__15_premises__rerun.E_input', t73_enumset1)).
fof(t7_mcart_1, axiom, (![X1]:![X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2)), file('out/t20_mcart_1__15_premises__rerun.E_input', t7_mcart_1)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t20_mcart_1__15_premises__rerun.E_input', d10_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t20_mcart_1__15_premises__rerun.E_input', d3_tarski)).
fof(t7_ordinal1, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1)))), file('out/t20_mcart_1__15_premises__rerun.E_input', t7_ordinal1)).
fof(d1_mcart_1, axiom, (![X1]:(?[X2]:?[X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k1_mcart_1(X1)<=>![X3]:![X4]:(X1=k4_tarski(X3,X4)=>X2=X3)))), file('out/t20_mcart_1__15_premises__rerun.E_input', d1_mcart_1)).
fof(d2_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:?[X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6)))))), file('out/t20_mcart_1__15_premises__rerun.E_input', d2_zfmisc_1)).
fof(d2_tarski, axiom, (![X1]:![X2]:![X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2)))), file('out/t20_mcart_1__15_premises__rerun.E_input', d2_tarski)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t20_mcart_1__15_premises__rerun.E_input', t70_enumset1)).
fof(t9_mcart_1, axiom, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3]:![X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4)))))))), file('out/t20_mcart_1__15_premises__rerun.E_input', t9_mcart_1)).
fof(c_0_16, plain, (![X3]:![X4]:![X3]:![X4]:((k2_zfmisc_1(X3,X4)!=k1_xboole_0|(X3=k1_xboole_0|X4=k1_xboole_0))&((X3!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)&(X4!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])])])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:((k1_mcart_1(X4)=X5|~r2_hidden(X4,k2_zfmisc_1(k1_tarski(X5),X6)))&(r2_hidden(k2_mcart_1(X4),X6)|~r2_hidden(X4,k2_zfmisc_1(k1_tarski(X5),X6))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_mcart_1])])])).
cnf(c_0_18,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, negated_conjecture, (~(![X1]:(?[X2]:?[X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))))), inference(assume_negation,[status(cth)],[t20_mcart_1])).
fof(c_0_20, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X10]:![X6]:![X7]:![X8]:![X9]:((((~r2_hidden(X10,X9)|(X10=X6|(X10=X7|X10=X8)))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X6|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X7|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&((X10!=X8|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8)))))&((((esk4_4(X6,X7,X8,X9)!=X6|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&(((esk4_4(X6,X7,X8,X9)!=X7|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&((esk4_4(X6,X7,X8,X9)!=X8|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))))&((r2_hidden(esk4_4(X6,X7,X8,X9),X9)|(esk4_4(X6,X7,X8,X9)=X6|(esk4_4(X6,X7,X8,X9)=X7|esk4_4(X6,X7,X8,X9)=X8)))|X9=k1_enumset1(X6,X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_22, plain, (![X5]:![X6]:![X7]:![X8]:k3_enumset1(X5,X5,X6,X7,X8)=k2_enumset1(X5,X6,X7,X8)), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_23, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k4_enumset1(X6,X6,X7,X8,X9,X10)=k3_enumset1(X6,X7,X8,X9,X10)), inference(variable_rename,[status(thm)],[t73_enumset1])).
cnf(c_0_24,plain,(k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_18])).
fof(c_0_26, plain, (![X3]:![X4]:![X3]:![X4]:(k1_mcart_1(k4_tarski(X3,X4))=X3&k2_mcart_1(k4_tarski(X3,X4))=X4)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t7_mcart_1])])])).
fof(c_0_27, negated_conjecture, ((esk1_0=k4_tarski(esk2_0,esk3_0)&(esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])])])).
fof(c_0_28, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_29,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X4), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31,plain,(k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,plain,(k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33,plain,(k1_mcart_1(X1)=X2|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_35,plain,(k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36,negated_conjecture,(esk1_0=k4_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_37, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~r1_tarski(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_38,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39,plain,(r2_hidden(X5,X1)|X4!=X5|X1!=k4_enumset1(X2,X2,X2,X2,X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_40,plain,(X1=X2|~r2_hidden(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_33, c_0_33])).
cnf(c_0_41,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42,negated_conjecture,(esk2_0=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43,plain,(~r1_tarski(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_45,plain,(r2_hidden(X1,X2)|X2!=k4_enumset1(X3,X3,X3,X3,X4,X1)), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_46,plain,(X1=X2|r1_tarski(k1_xboole_0,X3)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_47, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X8]:(((X8!=k1_mcart_1(X5)|(X5!=k4_tarski(X9,X10)|X8=X9))|X5!=k4_tarski(X6,X7))&(((X5=k4_tarski(esk6_2(X5,X8),esk7_2(X5,X8))|X8=k1_mcart_1(X5))|X5!=k4_tarski(X6,X7))&((X8!=esk6_2(X5,X8)|X8=k1_mcart_1(X5))|X5!=k4_tarski(X6,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_mcart_1])])])])])])])).
cnf(c_0_48,plain,(k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_49,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),esk3_0)=esk1_0), inference(rw,[status(thm)],[c_0_36, c_0_42])).
cnf(c_0_50,plain,(~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51,plain,(r1_tarski(k1_xboole_0,X1)|r2_hidden(X2,X3)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_52,plain,(X4=k1_mcart_1(X1)|X1=k4_tarski(esk6_2(X1,X4),esk7_2(X1,X4))|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53,negated_conjecture,(esk3_0=k2_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_54,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X2), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_55,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_56, plain, (![X7]:![X8]:![X9]:![X10]:![X10]:![X13]:![X14]:![X7]:![X8]:![X9]:![X16]:![X17]:(((((r2_hidden(esk8_4(X7,X8,X9,X10),X7)|~r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8))&(((r2_hidden(esk9_4(X7,X8,X9,X10),X8)|~r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8))&((X10=k4_tarski(esk8_4(X7,X8,X9,X10),esk9_4(X7,X8,X9,X10))|~r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8))))&(((~r2_hidden(X13,X7)|(~r2_hidden(X14,X8)|X10!=k4_tarski(X13,X14)))|r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8)))&(((~r2_hidden(esk10_3(X7,X8,X9),X9)|(~r2_hidden(X16,X7)|(~r2_hidden(X17,X8)|esk10_3(X7,X8,X9)!=k4_tarski(X16,X17))))|X9=k2_zfmisc_1(X7,X8))&(((r2_hidden(esk11_3(X7,X8,X9),X7)|r2_hidden(esk10_3(X7,X8,X9),X9))|X9=k2_zfmisc_1(X7,X8))&(((r2_hidden(esk12_3(X7,X8,X9),X8)|r2_hidden(esk10_3(X7,X8,X9),X9))|X9=k2_zfmisc_1(X7,X8))&((esk10_3(X7,X8,X9)=k4_tarski(esk11_3(X7,X8,X9),esk12_3(X7,X8,X9))|r2_hidden(esk10_3(X7,X8,X9),X9))|X9=k2_zfmisc_1(X7,X8))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])])).
cnf(c_0_57,plain,(k4_tarski(esk6_2(k4_tarski(X1,X2),X3),esk7_2(k4_tarski(X1,X2),X3))=k4_tarski(X1,X2)|X3=X1), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_52]), c_0_35])).
cnf(c_0_58,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_49, c_0_53])).
cnf(c_0_59,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X3), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_60,plain,(r2_hidden(X5,X1)|X2!=X5|X1!=k4_enumset1(X2,X2,X2,X2,X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_61,plain,(~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_43, c_0_55])).
cnf(c_0_62,plain,(r2_hidden(esk9_4(X2,X3,X1,X4),X3)|X1!=k2_zfmisc_1(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63,negated_conjecture,(k4_tarski(esk6_2(esk1_0,X1),esk7_2(esk1_0,X1))=esk1_0|X1=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_64, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(X8=X5|X8=X6))|X7!=k2_tarski(X5,X6))&(((X8!=X5|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))&((X8!=X6|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))))&((((esk13_3(X5,X6,X7)!=X5|~r2_hidden(esk13_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6))&((esk13_3(X5,X6,X7)!=X6|~r2_hidden(esk13_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6)))&((r2_hidden(esk13_3(X5,X6,X7),X7)|(esk13_3(X5,X6,X7)=X5|esk13_3(X5,X6,X7)=X6))|X7=k2_tarski(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])])).
fof(c_0_65, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_66,plain,(r2_hidden(X5,X1)|X3!=X5|X1!=k4_enumset1(X2,X2,X2,X2,X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_30]), c_0_31]), c_0_32])).
fof(c_0_67, plain, (![X5]:![X7]:![X8]:((r2_hidden(esk14_1(X5),X5)|X5=k1_xboole_0)&(((~r2_hidden(X7,X5)|esk14_1(X5)!=k4_tarski(X7,X8))|X5=k1_xboole_0)&((~r2_hidden(X8,X5)|esk14_1(X5)!=k4_tarski(X7,X8))|X5=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])])])).
cnf(c_0_68,plain,(r2_hidden(X1,X2)|X2!=k4_enumset1(X1,X1,X1,X1,X3,X4)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_69,plain,(X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_25])).
cnf(c_0_70,plain,(r2_hidden(X1,k4_enumset1(X2,X2,X2,X2,X3,X1))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_71,negated_conjecture,(esk7_2(esk1_0,X1)=k2_mcart_1(esk1_0)|X1=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_63])).
cnf(c_0_72,plain,(X4=X3|X4=X2|X1!=k2_tarski(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74,plain,(r2_hidden(X1,X2)|X2!=k4_enumset1(X3,X3,X3,X3,X1,X4)), inference(er,[status(thm)],[c_0_66])).
cnf(c_0_75,plain,(X1=k1_xboole_0|esk14_1(X1)!=k4_tarski(X2,X3)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76,plain,(r2_hidden(X1,k4_enumset1(X1,X1,X1,X1,X2,X3))), inference(er,[status(thm)],[c_0_68])).
cnf(c_0_77,plain,(k4_enumset1(X1,X1,X1,X1,X2,X3)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78,negated_conjecture,(k4_tarski(esk6_2(esk1_0,X1),k2_mcart_1(esk1_0))=esk1_0|X1=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_71])).
cnf(c_0_79,negated_conjecture,(esk1_0=k2_mcart_1(esk1_0)|esk1_0=k1_mcart_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_80,plain,(X3=X4|X2=X4|X1!=k4_enumset1(X2,X2,X2,X2,X2,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_81,plain,(r2_hidden(X1,k4_enumset1(X2,X2,X2,X2,X1,X3))), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_82,plain,(esk14_1(k4_enumset1(X1,X1,X1,X1,X2,X3))!=k4_tarski(X4,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])).
cnf(c_0_83,negated_conjecture,(k4_tarski(esk6_2(esk1_0,X1),esk1_0)=esk1_0|k1_mcart_1(esk1_0)=esk1_0|X1=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_84,plain,(X1=X2|X3=X2|~r2_hidden(X2,k4_enumset1(X1,X1,X1,X1,X1,X3))), inference(er,[status(thm)],[c_0_80])).
cnf(c_0_85,plain,(X1=k1_xboole_0|r2_hidden(esk14_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_86,plain,(esk14_1(k4_enumset1(X1,X1,X1,X1,X2,X3))!=k4_tarski(X4,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_81]), c_0_77])).
cnf(c_0_87,plain,(X1=k1_xboole_0|esk14_1(X1)!=k4_tarski(X2,X3)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_88,negated_conjecture,(esk6_2(esk1_0,X1)=k1_mcart_1(esk1_0)|X1=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_63])).
cnf(c_0_89,negated_conjecture,(k1_mcart_1(esk1_0)=esk1_0|X1=k1_mcart_1(esk1_0)|esk14_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X2,X3))!=esk1_0), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_90,plain,(esk14_1(k4_enumset1(X1,X1,X1,X1,X1,X2))=X2|esk14_1(k4_enumset1(X1,X1,X1,X1,X1,X2))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_77])).
cnf(c_0_91,negated_conjecture,(k1_mcart_1(esk1_0)=esk1_0|X1=k1_mcart_1(esk1_0)|esk14_1(k4_enumset1(X2,X2,X2,X2,esk1_0,X3))!=esk1_0), inference(spm,[status(thm)],[c_0_86, c_0_83])).
cnf(c_0_92,plain,(X1!=k2_zfmisc_1(X2,X3)|X3!=k1_xboole_0|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_69, c_0_62])).
cnf(c_0_93,plain,(X4=k4_tarski(esk8_4(X2,X3,X1,X4),esk9_4(X2,X3,X1,X4))|X1!=k2_zfmisc_1(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_94,plain,(esk14_1(k4_enumset1(X1,X1,X1,X1,X2,X3))!=k4_tarski(X1,X4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_76]), c_0_77])).
cnf(c_0_95,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),esk7_2(esk1_0,X1))=esk1_0|X1=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_88])).
cnf(c_0_96,negated_conjecture,(k1_mcart_1(esk1_0)=esk1_0|X1=k1_mcart_1(esk1_0)|X2!=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])).
cnf(c_0_97,plain,(k4_tarski(X1,esk9_4(X2,X3,X4,X5))!=esk14_1(X3)|X4!=k2_zfmisc_1(X2,X3)|~r2_hidden(X5,X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_62]), c_0_92])).
cnf(c_0_98,plain,(esk9_4(X1,X2,X3,X4)=k2_mcart_1(X4)|X3!=k2_zfmisc_1(X1,X2)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_48, c_0_93])).
cnf(c_0_99,negated_conjecture,(X1=k1_mcart_1(esk1_0)|esk14_1(k4_enumset1(k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),X2,X3))!=esk1_0), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_100,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),esk1_0)=esk1_0|k1_mcart_1(esk1_0)=esk1_0), inference(spm,[status(thm)],[c_0_58, c_0_79])).
cnf(c_0_101,negated_conjecture,(k1_mcart_1(esk1_0)=esk1_0|X1=k1_mcart_1(esk1_0)), inference(er,[status(thm)],[c_0_96])).
cnf(c_0_102,plain,(k4_tarski(X1,k2_mcart_1(X2))!=esk14_1(X3)|X4!=k2_zfmisc_1(X5,X3)|~r2_hidden(X2,X4)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_103,negated_conjecture,(esk14_1(k4_enumset1(k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),X1))=X1|X2=k1_mcart_1(esk1_0)|k1_mcart_1(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_99, c_0_90])).
cnf(c_0_104,negated_conjecture,(k1_mcart_1(esk1_0)=esk1_0), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_105,plain,(k4_tarski(X1,X2)!=esk14_1(X3)|X4!=k2_zfmisc_1(X5,X3)|~r2_hidden(k4_tarski(X6,X2),X4)), inference(spm,[status(thm)],[c_0_102, c_0_48])).
cnf(c_0_106,plain,(esk14_1(k4_enumset1(X1,X1,X1,X1,X2,X3))!=k4_tarski(X2,X4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_81]), c_0_77])).
cnf(c_0_107,negated_conjecture,(esk14_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1))=X1|X2=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104]), c_0_104]), c_0_104]), c_0_104]), c_0_104]), c_0_104]), c_0_104])])).
cnf(c_0_108,plain,(k4_enumset1(X1,X1,X1,X1,k4_tarski(X2,X3),X4)!=k2_zfmisc_1(X5,X6)|k4_tarski(X7,X3)!=esk14_1(X6)), inference(spm,[status(thm)],[c_0_105, c_0_81])).
cnf(c_0_109,negated_conjecture,(X1=esk1_0|X2!=k4_tarski(esk1_0,X3)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_110,negated_conjecture,(k4_enumset1(X1,X1,X1,X1,esk1_0,X2)!=k2_zfmisc_1(X3,X4)|k4_tarski(X5,k2_mcart_1(esk1_0))!=esk14_1(X4)), inference(spm,[status(thm)],[c_0_108, c_0_58])).
cnf(c_0_111,negated_conjecture,(X1=esk1_0), inference(er,[status(thm)],[c_0_109])).
cnf(c_0_112,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_111]), c_0_111]), c_0_111])]), c_0_111]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 113
# Proof object clause steps            : 80
# Proof object formula steps           : 33
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 42
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 48
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 44
# Processed clauses                    : 2181
# ...of these trivial                  : 8
# ...subsumed                          : 1487
# ...remaining for further processing  : 686
# Other redundant clauses eliminated   : 795
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 87
# Backward-rewritten                   : 373
# Generated clauses                    : 27978
# ...of the previous two non-trivial   : 26318
# Contextual simplify-reflections      : 1465
# Paramodulations                      : 26919
# Factorizations                       : 114
# Equation resolutions                 : 945
# Current number of processed clauses  : 178
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 165
# Current number of unprocessed clauses: 665
# ...number of literals in the above   : 2114
# Current number of archived formulas  : 0
# Current number of archived clauses   : 505
# Clause-clause subsumption calls (NU) : 120279
# Rec. Clause-clause subsumption calls : 58248
# Non-unit clause-clause subsumptions  : 2665
# Unit Clause-clause subsumption calls : 556
# Rewrite failures with RHS unbound    : 2
# BW rewrite match attempts            : 450
# BW rewrite match successes           : 442
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 273065

# -------------------------------------------------
# User time                : 0.432 s
# System time              : 0.008 s
# Total time               : 0.440 s
# Maximum resident set size: 44924 pages
