# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t20_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1)))))), file('out/t20_tex_2__17_premises__rerun.E_input', t20_tex_2)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t20_tex_2__17_premises__rerun.E_input', dt_l1_pre_topc)).
fof(t8_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~((v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))&v7_struct_0(X1)))))), file('out/t20_tex_2__17_premises__rerun.E_input', t8_tex_2)).
fof(dt_k1_tex_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v1_pre_topc(k1_tex_2(X1,X2))&m1_pre_topc(k1_tex_2(X1,X2),X1))))), file('out/t20_tex_2__17_premises__rerun.E_input', dt_k1_tex_2)).
fof(cc10_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v7_struct_0(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>(~(v2_struct_0(X2))=>(~(v2_struct_0(X2))&~(v1_tex_2(X2,X1))))))), file('out/t20_tex_2__17_premises__rerun.E_input', cc10_tex_2)).
fof(t9_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&(~(v7_struct_0(X1))&l1_struct_0(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))))), file('out/t20_tex_2__17_premises__rerun.E_input', t9_tex_2)).
fof(d3_tex_2, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v1_subset_1(X3,u1_struct_0(X1)))))))), file('out/t20_tex_2__17_premises__rerun.E_input', d3_tex_2)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t20_tex_2__17_premises__rerun.E_input', dt_m1_pre_topc)).
fof(cc4_subset_1, axiom, (![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>~(v1_subset_1(X2,X1))))), file('out/t20_tex_2__17_premises__rerun.E_input', cc4_subset_1)).
fof(t1_tsep_1, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t20_tex_2__17_premises__rerun.E_input', t1_tsep_1)).
fof(d7_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1)))), file('out/t20_tex_2__17_premises__rerun.E_input', d7_subset_1)).
fof(fc2_tex_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v7_struct_0(k1_tex_2(X1,X2))&v1_pre_topc(k1_tex_2(X1,X2)))))), file('out/t20_tex_2__17_premises__rerun.E_input', fc2_tex_2)).
fof(fc7_struct_0, axiom, (![X1]:((v7_struct_0(X1)&l1_struct_0(X1))=>v1_zfmisc_1(u1_struct_0(X1)))), file('out/t20_tex_2__17_premises__rerun.E_input', fc7_struct_0)).
fof(dt_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1)))), file('out/t20_tex_2__17_premises__rerun.E_input', dt_k6_domain_1)).
fof(t6_tex_2, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>~((v1_subset_1(k6_domain_1(X1,X2),X1)&v1_zfmisc_1(X1)))))), file('out/t20_tex_2__17_premises__rerun.E_input', t6_tex_2)).
fof(fc6_struct_0, axiom, (![X1]:((~(v7_struct_0(X1))&l1_struct_0(X1))=>~(v1_zfmisc_1(u1_struct_0(X1))))), file('out/t20_tex_2__17_premises__rerun.E_input', fc6_struct_0)).
fof(fc14_subset_1, axiom, (![X1]:~(v1_subset_1(k2_subset_1(X1),X1))), file('out/t20_tex_2__17_premises__rerun.E_input', fc14_subset_1)).
fof(d4_subset_1, axiom, (![X1]:k2_subset_1(X1)=X1), file('out/t20_tex_2__17_premises__rerun.E_input', d4_subset_1)).
fof(c_0_18, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t20_tex_2])).
fof(c_0_19, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_20, negated_conjecture, (((~v2_struct_0(esk1_0)&l1_pre_topc(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0)))&(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])])])).
fof(c_0_21, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_struct_0(X3))|(~m1_subset_1(X4,u1_struct_0(X3))|(~v1_subset_1(k6_domain_1(u1_struct_0(X3),X4),u1_struct_0(X3))|~v7_struct_0(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_tex_2])])])])])])).
cnf(c_0_22,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, plain, (![X3]:![X4]:((~v2_struct_0(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&((v1_pre_topc(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&(m1_pre_topc(k1_tex_2(X3,X4),X3)|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tex_2])])])])).
fof(c_0_25, plain, (![X3]:![X4]:((((~v2_struct_0(X4)|v2_struct_0(X4))|~m1_pre_topc(X4,X3))|(v2_struct_0(X3)|(~v7_struct_0(X3)|~l1_pre_topc(X3))))&(((~v1_tex_2(X4,X3)|v2_struct_0(X4))|~m1_pre_topc(X4,X3))|(v2_struct_0(X3)|(~v7_struct_0(X3)|~l1_pre_topc(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc10_tex_2])])])])])])])).
cnf(c_0_26,plain,(v2_struct_0(X1)|~v7_struct_0(X1)|~v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,negated_conjecture,(v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))|v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30,negated_conjecture,(l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31,plain,(v2_struct_0(X2)|m1_pre_topc(k1_tex_2(X2,X1),X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, (![X3]:![X4]:((v2_struct_0(X3)|(v7_struct_0(X3)|~l1_struct_0(X3)))|(~m1_subset_1(X4,u1_struct_0(X3))|v1_subset_1(k6_domain_1(u1_struct_0(X3),X4),u1_struct_0(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_tex_2])])])])])])).
fof(c_0_33, plain, (![X4]:![X5]:![X6]:((((~v1_tex_2(X5,X4)|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(X6!=u1_struct_0(X5)|v1_subset_1(X6,u1_struct_0(X4)))))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4))&((((m1_subset_1(esk3_2(X4,X5),k1_zfmisc_1(u1_struct_0(X4)))|v1_tex_2(X5,X4))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4))&((((esk3_2(X4,X5)=u1_struct_0(X5)|v1_tex_2(X5,X4))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4))&(((~v1_subset_1(esk3_2(X4,X5),u1_struct_0(X4))|v1_tex_2(X5,X4))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tex_2])])])])])])])).
cnf(c_0_34,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X1)|~v7_struct_0(X1)|~m1_pre_topc(X2,X1)|~v1_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35,negated_conjecture,(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v7_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])]), c_0_29]), c_0_30])])).
cnf(c_0_36,negated_conjecture,(m1_pre_topc(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_23])]), c_0_29])).
fof(c_0_37, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
fof(c_0_38, plain, (![X3]:![X4]:(~v1_xboole_0(X3)|(~m1_subset_1(X4,k1_zfmisc_1(X3))|~v1_subset_1(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc4_subset_1])])])])])])).
fof(c_0_39, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(X3)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])])])).
fof(c_0_40, plain, (![X3]:![X4]:(((~v1_subset_1(X4,X3)|X4!=X3)|~m1_subset_1(X4,k1_zfmisc_1(X3)))&((X4=X3|v1_subset_1(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(X3))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_41,negated_conjecture,(~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))|~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_42,plain,(v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))|v7_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43,plain,(v1_tex_2(X2,X1)|esk3_2(X1,X2)=u1_struct_0(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44,plain,(v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_struct_0(k1_tex_2(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_45,negated_conjecture,(v2_struct_0(k1_tex_2(esk1_0,esk2_0))|~v7_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_23])]), c_0_29])).
fof(c_0_46, plain, (![X3]:![X4]:((~v2_struct_0(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&((v7_struct_0(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&(v1_pre_topc(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_tex_2])])])])).
cnf(c_0_47,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48,plain,(~v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49,plain,(m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50,plain,(v1_subset_1(X1,X2)|X1=X2|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_51, plain, (![X2]:((~v7_struct_0(X2)|~l1_struct_0(X2))|v1_zfmisc_1(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_struct_0])])).
cnf(c_0_52,negated_conjecture,(v7_struct_0(esk1_0)|~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_28])]), c_0_29]), c_0_30])])).
cnf(c_0_53,negated_conjecture,(u1_struct_0(X1)=esk3_2(esk1_0,X1)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_43, c_0_23])).
cnf(c_0_54,negated_conjecture,(~v7_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_28]), c_0_23])]), c_0_29])).
cnf(c_0_55,plain,(v2_struct_0(X2)|v7_struct_0(k1_tex_2(X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56,negated_conjecture,(l1_pre_topc(k1_tex_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_36]), c_0_23])])).
cnf(c_0_57,plain,(~v1_xboole_0(u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~v1_subset_1(u1_struct_0(X2),u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_58,plain,(u1_struct_0(X1)=u1_struct_0(X2)|v1_subset_1(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_50, c_0_49])).
fof(c_0_59, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|m1_subset_1(k6_domain_1(X3,X4),k1_zfmisc_1(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_60,plain,(v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61,negated_conjecture,(u1_struct_0(k1_tex_2(esk1_0,esk2_0))=esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_36])]), c_0_54])).
cnf(c_0_62,negated_conjecture,(v7_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_28]), c_0_23])]), c_0_29])).
cnf(c_0_63,negated_conjecture,(l1_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_22, c_0_56])).
cnf(c_0_64,plain,(u1_struct_0(X1)=u1_struct_0(X2)|~v1_xboole_0(u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_65, plain, (![X3]:![X4]:(v1_xboole_0(X3)|(~m1_subset_1(X4,X3)|(~v1_subset_1(k6_domain_1(X3,X4),X3)|~v1_zfmisc_1(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_tex_2])])])])])])).
cnf(c_0_66,plain,(m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_67, plain, (![X2]:((v7_struct_0(X2)|~l1_struct_0(X2))|~v1_zfmisc_1(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_struct_0])])])).
cnf(c_0_68,negated_conjecture,(v1_zfmisc_1(esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])])).
cnf(c_0_69,negated_conjecture,(esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)|~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_36]), c_0_61]), c_0_23])])).
cnf(c_0_70,plain,(v1_tex_2(X2,X1)|m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_71,plain,(v1_tex_2(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_72,plain,(v1_xboole_0(X1)|~v1_zfmisc_1(X1)|~v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73,plain,(k6_domain_1(X1,X2)=X1|v1_xboole_0(X1)|v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_50, c_0_66])).
cnf(c_0_74,plain,(v7_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75,negated_conjecture,(v1_zfmisc_1(u1_struct_0(esk1_0))|~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_76, plain, (![X2]:~v1_subset_1(k2_subset_1(X2),X2)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
fof(c_0_77, plain, (![X2]:k2_subset_1(X2)=X2), inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_78,plain,(esk3_2(X1,X2)=u1_struct_0(X1)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_70]), c_0_71])).
cnf(c_0_79,plain,(k6_domain_1(X1,X2)=X1|v1_xboole_0(X1)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_80,negated_conjecture,(~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_30])]), c_0_54])).
cnf(c_0_81,plain,(~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_82,plain,(k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_83,negated_conjecture,(esk3_2(esk1_0,X1)=u1_struct_0(esk1_0)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_78, c_0_23])).
cnf(c_0_84,negated_conjecture,(k6_domain_1(u1_struct_0(esk1_0),esk2_0)=u1_struct_0(esk1_0)|~v1_zfmisc_1(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_28]), c_0_80])).
cnf(c_0_85,plain,(~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_86,negated_conjecture,(esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_83]), c_0_36])]), c_0_54])).
cnf(c_0_87,negated_conjecture,(~v1_zfmisc_1(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_84]), c_0_30]), c_0_28])]), c_0_54]), c_0_85]), c_0_29])).
cnf(c_0_88,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_86]), c_0_87]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 89
# Proof object clause steps            : 52
# Proof object formula steps           : 37
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 29
# Processed clauses                    : 105
# ...of these trivial                  : 0
# ...subsumed                          : 19
# ...remaining for further processing  : 86
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 4
# Generated clauses                    : 90
# ...of the previous two non-trivial   : 81
# Contextual simplify-reflections      : 4
# Paramodulations                      : 89
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 66
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 512
# Rec. Clause-clause subsumption calls : 218
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 25
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4351

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45308 pages
