# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t20_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1)))))), file('out/t20_tex_2__11_premises__rerun.E_input', t20_tex_2)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t20_tex_2__11_premises__rerun.E_input', dt_l1_pre_topc)).
fof(t8_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~((v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))&v7_struct_0(X1)))))), file('out/t20_tex_2__11_premises__rerun.E_input', t8_tex_2)).
fof(dt_k1_tex_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v1_pre_topc(k1_tex_2(X1,X2))&m1_pre_topc(k1_tex_2(X1,X2),X1))))), file('out/t20_tex_2__11_premises__rerun.E_input', dt_k1_tex_2)).
fof(cc10_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v7_struct_0(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>(~(v2_struct_0(X2))=>(~(v2_struct_0(X2))&~(v1_tex_2(X2,X1))))))), file('out/t20_tex_2__11_premises__rerun.E_input', cc10_tex_2)).
fof(d7_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1)))), file('out/t20_tex_2__11_premises__rerun.E_input', d7_subset_1)).
fof(d3_tex_2, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v1_subset_1(X3,u1_struct_0(X1)))))))), file('out/t20_tex_2__11_premises__rerun.E_input', d3_tex_2)).
fof(t9_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&(~(v7_struct_0(X1))&l1_struct_0(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))))), file('out/t20_tex_2__11_premises__rerun.E_input', t9_tex_2)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t20_tex_2__11_premises__rerun.E_input', dt_m1_pre_topc)).
fof(fc2_tex_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v7_struct_0(k1_tex_2(X1,X2))&v1_pre_topc(k1_tex_2(X1,X2)))))), file('out/t20_tex_2__11_premises__rerun.E_input', fc2_tex_2)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t20_tex_2__11_premises__rerun.E_input', fc2_struct_0)).
fof(fc1_struct_0, axiom, (![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1)))), file('out/t20_tex_2__11_premises__rerun.E_input', fc1_struct_0)).
fof(c_0_12, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t20_tex_2])).
fof(c_0_13, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_14, negated_conjecture, (((~v2_struct_0(esk1_0)&l1_pre_topc(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0)))&(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])).
fof(c_0_15, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l1_struct_0(X3))|(~m1_subset_1(X4,u1_struct_0(X3))|(~v1_subset_1(k6_domain_1(u1_struct_0(X3),X4),u1_struct_0(X3))|~v7_struct_0(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_tex_2])])])])])])).
cnf(c_0_16,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, (![X3]:![X4]:((~v2_struct_0(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&((v1_pre_topc(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&(m1_pre_topc(k1_tex_2(X3,X4),X3)|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tex_2])])])])).
fof(c_0_19, plain, (![X3]:![X4]:((((~v2_struct_0(X4)|v2_struct_0(X4))|~m1_pre_topc(X4,X3))|(v2_struct_0(X3)|(~v7_struct_0(X3)|~l1_pre_topc(X3))))&(((~v1_tex_2(X4,X3)|v2_struct_0(X4))|~m1_pre_topc(X4,X3))|(v2_struct_0(X3)|(~v7_struct_0(X3)|~l1_pre_topc(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc10_tex_2])])])])])])])).
cnf(c_0_20,plain,(v2_struct_0(X1)|~v7_struct_0(X1)|~v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,negated_conjecture,(v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))|v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,negated_conjecture,(l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_25,plain,(v2_struct_0(X2)|m1_pre_topc(k1_tex_2(X2,X1),X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, plain, (![X3]:![X4]:(((~v1_subset_1(X4,X3)|X4!=X3)|~m1_subset_1(X4,k1_zfmisc_1(X3)))&((X4=X3|v1_subset_1(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(X3))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
fof(c_0_27, plain, (![X4]:![X5]:![X6]:((((~v1_tex_2(X5,X4)|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(X6!=u1_struct_0(X5)|v1_subset_1(X6,u1_struct_0(X4)))))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4))&((((m1_subset_1(esk3_2(X4,X5),k1_zfmisc_1(u1_struct_0(X4)))|v1_tex_2(X5,X4))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4))&((((esk3_2(X4,X5)=u1_struct_0(X5)|v1_tex_2(X5,X4))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4))&(((~v1_subset_1(esk3_2(X4,X5),u1_struct_0(X4))|v1_tex_2(X5,X4))|~m1_pre_topc(X5,X4))|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tex_2])])])])])])])).
fof(c_0_28, plain, (![X3]:![X4]:((v2_struct_0(X3)|(v7_struct_0(X3)|~l1_struct_0(X3)))|(~m1_subset_1(X4,u1_struct_0(X3))|v1_subset_1(k6_domain_1(u1_struct_0(X3),X4),u1_struct_0(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_tex_2])])])])])])).
cnf(c_0_29,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X1)|~v7_struct_0(X1)|~m1_pre_topc(X2,X1)|~v1_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,negated_conjecture,(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v7_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])]), c_0_23]), c_0_24])])).
cnf(c_0_31,negated_conjecture,(m1_pre_topc(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_17])]), c_0_23])).
cnf(c_0_32,plain,(v1_subset_1(X1,X2)|X1=X2|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,plain,(v1_tex_2(X2,X1)|m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34,plain,(v1_tex_2(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35,negated_conjecture,(~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))|~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36,plain,(v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))|v7_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37,plain,(v1_tex_2(X2,X1)|esk3_2(X1,X2)=u1_struct_0(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38,plain,(v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_struct_0(k1_tex_2(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_39,negated_conjecture,(v2_struct_0(k1_tex_2(esk1_0,esk2_0))|~v7_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_17])]), c_0_23])).
cnf(c_0_40,plain,(esk3_2(X1,X2)=u1_struct_0(X1)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
fof(c_0_41, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
cnf(c_0_42,negated_conjecture,(v7_struct_0(esk1_0)|~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_22])]), c_0_23]), c_0_24])])).
cnf(c_0_43,negated_conjecture,(u1_struct_0(X1)=esk3_2(esk1_0,X1)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_17])).
cnf(c_0_44,negated_conjecture,(~v7_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_22]), c_0_17])]), c_0_23])).
cnf(c_0_45,negated_conjecture,(esk3_2(esk1_0,X1)=u1_struct_0(esk1_0)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_17])).
fof(c_0_46, plain, (![X3]:![X4]:((~v2_struct_0(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&((v7_struct_0(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&(v1_pre_topc(k1_tex_2(X3,X4))|((v2_struct_0(X3)|~l1_pre_topc(X3))|~m1_subset_1(X4,u1_struct_0(X3))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_tex_2])])])])).
cnf(c_0_47,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48,negated_conjecture,(u1_struct_0(k1_tex_2(esk1_0,esk2_0))=esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_31])]), c_0_44])).
cnf(c_0_49,negated_conjecture,(esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_45]), c_0_31])]), c_0_44])).
cnf(c_0_50,plain,(v2_struct_0(X2)|v7_struct_0(k1_tex_2(X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_51,negated_conjecture,(l1_pre_topc(k1_tex_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_31]), c_0_17])])).
cnf(c_0_52,negated_conjecture,(u1_struct_0(k1_tex_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53,negated_conjecture,(v7_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_22]), c_0_17])]), c_0_23])).
cnf(c_0_54,negated_conjecture,(l1_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_16, c_0_51])).
fof(c_0_55, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_56, plain, (![X2]:((~v2_struct_0(X2)|~l1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_57,negated_conjecture,(v2_struct_0(k1_tex_2(esk1_0,esk2_0))|~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_52]), c_0_53]), c_0_54])])).
cnf(c_0_58,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_59,plain,(v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_60,negated_conjecture,(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|v2_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_21]), c_0_22])])).
cnf(c_0_61,negated_conjecture,(~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_58]), c_0_24])])).
cnf(c_0_62,negated_conjecture,(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_48]), c_0_49]), c_0_54])]), c_0_61])).
cnf(c_0_63,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_62])]), c_0_44]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 39
# Proof object formula steps           : 25
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 88
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 17
# Generated clauses                    : 67
# ...of the previous two non-trivial   : 71
# Contextual simplify-reflections      : 2
# Paramodulations                      : 66
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 27
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 71
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 433
# Rec. Clause-clause subsumption calls : 226
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 19
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3858

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45300 pages
