# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3))))))), file('out/t21_waybel_0__7_premises__rerun.E_input', t21_waybel_0)).
fof(dt_k5_waybel_0, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t21_waybel_0__7_premises__rerun.E_input', dt_k5_waybel_0)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t21_waybel_0__7_premises__rerun.E_input', t4_subset)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t21_waybel_0__7_premises__rerun.E_input', d3_tarski)).
fof(d8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4))))))), file('out/t21_waybel_0__7_premises__rerun.E_input', d8_lattice3)).
fof(t18_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3)))))), file('out/t21_waybel_0__7_premises__rerun.E_input', t18_waybel_0)).
fof(t17_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), file('out/t21_waybel_0__7_premises__rerun.E_input', t17_waybel_0)).
fof(t12_yellow_0, axiom, (![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X3)&r1_orders_2(X1,X4,X3))=>r1_lattice3(X1,X2,X4)))))), file('out/t21_waybel_0__7_premises__rerun.E_input', t12_yellow_0)).
fof(c_0_8, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3)))))))), inference(assume_negation,[status(cth)],[t21_waybel_0])).
fof(c_0_9, plain, (![X3]:![X4]:(((v2_struct_0(X3)|~l1_orders_2(X3))|~m1_subset_1(X4,u1_struct_0(X3)))|m1_subset_1(k5_waybel_0(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
fof(c_0_10, negated_conjecture, (((~v2_struct_0(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(r1_orders_2(esk1_0,esk2_0,esk3_0)&~r1_tarski(k5_waybel_0(esk1_0,esk2_0),k5_waybel_0(esk1_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_12,plain,(m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])]), c_0_15])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk4_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk4_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:((((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk5_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk5_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,X7,esk5_3(X5,X6,X7))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(((((~r2_hidden(X6,k6_waybel_0(X4,X5))|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((~r1_orders_2(X4,X5,X6)|r2_hidden(X6,k6_waybel_0(X4,X5)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])])])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:(((((~r2_hidden(X6,k5_waybel_0(X4,X5))|r1_orders_2(X4,X6,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((~r1_orders_2(X4,X6,X5)|r2_hidden(X6,k5_waybel_0(X4,X5)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])])])).
cnf(c_0_22,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,plain,(r1_tarski(X1,X2)|r2_hidden(esk4_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(r1_lattice3(X1,X3,X2)|m1_subset_1(esk5_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,plain,(v2_struct_0(X1)|r2_hidden(X3,k6_waybel_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_27, plain, (![X5]:![X6]:![X7]:![X8]:((~v4_orders_2(X5)|~l1_orders_2(X5))|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|((~r1_lattice3(X5,X6,X7)|~r1_orders_2(X5,X8,X7))|r1_lattice3(X5,X6,X8)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_yellow_0])])])])])).
cnf(c_0_28,plain,(v2_struct_0(X1)|r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k5_waybel_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|m1_subset_1(esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30,plain,(v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k6_waybel_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31,negated_conjecture,(r1_lattice3(esk1_0,X1,esk2_0)|m1_subset_1(esk5_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_13]), c_0_14])])).
cnf(c_0_32,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk5_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34,negated_conjecture,(r2_hidden(esk3_0,k6_waybel_0(esk1_0,X1))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_14])]), c_0_15])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36,plain,(r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_38,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),X2)|~r2_hidden(esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_14])]), c_0_15])).
cnf(c_0_39,negated_conjecture,(r1_lattice3(esk1_0,X1,esk2_0)|r1_orders_2(esk1_0,X2,esk5_3(esk1_0,X1,esk2_0))|~r2_hidden(esk5_3(esk1_0,X1,esk2_0),k6_waybel_0(esk1_0,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_14])]), c_0_15])).
cnf(c_0_40,negated_conjecture,(r1_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk5_3(esk1_0,X1,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_13]), c_0_14])])).
cnf(c_0_41,plain,(r1_lattice3(X1,X3,X2)|r2_hidden(esk5_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42,negated_conjecture,(r1_orders_2(esk1_0,X1,esk3_0)|~r1_lattice3(esk1_0,X2,X1)|~r2_hidden(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_26]), c_0_14])])).
cnf(c_0_43,negated_conjecture,(r2_hidden(esk3_0,k6_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_13])])).
cnf(c_0_44,negated_conjecture,(r1_lattice3(esk1_0,X1,X2)|~r1_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X2,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_13]), c_0_14]), c_0_37])])).
cnf(c_0_45,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_13]), c_0_23])).
cnf(c_0_46,negated_conjecture,(r1_lattice3(esk1_0,X1,esk2_0)|~r2_hidden(esk5_3(esk1_0,X1,esk2_0),k6_waybel_0(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_13]), c_0_40])).
cnf(c_0_47,negated_conjecture,(r1_lattice3(esk1_0,X1,esk2_0)|r2_hidden(esk5_3(esk1_0,X1,esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_13]), c_0_14])])).
cnf(c_0_48,negated_conjecture,(r1_orders_2(esk1_0,X1,esk3_0)|~r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49,negated_conjecture,(r1_lattice3(esk1_0,X1,esk4_2(k5_waybel_0(esk1_0,esk2_0),X2))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X2)|~r1_lattice3(esk1_0,X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_29])).
cnf(c_0_50,negated_conjecture,(r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51,plain,(v2_struct_0(X1)|r2_hidden(X3,k5_waybel_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_52,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),esk3_0)|~r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk4_2(k5_waybel_0(esk1_0,esk2_0),X1))), inference(spm,[status(thm)],[c_0_48, c_0_29])).
cnf(c_0_53,negated_conjecture,(r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk4_2(k5_waybel_0(esk1_0,esk2_0),X1))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54,negated_conjecture,(r2_hidden(esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,X2))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|~r1_orders_2(esk1_0,esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_29]), c_0_14])]), c_0_15])).
cnf(c_0_55,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),esk3_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_56,plain,(r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_57,negated_conjecture,(r2_hidden(esk4_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,esk3_0))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_26]), c_0_55])).
cnf(c_0_58,negated_conjecture,(~r1_tarski(k5_waybel_0(esk1_0,esk2_0),k5_waybel_0(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_59,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 43
# Proof object formula steps           : 17
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 266
# ...of these trivial                  : 0
# ...subsumed                          : 32
# ...remaining for further processing  : 234
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 0
# Generated clauses                    : 485
# ...of the previous two non-trivial   : 471
# Contextual simplify-reflections      : 32
# Paramodulations                      : 485
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 207
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 192
# Current number of unprocessed clauses: 247
# ...number of literals in the above   : 1070
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 6013
# Rec. Clause-clause subsumption calls : 3735
# Non-unit clause-clause subsumptions  : 70
# Unit Clause-clause subsumption calls : 121
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17691

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 45232 pages
