# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_compts_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v9_pre_topc(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((~(X3=k1_xboole_0)&(v4_pre_topc(X3,X1)&(r2_hidden(X2,k3_subset_1(u1_struct_0(X1),X3))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(v3_pre_topc(X5,X1)&(r2_hidden(X2,X4)&(r1_tarski(X3,X5)&r1_xboole_0(X4,X5))))))))))))))))), file('out/t21_compts_1__3_premises__rerun.E_input', d5_compts_1)).
fof(t15_compts_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v8_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_compts_1(X2,X1)=>(X2=k1_xboole_0|![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(v3_pre_topc(X5,X1)&(r2_hidden(X3,X4)&(r1_tarski(X2,X5)&r1_xboole_0(X4,X5))))))))))))))))), file('out/t21_compts_1__3_premises__rerun.E_input', t15_compts_1)).
fof(t21_compts_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>((v8_pre_topc(X1)&v1_compts_1(X1))=>v9_pre_topc(X1)))), file('out/t21_compts_1__3_premises__rerun.E_input', t21_compts_1)).
fof(t17_compts_1, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v1_compts_1(X1)&v4_pre_topc(X2,X1))=>v2_compts_1(X2,X1))))), file('out/t21_compts_1__3_premises__rerun.E_input', t17_compts_1)).
fof(c_0_4, plain, (![X6]:![X7]:![X8]:![X13]:![X14]:(((((((((m1_subset_1(esk1_3(X6,X7,X8),k1_zfmisc_1(u1_struct_0(X6)))|~r2_hidden(X7,k3_subset_1(u1_struct_0(X6),X8)))|~v4_pre_topc(X8,X6))|X8=k1_xboole_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,u1_struct_0(X6)))|~v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((((((((m1_subset_1(esk2_3(X6,X7,X8),k1_zfmisc_1(u1_struct_0(X6)))|~r2_hidden(X7,k3_subset_1(u1_struct_0(X6),X8)))|~v4_pre_topc(X8,X6))|X8=k1_xboole_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,u1_struct_0(X6)))|~v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((((((((v3_pre_topc(esk1_3(X6,X7,X8),X6)|~r2_hidden(X7,k3_subset_1(u1_struct_0(X6),X8)))|~v4_pre_topc(X8,X6))|X8=k1_xboole_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,u1_struct_0(X6)))|~v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((((((((v3_pre_topc(esk2_3(X6,X7,X8),X6)|~r2_hidden(X7,k3_subset_1(u1_struct_0(X6),X8)))|~v4_pre_topc(X8,X6))|X8=k1_xboole_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,u1_struct_0(X6)))|~v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((((((((r2_hidden(X7,esk1_3(X6,X7,X8))|~r2_hidden(X7,k3_subset_1(u1_struct_0(X6),X8)))|~v4_pre_topc(X8,X6))|X8=k1_xboole_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,u1_struct_0(X6)))|~v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((((((((r1_tarski(X8,esk2_3(X6,X7,X8))|~r2_hidden(X7,k3_subset_1(u1_struct_0(X6),X8)))|~v4_pre_topc(X8,X6))|X8=k1_xboole_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,u1_struct_0(X6)))|~v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(((((((r1_xboole_0(esk1_3(X6,X7,X8),esk2_3(X6,X7,X8))|~r2_hidden(X7,k3_subset_1(u1_struct_0(X6),X8)))|~v4_pre_topc(X8,X6))|X8=k1_xboole_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,u1_struct_0(X6)))|~v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))))))))&(((m1_subset_1(esk3_1(X6),u1_struct_0(X6))|v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(((m1_subset_1(esk4_1(X6),k1_zfmisc_1(u1_struct_0(X6)))|v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(((esk4_1(X6)!=k1_xboole_0|v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(((v4_pre_topc(esk4_1(X6),X6)|v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(((r2_hidden(esk3_1(X6),k3_subset_1(u1_struct_0(X6),esk4_1(X6)))|v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(((~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X6)))|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X6)))|(~v3_pre_topc(X13,X6)|(~v3_pre_topc(X14,X6)|(~r2_hidden(esk3_1(X6),X13)|(~r1_tarski(esk4_1(X6),X14)|~r1_xboole_0(X13,X14)))))))|v9_pre_topc(X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_compts_1])])])])])])])])).
fof(c_0_5, plain, (![X6]:![X7]:![X8]:((((((((m1_subset_1(esk5_3(X6,X7,X8),k1_zfmisc_1(u1_struct_0(X6)))|~r2_hidden(X8,k3_subset_1(u1_struct_0(X6),X7)))|~m1_subset_1(X8,u1_struct_0(X6)))|X7=k1_xboole_0)|~v2_compts_1(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~v8_pre_topc(X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&((((((((m1_subset_1(esk6_3(X6,X7,X8),k1_zfmisc_1(u1_struct_0(X6)))|~r2_hidden(X8,k3_subset_1(u1_struct_0(X6),X7)))|~m1_subset_1(X8,u1_struct_0(X6)))|X7=k1_xboole_0)|~v2_compts_1(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~v8_pre_topc(X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&((((((((v3_pre_topc(esk5_3(X6,X7,X8),X6)|~r2_hidden(X8,k3_subset_1(u1_struct_0(X6),X7)))|~m1_subset_1(X8,u1_struct_0(X6)))|X7=k1_xboole_0)|~v2_compts_1(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~v8_pre_topc(X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&((((((((v3_pre_topc(esk6_3(X6,X7,X8),X6)|~r2_hidden(X8,k3_subset_1(u1_struct_0(X6),X7)))|~m1_subset_1(X8,u1_struct_0(X6)))|X7=k1_xboole_0)|~v2_compts_1(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~v8_pre_topc(X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&((((((((r2_hidden(X8,esk5_3(X6,X7,X8))|~r2_hidden(X8,k3_subset_1(u1_struct_0(X6),X7)))|~m1_subset_1(X8,u1_struct_0(X6)))|X7=k1_xboole_0)|~v2_compts_1(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~v8_pre_topc(X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&((((((((r1_tarski(X7,esk6_3(X6,X7,X8))|~r2_hidden(X8,k3_subset_1(u1_struct_0(X6),X7)))|~m1_subset_1(X8,u1_struct_0(X6)))|X7=k1_xboole_0)|~v2_compts_1(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~v8_pre_topc(X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&(((((((r1_xboole_0(esk5_3(X6,X7,X8),esk6_3(X6,X7,X8))|~r2_hidden(X8,k3_subset_1(u1_struct_0(X6),X7)))|~m1_subset_1(X8,u1_struct_0(X6)))|X7=k1_xboole_0)|~v2_compts_1(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~v8_pre_topc(X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_compts_1])])])])])])])).
cnf(c_0_6,plain,(v2_struct_0(X1)|v9_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r1_xboole_0(X2,X3)|~r1_tarski(esk4_1(X1),X3)|~r2_hidden(esk3_1(X1),X2)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(X2=k1_xboole_0|r1_tarski(X2,esk6_3(X1,X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(v2_struct_0(X1)|v9_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|esk4_1(X1)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,plain,(v9_pre_topc(X1)|v2_struct_0(X1)|~v2_compts_1(esk4_1(X1),X2)|~v8_pre_topc(X2)|~r1_xboole_0(X3,esk6_3(X2,esk4_1(X1),X4))|~v3_pre_topc(esk6_3(X2,esk4_1(X1),X4),X1)|~v3_pre_topc(X3,X1)|~r2_hidden(X4,k3_subset_1(u1_struct_0(X2),esk4_1(X1)))|~r2_hidden(esk3_1(X1),X3)|~m1_subset_1(esk6_3(X2,esk4_1(X1),X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])).
cnf(c_0_10,plain,(X2=k1_xboole_0|r1_xboole_0(esk5_3(X1,X2,X3),esk6_3(X1,X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,plain,(v9_pre_topc(X1)|v2_struct_0(X1)|~v2_compts_1(esk4_1(X1),X2)|~v8_pre_topc(X2)|~v3_pre_topc(esk6_3(X2,esk4_1(X1),X3),X1)|~v3_pre_topc(esk5_3(X2,esk4_1(X1),X3),X1)|~r2_hidden(esk3_1(X1),esk5_3(X2,esk4_1(X1),X3))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X2),esk4_1(X1)))|~m1_subset_1(esk6_3(X2,esk4_1(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk5_3(X2,esk4_1(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_8])).
cnf(c_0_12,plain,(X2=k1_xboole_0|r2_hidden(X3,esk5_3(X1,X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13,plain,(v9_pre_topc(X1)|v2_struct_0(X1)|~v2_compts_1(esk4_1(X1),X2)|~v8_pre_topc(X2)|~v3_pre_topc(esk6_3(X2,esk4_1(X1),esk3_1(X1)),X1)|~v3_pre_topc(esk5_3(X2,esk4_1(X1),esk3_1(X1)),X1)|~r2_hidden(esk3_1(X1),k3_subset_1(u1_struct_0(X2),esk4_1(X1)))|~m1_subset_1(esk6_3(X2,esk4_1(X1),esk3_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk5_3(X2,esk4_1(X1),esk3_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk3_1(X1),u1_struct_0(X2))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_8])).
cnf(c_0_14,plain,(X2=k1_xboole_0|v3_pre_topc(esk6_3(X1,X2,X3),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v9_pre_topc(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,plain,(v2_struct_0(X1)|v9_pre_topc(X1)|m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,plain,(v2_struct_0(X1)|v9_pre_topc(X1)|r2_hidden(esk3_1(X1),k3_subset_1(u1_struct_0(X1),esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18,plain,(v9_pre_topc(X1)|v2_struct_0(X1)|~v2_compts_1(esk4_1(X1),X1)|~v8_pre_topc(X1)|~v3_pre_topc(esk5_3(X1,esk4_1(X1),esk3_1(X1)),X1)|~m1_subset_1(esk6_3(X1,esk4_1(X1),esk3_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk5_3(X1,esk4_1(X1),esk3_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_8])).
cnf(c_0_19,plain,(X2=k1_xboole_0|v3_pre_topc(esk5_3(X1,X2,X3),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_20, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>((v8_pre_topc(X1)&v1_compts_1(X1))=>v9_pre_topc(X1))))), inference(assume_negation,[status(cth)],[t21_compts_1])).
cnf(c_0_21,plain,(v9_pre_topc(X1)|v2_struct_0(X1)|~v2_compts_1(esk4_1(X1),X1)|~v8_pre_topc(X1)|~m1_subset_1(esk6_3(X1,esk4_1(X1),esk3_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk5_3(X1,esk4_1(X1),esk3_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_15]), c_0_16]), c_0_17]), c_0_8])).
cnf(c_0_22,plain,(X2=k1_xboole_0|m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_23, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|((~v1_compts_1(X3)|~v4_pre_topc(X4,X3))|v2_compts_1(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_compts_1])])])])])).
fof(c_0_24, negated_conjecture, (((~v2_struct_0(esk7_0)&(v2_pre_topc(esk7_0)&l1_pre_topc(esk7_0)))&((v8_pre_topc(esk7_0)&v1_compts_1(esk7_0))&~v9_pre_topc(esk7_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
cnf(c_0_25,plain,(v9_pre_topc(X1)|v2_struct_0(X1)|~v2_compts_1(esk4_1(X1),X1)|~v8_pre_topc(X1)|~m1_subset_1(esk5_3(X1,esk4_1(X1),esk3_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_15]), c_0_16]), c_0_17]), c_0_8])).
cnf(c_0_26,plain,(X2=k1_xboole_0|m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k3_subset_1(u1_struct_0(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_27,plain,(v2_compts_1(X1,X2)|~v4_pre_topc(X1,X2)|~v1_compts_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,negated_conjecture,(v1_compts_1(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29,negated_conjecture,(l1_pre_topc(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30,plain,(v9_pre_topc(X1)|v2_struct_0(X1)|~v2_compts_1(esk4_1(X1),X1)|~v8_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_15]), c_0_16]), c_0_17]), c_0_8])).
cnf(c_0_31,negated_conjecture,(v2_compts_1(X1,esk7_0)|~v4_pre_topc(X1,esk7_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk7_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])])).
cnf(c_0_32,negated_conjecture,(v8_pre_topc(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,negated_conjecture,(v2_pre_topc(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34,negated_conjecture,(~v9_pre_topc(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35,negated_conjecture,(~v2_struct_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36,negated_conjecture,(~v4_pre_topc(esk4_1(esk7_0),esk7_0)|~m1_subset_1(esk4_1(esk7_0),k1_zfmisc_1(u1_struct_0(esk7_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_29]), c_0_33])]), c_0_34]), c_0_35])).
cnf(c_0_37,plain,(v2_struct_0(X1)|v9_pre_topc(X1)|v4_pre_topc(esk4_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_38,negated_conjecture,(~m1_subset_1(esk4_1(esk7_0),k1_zfmisc_1(u1_struct_0(esk7_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_29]), c_0_33])]), c_0_34]), c_0_35])).
cnf(c_0_39,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_16]), c_0_29]), c_0_33])]), c_0_34]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 31
# Proof object formula steps           : 9
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 66
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 66
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 14
# ...of the previous two non-trivial   : 13
# Contextual simplify-reflections      : 21
# Paramodulations                      : 14
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 36
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 29
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 17
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 1387
# Rec. Clause-clause subsumption calls : 24
# Non-unit clause-clause subsumptions  : 24
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5429

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43040 pages
