# No SInE strategy applied
# Auto-Mode selected heuristic H_____102_C18_F1_PI_AE_CS_SP_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3))))))), file('out/t21_waybel_0__5_premises__rerun.E_input', t21_waybel_0)).
fof(dt_k5_waybel_0, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t21_waybel_0__5_premises__rerun.E_input', dt_k5_waybel_0)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t21_waybel_0__5_premises__rerun.E_input', t4_subset)).
fof(d2_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>(v4_orders_2(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X4))=>r1_orders_2(X1,X2,X4)))))))), file('out/t21_waybel_0__5_premises__rerun.E_input', d2_yellow_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t21_waybel_0__5_premises__rerun.E_input', d3_tarski)).
fof(t17_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), file('out/t21_waybel_0__5_premises__rerun.E_input', t17_waybel_0)).
fof(c_0_6, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3)))))))), inference(assume_negation,[status(cth)],[t21_waybel_0])).
fof(c_0_7, plain, (![X3]:![X4]:(((v2_struct_0(X3)|~l1_orders_2(X3))|~m1_subset_1(X4,u1_struct_0(X3)))|m1_subset_1(k5_waybel_0(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
fof(c_0_8, negated_conjecture, (((~v2_struct_0(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(r1_orders_2(esk1_0,esk2_0,esk3_0)&~r1_tarski(k5_waybel_0(esk1_0,esk2_0),k5_waybel_0(esk1_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_10,plain,(m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X8]:(((~v4_orders_2(X5)|(~m1_subset_1(X6,u1_struct_0(X5))|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|((~r1_orders_2(X5,X6,X7)|~r1_orders_2(X5,X7,X8))|r1_orders_2(X5,X6,X8))))))|~l1_orders_2(X5))&(((m1_subset_1(esk4_1(X5),u1_struct_0(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&(((m1_subset_1(esk5_1(X5),u1_struct_0(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&(((m1_subset_1(esk6_1(X5),u1_struct_0(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&((((r1_orders_2(X5,esk4_1(X5),esk5_1(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&((r1_orders_2(X5,esk5_1(X5),esk6_1(X5))|v4_orders_2(X5))|~l1_orders_2(X5)))&((~r1_orders_2(X5,esk4_1(X5),esk6_1(X5))|v4_orders_2(X5))|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_yellow_0])])])])])])])).
cnf(c_0_15,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])]), c_0_13])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk7_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk7_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_18,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(((((~r2_hidden(X6,k5_waybel_0(X4,X5))|r1_orders_2(X4,X6,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((~r1_orders_2(X4,X6,X5)|r2_hidden(X6,k5_waybel_0(X4,X5)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])])])).
cnf(c_0_21,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22,plain,(r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,negated_conjecture,(r1_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,esk2_0,X2)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_11]), c_0_12]), c_0_19])])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_26,plain,(v2_struct_0(X1)|r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k5_waybel_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|m1_subset_1(esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28,plain,(v2_struct_0(X1)|r2_hidden(X3,k5_waybel_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,negated_conjecture,(r1_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])])).
cnf(c_0_30,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),X2)|~r2_hidden(esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_12])]), c_0_13])).
cnf(c_0_31,negated_conjecture,(r2_hidden(esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,X2))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|~r1_orders_2(esk1_0,esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_27]), c_0_12])]), c_0_13])).
cnf(c_0_32,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),esk3_0)|~r1_orders_2(esk1_0,esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_29, c_0_27])).
cnf(c_0_33,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_11]), c_0_22])).
cnf(c_0_34,negated_conjecture,(r2_hidden(esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,esk3_0))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|~r1_orders_2(esk1_0,esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),esk3_0)), inference(spm,[status(thm)],[c_0_31, c_0_24])).
cnf(c_0_35,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,plain,(r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37,negated_conjecture,(r2_hidden(esk7_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,esk3_0))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38,negated_conjecture,(~r1_tarski(k5_waybel_0(esk1_0,esk2_0),k5_waybel_0(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 27
# Proof object formula steps           : 13
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 131
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 124
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 260
# ...of the previous two non-trivial   : 157
# Contextual simplify-reflections      : 15
# Paramodulations                      : 260
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 99
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 88
# Current number of unprocessed clauses: 68
# ...number of literals in the above   : 364
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 1446
# Rec. Clause-clause subsumption calls : 361
# Non-unit clause-clause subsumptions  : 26
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10224

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45216 pages
