# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t21_wellord1__50_premises__rerun__rerun.E_input', d4_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t21_wellord1__50_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(d6_wellord1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2)))), file('out/t21_wellord1__50_premises__rerun__rerun.E_input', d6_wellord1)).
fof(t21_wellord1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>k2_wellord1(k2_wellord1(X2,X1),X1)=k2_wellord1(X2,X1))), file('out/t21_wellord1__50_premises__rerun__rerun.E_input', t21_wellord1)).
fof(dt_k2_wellord1, axiom, (![X1]:![X2]:(v1_relat_1(X1)=>v1_relat_1(k2_wellord1(X1,X2)))), file('out/t21_wellord1__50_premises__rerun__rerun.E_input', dt_k2_wellord1)).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|~r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
cnf(c_0_6,plain,(r2_hidden(X4,X3)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_8,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|r2_hidden(esk1_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X3,X4)|r2_hidden(esk1_3(X3,X4,k3_xboole_0(X1,X2)),X4)|r2_hidden(esk1_3(X3,X4,k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_10,plain,(X1=k3_xboole_0(X2,X3)|~r2_hidden(esk1_3(X2,X3,X1),X3)|~r2_hidden(esk1_3(X2,X3,X1),X2)|~r2_hidden(esk1_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X3,X2)|r2_hidden(esk1_3(X3,X2,k3_xboole_0(X1,X2)),X2)), inference(ef,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|r2_hidden(esk1_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_13, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_14,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X3,X2)|~r2_hidden(esk1_3(X3,X2,k3_xboole_0(X1,X2)),k3_xboole_0(X1,X2))|~r2_hidden(esk1_3(X3,X2,k3_xboole_0(X1,X2)),X3)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15,plain,(k3_xboole_0(X1,X2)=X1|r2_hidden(esk1_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, (![X3]:![X4]:(~v1_relat_1(X3)|k2_wellord1(X3,X4)=k3_xboole_0(X3,k2_zfmisc_1(X4,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_wellord1])])])])])).
cnf(c_0_18,plain,(k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)|~r2_hidden(esk1_3(k3_xboole_0(X2,X1),X1,k3_xboole_0(X2,X1)),k3_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
fof(c_0_19, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>k2_wellord1(k2_wellord1(X2,X1),X1)=k2_wellord1(X2,X1)))), inference(assume_negation,[status(cth)],[t21_wellord1])).
cnf(c_0_20,plain,(k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21,plain,(k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_15]), c_0_16])])).
fof(c_0_22, plain, (![X3]:![X4]:(~v1_relat_1(X3)|v1_relat_1(k2_wellord1(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_wellord1])])])])).
fof(c_0_23, negated_conjecture, ((v1_relat_1(esk3_0)&k2_wellord1(k2_wellord1(esk3_0,esk2_0),esk2_0)!=k2_wellord1(esk3_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
cnf(c_0_24,plain,(k3_xboole_0(k2_zfmisc_1(X1,X1),X2)=k2_wellord1(X2,X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_16, c_0_20])).
cnf(c_0_25,plain,(k3_xboole_0(k2_zfmisc_1(X1,X1),k2_wellord1(X2,X1))=k2_wellord1(X2,X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_21, c_0_20])).
cnf(c_0_26,plain,(v1_relat_1(k2_wellord1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27,negated_conjecture,(k2_wellord1(k2_wellord1(esk3_0,esk2_0),esk2_0)!=k2_wellord1(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,plain,(k2_wellord1(k2_wellord1(X1,X2),X2)=k2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_29,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 20
# Proof object formula steps           : 11
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 11
# Processed clauses                    : 122
# ...of these trivial                  : 6
# ...subsumed                          : 63
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 566
# ...of the previous two non-trivial   : 515
# Contextual simplify-reflections      : 7
# Paramodulations                      : 539
# Factorizations                       : 18
# Equation resolutions                 : 9
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 45
# Current number of unprocessed clauses: 371
# ...number of literals in the above   : 1239
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 695
# Rec. Clause-clause subsumption calls : 556
# Non-unit clause-clause subsumptions  : 70
# Unit Clause-clause subsumption calls : 24
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7399

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 45272 pages
