# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_r3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3)))), file('out/t23_lattice3__8_premises__rerun.E_input', redefinition_r3_lattices)).
fof(t26_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)=>r3_lattices(X1,k7_lattices(X1,X3),k7_lattices(X1,X2))))))), file('out/t23_lattice3__8_premises__rerun.E_input', t26_lattices)).
fof(cc1_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1)))))))))), file('out/t23_lattice3__8_premises__rerun.E_input', cc1_lattices)).
fof(dt_k7_lattices, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l3_lattices(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k7_lattices(X1,X2),u1_struct_0(X1)))), file('out/t23_lattice3__8_premises__rerun.E_input', dt_k7_lattices)).
fof(rd7_lattices, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))&m1_subset_1(X2,u1_struct_0(X1)))=>k7_lattices(X1,k7_lattices(X1,X2))=X2)), file('out/t23_lattice3__8_premises__rerun.E_input', rd7_lattices)).
fof(t23_lattice3, conjecture, (![X1]:![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v17_lattices(X2)&l3_lattices(X2))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>(r4_lattice3(X2,X3,X1)<=>r3_lattice3(X2,k7_lattices(X2,X3),a_2_0_lattice3(X1,X2)))))), file('out/t23_lattice3__8_premises__rerun.E_input', t23_lattice3)).
fof(d16_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4))))))), file('out/t23_lattice3__8_premises__rerun.E_input', d16_lattice3)).
fof(fraenkel_a_2_0_lattice3, axiom, (![X1]:![X2]:![X3]:((~(v2_struct_0(X3))&(v10_lattices(X3)&(v17_lattices(X3)&l3_lattices(X3))))=>(r2_hidden(X1,a_2_0_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X3))&(X1=k7_lattices(X3,X4)&r2_hidden(X4,X2)))))), file('out/t23_lattice3__8_premises__rerun.E_input', fraenkel_a_2_0_lattice3)).
fof(d17_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2))))))), file('out/t23_lattice3__8_premises__rerun.E_input', d17_lattice3)).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:(((~r3_lattices(X4,X5,X6)|r1_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_lattices(X4,X5,X6)|r3_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|(~v10_lattices(X4)|(~v17_lattices(X4)|~l3_lattices(X4))))|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|(~r3_lattices(X4,X5,X6)|r3_lattices(X4,k7_lattices(X4,X6),k7_lattices(X4,X5))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_lattices])])])])])])).
fof(c_0_11, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v4_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v5_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v6_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v7_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v8_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&((v9_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_12, plain, (![X3]:![X4]:(((v2_struct_0(X3)|~l3_lattices(X3))|~m1_subset_1(X4,u1_struct_0(X3)))|m1_subset_1(k7_lattices(X3,X4),u1_struct_0(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_lattices])])])).
cnf(c_0_13,plain,(v2_struct_0(X2)|r1_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r3_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(r3_lattices(X1,k7_lattices(X1,X2),k7_lattices(X1,X3))|v2_struct_0(X1)|~r3_lattices(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v17_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v6_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,plain,(m1_subset_1(k7_lattices(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_19, plain, (![X3]:![X4]:(((v2_struct_0(X3)|(~v10_lattices(X3)|(~v17_lattices(X3)|~l3_lattices(X3))))|~m1_subset_1(X4,u1_struct_0(X3)))|k7_lattices(X3,k7_lattices(X3,X4))=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rd7_lattices])])])).
fof(c_0_20, negated_conjecture, (~(![X1]:![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v17_lattices(X2)&l3_lattices(X2))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>(r4_lattice3(X2,X3,X1)<=>r3_lattice3(X2,k7_lattices(X2,X3),a_2_0_lattice3(X1,X2))))))), inference(assume_negation,[status(cth)],[t23_lattice3])).
cnf(c_0_21,plain,(r1_lattices(X1,k7_lattices(X1,X2),k7_lattices(X1,X3))|v2_struct_0(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~r3_lattices(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_18])).
cnf(c_0_22,plain,(k7_lattices(X1,k7_lattices(X1,X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v17_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:((((~r3_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|r1_lattices(X5,X6,X8))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((r2_hidden(esk4_3(X5,X6,X7),X7)|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((~r1_lattices(X5,X6,esk4_3(X5,X6,X7))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])])).
fof(c_0_24, negated_conjecture, (((~v2_struct_0(esk2_0)&(v10_lattices(esk2_0)&(v17_lattices(esk2_0)&l3_lattices(esk2_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&((~r4_lattice3(esk2_0,esk3_0,esk1_0)|~r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0)))&(r4_lattice3(esk2_0,esk3_0,esk1_0)|r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])])])).
cnf(c_0_25,plain,(r1_lattices(X1,k7_lattices(X1,X2),X3)|v2_struct_0(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~r3_lattices(X1,k7_lattices(X1,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_18])).
cnf(c_0_26,plain,(v2_struct_0(X2)|r3_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r1_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27,plain,(v2_struct_0(X1)|r1_lattices(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r3_lattice3(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,negated_conjecture,(r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0))|r4_lattice3(esk2_0,esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29,negated_conjecture,(l3_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(r1_lattices(X1,k7_lattices(X1,X2),X3)|v2_struct_0(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~r1_lattices(X1,k7_lattices(X1,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_15]), c_0_16]), c_0_17]), c_0_18])).
cnf(c_0_32,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|r1_lattices(esk2_0,k7_lattices(esk2_0,esk3_0),X1)|~r2_hidden(X1,a_2_0_lattice3(esk1_0,esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_33,negated_conjecture,(v17_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34,negated_conjecture,(v10_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_36, plain, (![X5]:![X6]:![X7]:![X9]:((((m1_subset_1(esk5_3(X5,X6,X7),u1_struct_0(X7))|~r2_hidden(X5,a_2_0_lattice3(X6,X7)))|(v2_struct_0(X7)|(~v10_lattices(X7)|(~v17_lattices(X7)|~l3_lattices(X7)))))&(((X5=k7_lattices(X7,esk5_3(X5,X6,X7))|~r2_hidden(X5,a_2_0_lattice3(X6,X7)))|(v2_struct_0(X7)|(~v10_lattices(X7)|(~v17_lattices(X7)|~l3_lattices(X7)))))&((r2_hidden(esk5_3(X5,X6,X7),X6)|~r2_hidden(X5,a_2_0_lattice3(X6,X7)))|(v2_struct_0(X7)|(~v10_lattices(X7)|(~v17_lattices(X7)|~l3_lattices(X7)))))))&(((~m1_subset_1(X9,u1_struct_0(X7))|(X5!=k7_lattices(X7,X9)|~r2_hidden(X9,X6)))|r2_hidden(X5,a_2_0_lattice3(X6,X7)))|(v2_struct_0(X7)|(~v10_lattices(X7)|(~v17_lattices(X7)|~l3_lattices(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_lattice3])])])])])])])])).
cnf(c_0_37,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|r1_lattices(esk2_0,k7_lattices(esk2_0,X1),esk3_0)|~r2_hidden(X1,a_2_0_lattice3(esk1_0,esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_29])]), c_0_30])).
cnf(c_0_38,plain,(v2_struct_0(X1)|r2_hidden(X2,a_2_0_lattice3(X3,X1))|~l3_lattices(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~r2_hidden(X4,X3)|X2!=k7_lattices(X1,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_39, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:((((~r4_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|r1_lattices(X5,X8,X6))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((m1_subset_1(esk6_3(X5,X6,X7),u1_struct_0(X5))|r4_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((r2_hidden(esk6_3(X5,X6,X7),X7)|r4_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((~r1_lattices(X5,esk6_3(X5,X6,X7),X6)|r4_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])])).
cnf(c_0_40,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|r1_lattices(esk2_0,X1,esk3_0)|~r2_hidden(k7_lattices(esk2_0,X1),a_2_0_lattice3(esk1_0,esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_22]), c_0_33]), c_0_34]), c_0_29])]), c_0_30])).
cnf(c_0_41,plain,(r2_hidden(k7_lattices(X1,X2),a_2_0_lattice3(X3,X1))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~v17_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_42,plain,(v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,esk6_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_43,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|r1_lattices(esk2_0,X1,esk3_0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_33]), c_0_34]), c_0_29])]), c_0_30])).
cnf(c_0_44,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|r4_lattice3(esk2_0,esk3_0,X1)|~r2_hidden(esk6_3(esk2_0,esk3_0,X1),esk1_0)|~m1_subset_1(k7_lattices(esk2_0,esk6_3(esk2_0,esk3_0,X1)),u1_struct_0(esk2_0))|~m1_subset_1(esk6_3(esk2_0,esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_35]), c_0_29])]), c_0_30])).
cnf(c_0_45,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|r4_lattice3(esk2_0,esk3_0,X1)|~r2_hidden(esk6_3(esk2_0,esk3_0,X1),esk1_0)|~m1_subset_1(esk6_3(esk2_0,esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_18]), c_0_29])]), c_0_30])).
cnf(c_0_46,plain,(v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|r2_hidden(esk6_3(X1,X2,X3),X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk1_0),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_35]), c_0_29])]), c_0_30])).
cnf(c_0_48,plain,(v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49,plain,(v2_struct_0(X1)|r1_lattices(X1,X3,X2)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r4_lattice3(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_35]), c_0_29])]), c_0_30])).
cnf(c_0_51,plain,(v2_struct_0(X1)|X2=k7_lattices(X1,esk5_3(X2,X3,X1))|~l3_lattices(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_0_lattice3(X3,X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_52,plain,(v2_struct_0(X1)|m1_subset_1(esk5_3(X2,X3,X1),u1_struct_0(X1))|~l3_lattices(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_0_lattice3(X3,X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53,negated_conjecture,(r1_lattices(esk2_0,X1,esk3_0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_35]), c_0_29])]), c_0_30])).
cnf(c_0_54,plain,(v2_struct_0(X1)|r2_hidden(esk5_3(X2,X3,X1),X3)|~l3_lattices(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_0_lattice3(X3,X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_55,plain,(esk5_3(X1,X2,X3)=k7_lattices(X3,X1)|v2_struct_0(X3)|~r2_hidden(X1,a_2_0_lattice3(X2,X3))|~v17_lattices(X3)|~v10_lattices(X3)|~l3_lattices(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_51]), c_0_52])).
cnf(c_0_56,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,X2,esk4_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_57,negated_conjecture,(r1_lattices(esk2_0,k7_lattices(esk2_0,esk3_0),X1)|~r2_hidden(k7_lattices(esk2_0,X1),esk1_0)|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_53]), c_0_33]), c_0_34]), c_0_35]), c_0_29])]), c_0_30])).
cnf(c_0_58,plain,(r2_hidden(k7_lattices(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_lattice3(X3,X1))|~v17_lattices(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|r2_hidden(esk4_3(X1,X2,X3),X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_60,negated_conjecture,(r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),X1)|~r2_hidden(k7_lattices(esk2_0,esk4_3(esk2_0,k7_lattices(esk2_0,esk3_0),X1)),esk1_0)|~m1_subset_1(k7_lattices(esk2_0,esk4_3(esk2_0,k7_lattices(esk2_0,esk3_0),X1)),u1_struct_0(esk2_0))|~m1_subset_1(esk4_3(esk2_0,k7_lattices(esk2_0,esk3_0),X1),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_29])]), c_0_30])).
cnf(c_0_61,plain,(r2_hidden(k7_lattices(X1,esk4_3(X2,X3,a_2_0_lattice3(X4,X1))),X4)|r3_lattice3(X2,X3,a_2_0_lattice3(X4,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v17_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X1)|~l3_lattices(X2)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62,negated_conjecture,(r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk4_3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0))),u1_struct_0(esk2_0))|~m1_subset_1(esk4_3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0)),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_33]), c_0_34]), c_0_29])]), c_0_30])).
cnf(c_0_63,negated_conjecture,(r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0))|~m1_subset_1(esk4_3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0)),u1_struct_0(esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_18]), c_0_29])]), c_0_30])).
cnf(c_0_64,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_65,negated_conjecture,(~r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0))|~r4_lattice3(esk2_0,esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_66,negated_conjecture,(r3_lattice3(esk2_0,k7_lattices(esk2_0,esk3_0),a_2_0_lattice3(esk1_0,esk2_0))|~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_29])]), c_0_30])).
cnf(c_0_67,negated_conjecture,(~m1_subset_1(k7_lattices(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_50])).
cnf(c_0_68,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_18]), c_0_35]), c_0_29])]), c_0_30]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 50
# Proof object formula steps           : 19
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 75
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 30
# Processed clauses                    : 225
# ...of these trivial                  : 0
# ...subsumed                          : 56
# ...remaining for further processing  : 169
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 0
# Generated clauses                    : 822
# ...of the previous two non-trivial   : 761
# Contextual simplify-reflections      : 126
# Paramodulations                      : 821
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 123
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 117
# Current number of unprocessed clauses: 563
# ...number of literals in the above   : 7882
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 16297
# Rec. Clause-clause subsumption calls : 931
# Non-unit clause-clause subsumptions  : 197
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 36722

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.000 s
# Total time               : 0.032 s
# Maximum resident set size: 45352 pages
