# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_yellow_0, conjecture, (![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k12_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))))))), file('out/t23_yellow_0__4_premises__rerun.E_input', t23_yellow_0)).
fof(l28_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))))))), file('out/t23_yellow_0__4_premises__rerun.E_input', l28_lattice3)).
fof(cc2_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t23_yellow_0__4_premises__rerun.E_input', cc2_lattice3)).
fof(redefinition_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3))), file('out/t23_yellow_0__4_premises__rerun.E_input', redefinition_k12_lattice3)).
fof(t15_lattice3, axiom, (![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2))))), file('out/t23_yellow_0__4_premises__rerun.E_input', t15_lattice3)).
fof(c_0_5, negated_conjecture, (~(![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k12_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))))))))), inference(assume_negation,[status(cth)],[t23_yellow_0])).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X9,X7)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X9,X8)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X10,X7)|~r1_orders_2(X6,X10,X8))|r1_orders_2(X6,X10,X9)))|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk6_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,esk6_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,esk6_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,esk6_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l28_lattice3])])])])])])])])).
fof(c_0_7, plain, (![X2]:(~l1_orders_2(X2)|(~v2_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_8, negated_conjecture, (![X11]:((v5_orders_2(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(((((m1_subset_1(esk5_0,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk4_0,esk3_0))|~r1_orders_2(esk1_0,esk4_0,esk2_0))|esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0))&(((((r1_orders_2(esk1_0,esk5_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk3_0))|~r1_orders_2(esk1_0,esk4_0,esk2_0))|esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0))&(((r1_orders_2(esk1_0,esk5_0,esk3_0)|~r1_orders_2(esk1_0,esk4_0,esk3_0))|~r1_orders_2(esk1_0,esk4_0,esk2_0))|esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)))&(((~r1_orders_2(esk1_0,esk5_0,esk4_0)|~r1_orders_2(esk1_0,esk4_0,esk3_0))|~r1_orders_2(esk1_0,esk4_0,esk2_0))|esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0))))&((r1_orders_2(esk1_0,esk4_0,esk2_0)|esk4_0=k12_lattice3(esk1_0,esk2_0,esk3_0))&((r1_orders_2(esk1_0,esk4_0,esk3_0)|esk4_0=k12_lattice3(esk1_0,esk2_0,esk3_0))&((~m1_subset_1(X11,u1_struct_0(esk1_0))|((~r1_orders_2(esk1_0,X11,esk2_0)|~r1_orders_2(esk1_0,X11,esk3_0))|r1_orders_2(esk1_0,X11,esk4_0)))|esk4_0=k12_lattice3(esk1_0,esk2_0,esk3_0)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k12_lattice3(X4,X5,X6)=k11_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_10,plain,(v2_struct_0(X1)|r1_orders_2(X1,X4,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X4!=k11_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(~v2_struct_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|k11_lattice3(X4,X5,X6)=k11_lattice3(X4,X6,X5))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_lattice3])])])])])).
cnf(c_0_13,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|r1_orders_2(X1,esk6_4(X1,X2,X3,X4),X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,negated_conjecture,(esk4_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,plain,(k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21,plain,(r1_orders_2(X1,X2,X3)|X2!=k11_lattice3(X1,X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_22,plain,(k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(esk4_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_24,negated_conjecture,(esk4_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,X1,esk4_0)|~r1_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25,plain,(X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk6_4(X2,X3,X4,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_26,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|r1_orders_2(X1,esk6_4(X1,X2,X3,X4),X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_27,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0|r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_28,plain,(r1_orders_2(X1,X2,X3)|X2!=k11_lattice3(X1,X3,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0|r1_orders_2(esk1_0,esk4_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_30,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,esk6_4(X1,X2,X3,X4),X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0|X1=k11_lattice3(esk1_0,X2,esk3_0)|r1_orders_2(esk1_0,esk6_4(esk1_0,X2,esk3_0,X1),esk4_0)|~r1_orders_2(esk1_0,esk6_4(esk1_0,X2,esk3_0,X1),esk2_0)|~r1_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(esk6_4(esk1_0,X2,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_16]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_32,plain,(X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk6_4(X2,X3,X4,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_26, c_0_11])).
cnf(c_0_33,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk3_0)|r1_orders_2(esk1_0,X1,esk3_0)|X1!=esk4_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_27]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_34,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk2_0)|r1_orders_2(esk1_0,X1,esk2_0)|X1!=esk4_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_17]), c_0_16]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_36,plain,(X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk6_4(X2,X3,X4,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_30, c_0_11])).
cnf(c_0_37,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0|X1=k11_lattice3(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk6_4(esk1_0,esk2_0,esk3_0,X1),esk4_0)|~r1_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(esk6_4(esk1_0,esk2_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_17]), c_0_16]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_38,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_34])).
cnf(c_0_40,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|m1_subset_1(esk6_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_41,plain,(v2_struct_0(X1)|r1_orders_2(X1,X5,X4)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X4!=k11_lattice3(X1,X2,X3)|~r1_orders_2(X1,X5,X3)|~r1_orders_2(X1,X5,X2)|~m1_subset_1(X5,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_42,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0|k11_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0|~m1_subset_1(esk6_4(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39]), c_0_34]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_43,plain,(X1=k11_lattice3(X2,X3,X4)|m1_subset_1(esk6_4(X2,X3,X4,X1),u1_struct_0(X2))|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_40, c_0_11])).
cnf(c_0_44,plain,(r1_orders_2(X1,X2,X3)|X3!=k11_lattice3(X1,X4,X5)|~r1_orders_2(X1,X2,X5)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[c_0_41, c_0_11])).
cnf(c_0_45,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0|k12_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_38]), c_0_39]), c_0_34]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_46,negated_conjecture,(r1_orders_2(esk1_0,esk5_0,esk3_0)|esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_47,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk1_0))|esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_48,negated_conjecture,(r1_orders_2(esk1_0,esk5_0,esk2_0)|esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_49,negated_conjecture,(esk4_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk3_0)|~r1_orders_2(esk1_0,esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_50,plain,(r1_orders_2(X1,X2,k11_lattice3(X1,X3,X4))|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(k11_lattice3(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_51,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,esk3_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_45]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_52,negated_conjecture,(r1_orders_2(esk1_0,esk5_0,esk3_0)|k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0|~r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_39])])).
cnf(c_0_53,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk1_0))|k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0|~r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_39])])).
cnf(c_0_54,negated_conjecture,(r1_orders_2(esk1_0,esk5_0,esk2_0)|k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0|~r1_orders_2(esk1_0,esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_39])])).
cnf(c_0_55,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0|~r1_orders_2(esk1_0,esk4_0,esk3_0)|~r1_orders_2(esk1_0,esk5_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_39])])).
cnf(c_0_56,negated_conjecture,(r1_orders_2(esk1_0,X1,esk4_0)|~r1_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_34]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_57,negated_conjecture,(r1_orders_2(esk1_0,esk5_0,esk3_0)|k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_38])])).
cnf(c_0_58,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk1_0))|k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_38])])).
cnf(c_0_59,negated_conjecture,(r1_orders_2(esk1_0,esk5_0,esk2_0)|k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_38])])).
cnf(c_0_60,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0|~r1_orders_2(esk1_0,esk5_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_38])])).
cnf(c_0_61,negated_conjecture,(k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_59]), c_0_60])).
cnf(c_0_62,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_15]), c_0_51]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 52
# Proof object formula steps           : 11
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 98
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 137
# ...of these trivial                  : 0
# ...subsumed                          : 29
# ...remaining for further processing  : 108
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 26
# Generated clauses                    : 171
# ...of the previous two non-trivial   : 135
# Contextual simplify-reflections      : 17
# Paramodulations                      : 162
# Factorizations                       : 0
# Equation resolutions                 : 9
# Current number of processed clauses  : 46
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 159
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 1355
# Rec. Clause-clause subsumption calls : 199
# Non-unit clause-clause subsumptions  : 57
# Unit Clause-clause subsumption calls : 20
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8456

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43032 pages
