# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2)))))))), file('out/t23_funct_1__6_premises__rerun.E_input', d8_relat_1)).
fof(dt_k5_relat_1, axiom, (![X1]:![X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2)))), file('out/t23_funct_1__6_premises__rerun.E_input', dt_k5_relat_1)).
fof(d4_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:![X3]:((r2_hidden(X2,k1_relat_1(X1))=>(X3=k1_funct_1(X1,X2)<=>r2_hidden(k4_tarski(X2,X3),X1)))&(~(r2_hidden(X2,k1_relat_1(X1)))=>(X3=k1_funct_1(X1,X2)<=>X3=k1_xboole_0))))), file('out/t23_funct_1__6_premises__rerun.E_input', d4_funct_1)).
fof(t8_funct_1, axiom, (![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1))))), file('out/t23_funct_1__6_premises__rerun.E_input', t8_funct_1)).
fof(fc2_funct_1, axiom, (![X1]:![X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&(v1_relat_1(X2)&v1_funct_1(X2)))=>(v1_relat_1(k5_relat_1(X1,X2))&v1_funct_1(k5_relat_1(X1,X2))))), file('out/t23_funct_1__6_premises__rerun.E_input', fc2_funct_1)).
fof(t23_funct_1, conjecture, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(X2))=>k1_funct_1(k5_relat_1(X2,X3),X1)=k1_funct_1(X3,k1_funct_1(X2,X1)))))), file('out/t23_funct_1__6_premises__rerun.E_input', t23_funct_1)).
fof(t22_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))=>k1_funct_1(k5_relat_1(X3,X2),X1)=k1_funct_1(X2,k1_funct_1(X3,X1)))))), file('out/t23_funct_1__6_premises__rerun.E_input', t22_funct_1)).
fof(c_0_7, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X10]:![X11]:![X13]:![X16]:((((((((r2_hidden(k4_tarski(X10,esk1_5(X7,X8,X9,X10,X11)),X7)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7))&(((((r2_hidden(k4_tarski(esk1_5(X7,X8,X9,X10,X11),X11),X8)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7)))&((((((~r2_hidden(k4_tarski(X10,X13),X7)|~r2_hidden(k4_tarski(X13,X11),X8))|r2_hidden(k4_tarski(X10,X11),X9))|X9!=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7)))&((((((~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9)|(~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),X16),X7)|~r2_hidden(k4_tarski(X16,esk3_3(X7,X8,X9)),X8)))|X9=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7))&((((((r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk4_3(X7,X8,X9)),X7)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9))|X9=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7))&(((((r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk3_3(X7,X8,X9)),X8)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9))|X9=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])])).
fof(c_0_8, plain, (![X3]:![X4]:((~v1_relat_1(X3)|~v1_relat_1(X4))|v1_relat_1(k5_relat_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:![X6]:![X5]:![X6]:![X6]:(((((X6!=k1_funct_1(X4,X5)|r2_hidden(k4_tarski(X5,X6),X4))|~r2_hidden(X5,k1_relat_1(X4)))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&(((~r2_hidden(k4_tarski(X5,X6),X4)|X6=k1_funct_1(X4,X5))|~r2_hidden(X5,k1_relat_1(X4)))|(~v1_relat_1(X4)|~v1_funct_1(X4))))&((((X6!=k1_funct_1(X4,X5)|X6=k1_xboole_0)|r2_hidden(X5,k1_relat_1(X4)))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&(((X6!=k1_xboole_0|X6=k1_funct_1(X4,X5))|r2_hidden(X5,k1_relat_1(X4)))|(~v1_relat_1(X4)|~v1_funct_1(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_funct_1])])])])])])])).
cnf(c_0_10,plain,(r2_hidden(k4_tarski(X4,X5),X3)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)|X3!=k5_relat_1(X1,X2)|~r2_hidden(k4_tarski(X6,X5),X2)|~r2_hidden(k4_tarski(X4,X6),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(r2_hidden(k4_tarski(X2,X3),X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|X3!=k1_funct_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:((((r2_hidden(X4,k1_relat_1(X6))|~r2_hidden(k4_tarski(X4,X5),X6))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((X5=k1_funct_1(X6,X4)|~r2_hidden(k4_tarski(X4,X5),X6))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&(((~r2_hidden(X4,k1_relat_1(X6))|X5!=k1_funct_1(X6,X4))|r2_hidden(k4_tarski(X4,X5),X6))|(~v1_relat_1(X6)|~v1_funct_1(X6))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_14,plain,(r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,X4))|~r2_hidden(k4_tarski(X5,X2),X4)|~r2_hidden(k4_tarski(X1,X5),X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])).
cnf(c_0_15,plain,(r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~v1_funct_1(X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(r2_hidden(k4_tarski(X1,k1_funct_1(X2,X3)),k5_relat_1(X4,X2))|~v1_funct_1(X2)|~r2_hidden(k4_tarski(X1,X3),X4)|~r2_hidden(X3,k1_relat_1(X2))|~v1_relat_1(X2)|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_18, plain, (![X3]:![X4]:((v1_relat_1(k5_relat_1(X3,X4))|((~v1_relat_1(X3)|~v1_funct_1(X3))|(~v1_relat_1(X4)|~v1_funct_1(X4))))&(v1_funct_1(k5_relat_1(X3,X4))|((~v1_relat_1(X3)|~v1_funct_1(X3))|(~v1_relat_1(X4)|~v1_funct_1(X4)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_funct_1])])])).
cnf(c_0_19,plain,(r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~v1_funct_1(k5_relat_1(X2,X3))|~v1_funct_1(X3)|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,k1_relat_1(X3))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_11])).
cnf(c_0_20,plain,(v1_funct_1(k5_relat_1(X2,X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21,plain,(r2_hidden(X2,k1_relat_1(X1))|X3=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)|X3!=k1_funct_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_22, negated_conjecture, (~(![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(X2))=>k1_funct_1(k5_relat_1(X2,X3),X1)=k1_funct_1(X3,k1_funct_1(X2,X1))))))), inference(assume_negation,[status(cth)],[t23_funct_1])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:((~v1_relat_1(X5)|~v1_funct_1(X5))|((~v1_relat_1(X6)|~v1_funct_1(X6))|(~r2_hidden(X4,k1_relat_1(k5_relat_1(X6,X5)))|k1_funct_1(k5_relat_1(X6,X5),X4)=k1_funct_1(X5,k1_funct_1(X6,X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_funct_1])])])])])).
cnf(c_0_24,plain,(r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_hidden(k1_funct_1(X2,X1),k1_relat_1(X3))|~r2_hidden(X1,k1_relat_1(X2))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_15]), c_0_20])).
cnf(c_0_25,plain,(k1_funct_1(X1,X2)=k1_xboole_0|r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_21])).
fof(c_0_26, negated_conjecture, (((v1_relat_1(esk6_0)&v1_funct_1(esk6_0))&((v1_relat_1(esk7_0)&v1_funct_1(esk7_0))&(r2_hidden(esk5_0,k1_relat_1(esk6_0))&k1_funct_1(k5_relat_1(esk6_0,esk7_0),esk5_0)!=k1_funct_1(esk7_0,k1_funct_1(esk6_0,esk5_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).
cnf(c_0_27,plain,(k1_funct_1(k5_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|~r2_hidden(X3,k1_relat_1(k5_relat_1(X1,X2)))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,plain,(k1_funct_1(X1,k1_funct_1(X2,X3))=k1_xboole_0|r2_hidden(X3,k1_relat_1(k5_relat_1(X2,X1)))|~v1_funct_1(X1)|~v1_funct_1(X2)|~r2_hidden(X3,k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,negated_conjecture,(k1_funct_1(k5_relat_1(esk6_0,esk7_0),esk5_0)!=k1_funct_1(esk7_0,k1_funct_1(esk6_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30,plain,(k1_funct_1(k5_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|k1_funct_1(X2,k1_funct_1(X1,X3))=k1_xboole_0|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_hidden(X3,k1_relat_1(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31,negated_conjecture,(v1_funct_1(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,negated_conjecture,(r2_hidden(esk5_0,k1_relat_1(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,negated_conjecture,(v1_relat_1(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35,negated_conjecture,(v1_relat_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36,negated_conjecture,(k1_funct_1(esk7_0,k1_funct_1(esk6_0,esk5_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])])).
cnf(c_0_37,negated_conjecture,(k1_funct_1(k5_relat_1(esk6_0,esk7_0),esk5_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_29, c_0_36])).
cnf(c_0_38,plain,(k1_funct_1(k5_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|k1_funct_1(k5_relat_1(X1,X2),X3)=k1_xboole_0|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_25]), c_0_11]), c_0_20])).
cnf(c_0_39,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_36]), c_0_31]), c_0_32]), c_0_34]), c_0_35])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 25
# Proof object formula steps           : 15
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 83
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 74
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 149
# ...of the previous two non-trivial   : 133
# Contextual simplify-reflections      : 30
# Paramodulations                      : 141
# Factorizations                       : 1
# Equation resolutions                 : 7
# Current number of processed clauses  : 47
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 92
# ...number of literals in the above   : 819
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 603
# Rec. Clause-clause subsumption calls : 126
# Non-unit clause-clause subsumptions  : 39
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7339

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 45476 pages
