# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_zfmisc_1, conjecture, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),k1_tarski(X2))=>X1=X2)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t24_zfmisc_1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t73_enumset1)).
fof(t22_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t22_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t28_xboole_1)).
fof(t46_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4))), file('out/t24_zfmisc_1__10_premises__rerun.E_input', t46_enumset1)).
fof(d1_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3))))))), file('out/t24_zfmisc_1__10_premises__rerun.E_input', d1_enumset1)).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:(r1_tarski(k1_tarski(X1),k1_tarski(X2))=>X1=X2))), inference(assume_negation,[status(cth)],[t24_zfmisc_1])).
fof(c_0_12, negated_conjecture, ((r1_tarski(k1_tarski(esk1_0),k1_tarski(esk2_0))&esk1_0!=esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_13, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_14, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_16, plain, (![X5]:![X6]:![X7]:![X8]:k3_enumset1(X5,X5,X6,X7,X8)=k2_enumset1(X5,X6,X7,X8)), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_17, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k4_enumset1(X6,X6,X7,X8,X9,X10)=k3_enumset1(X6,X7,X8,X9,X10)), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_18, plain, (![X3]:![X4]:k2_xboole_0(X3,k3_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t22_xboole_1])).
fof(c_0_19, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_20, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_21,negated_conjecture,(r1_tarski(k1_tarski(esk1_0),k1_tarski(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26,plain,(k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_27, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X5,X6,X7),k1_tarski(X8))), inference(variable_rename,[status(thm)],[t46_enumset1])).
cnf(c_0_28,plain,(k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31,negated_conjecture,(r1_tarski(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_22]), c_0_23]), c_0_23]), c_0_24]), c_0_24]), c_0_25]), c_0_25]), c_0_26]), c_0_26])).
cnf(c_0_32,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_33, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X10]:![X6]:![X7]:![X8]:![X9]:((((~r2_hidden(X10,X9)|(X10=X6|(X10=X7|X10=X8)))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X6|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X7|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&((X10!=X8|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8)))))&((((esk3_4(X6,X7,X8,X9)!=X6|~r2_hidden(esk3_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&(((esk3_4(X6,X7,X8,X9)!=X7|~r2_hidden(esk3_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&((esk3_4(X6,X7,X8,X9)!=X8|~r2_hidden(esk3_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))))&((r2_hidden(esk3_4(X6,X7,X8,X9),X9)|(esk3_4(X6,X7,X8,X9)=X6|(esk3_4(X6,X7,X8,X9)=X7|esk3_4(X6,X7,X8,X9)=X8)))|X9=k1_enumset1(X6,X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])])).
cnf(c_0_34,plain,(k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35,negated_conjecture,(k3_xboole_0(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36,plain,(k4_enumset1(X1,X1,X1,X2,X3,X4)=k2_xboole_0(k4_enumset1(X1,X1,X1,X1,X2,X3),k4_enumset1(X4,X4,X4,X4,X4,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_22]), c_0_23]), c_0_24]), c_0_24]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_26])).
cnf(c_0_37,plain,(X5=X4|X5=X3|X5=X2|X1!=k1_enumset1(X2,X3,X4)|~r2_hidden(X5,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_38,negated_conjecture,(k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_39,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X3), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40,plain,(X4=X5|X3=X5|X2=X5|X1!=k4_enumset1(X2,X2,X2,X2,X3,X4)|~r2_hidden(X5,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_41,negated_conjecture,(k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1)=k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk1_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_38]), c_0_36])).
cnf(c_0_42,plain,(r2_hidden(X5,X1)|X3!=X5|X1!=k4_enumset1(X2,X2,X2,X2,X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_43,plain,(X1=X2|X3=X2|X4=X2|~r2_hidden(X2,k4_enumset1(X4,X4,X4,X4,X3,X1))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_44,negated_conjecture,(k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,X1,X2)=k4_enumset1(esk2_0,esk2_0,esk2_0,esk1_0,X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_41]), c_0_36])).
cnf(c_0_45,plain,(r2_hidden(X1,X2)|X2!=k4_enumset1(X3,X3,X3,X3,X1,X4)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_46,negated_conjecture,(esk2_0=X1|X2=X1|X3=X1|~r2_hidden(X1,k4_enumset1(esk2_0,esk2_0,esk2_0,esk1_0,X2,X3))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47,negated_conjecture,(k4_enumset1(esk2_0,esk2_0,esk2_0,esk1_0,esk2_0,X1)=k4_enumset1(esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_44]), c_0_44])).
cnf(c_0_48,plain,(r2_hidden(X1,k4_enumset1(X2,X2,X2,X2,X1,X3))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_49,negated_conjecture,(esk2_0=X1|X2=X1|~r2_hidden(X1,k4_enumset1(esk2_0,esk2_0,esk2_0,esk1_0,esk1_0,X2))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_50,negated_conjecture,(r2_hidden(X1,k4_enumset1(esk2_0,esk2_0,esk2_0,esk1_0,X1,X2))), inference(spm,[status(thm)],[c_0_48, c_0_44])).
cnf(c_0_51,negated_conjecture,(esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_52,negated_conjecture,(X1=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
cnf(c_0_53,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 31
# Proof object formula steps           : 23
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 14
# Processed clauses                    : 49
# ...of these trivial                  : 1
# ...subsumed                          : 6
# ...remaining for further processing  : 42
# Other redundant clauses eliminated   : 48
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 38
# Generated clauses                    : 185
# ...of the previous two non-trivial   : 136
# Contextual simplify-reflections      : 0
# Paramodulations                      : 100
# Factorizations                       : 30
# Equation resolutions                 : 55
# Current number of processed clauses  : 1
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 32
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 60
# Rec. Clause-clause subsumption calls : 58
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 2
# BW rewrite match attempts            : 65
# BW rewrite match successes           : 46
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3574

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 47680 pages
