# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t45_setfam_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(m1_setfam_1(X2,X1)<=>k5_setfam_1(X1,X2)=X1))), file('out/t24_yellow_1__9_premises__rerun.E_input', t45_setfam_1)).
fof(dt_u1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t24_yellow_1__9_premises__rerun.E_input', dt_u1_pre_topc)).
fof(d11_setfam_1, axiom, (![X1]:![X2]:(m1_setfam_1(X2,X1)<=>r1_tarski(X1,k3_tarski(X2)))), file('out/t24_yellow_1__9_premises__rerun.E_input', d11_setfam_1)).
fof(t74_zfmisc_1, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2)))), file('out/t24_yellow_1__9_premises__rerun.E_input', t74_zfmisc_1)).
fof(t24_yellow_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), file('out/t24_yellow_1__9_premises__rerun.E_input', t24_yellow_1)).
fof(t14_yellow_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)))), file('out/t24_yellow_1__9_premises__rerun.E_input', t14_yellow_1)).
fof(d1_yellow_1, axiom, (![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), file('out/t24_yellow_1__9_premises__rerun.E_input', d1_yellow_1)).
fof(redefinition_k5_setfam_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2))), file('out/t24_yellow_1__9_premises__rerun.E_input', redefinition_k5_setfam_1)).
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t24_yellow_1__9_premises__rerun.E_input', t7_boole)).
fof(d1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1)))))))))), file('out/t24_yellow_1__9_premises__rerun.E_input', d1_pre_topc)).
fof(c_0_10, plain, (![X3]:![X4]:(((~m1_setfam_1(X4,X3)|k5_setfam_1(X3,X4)=X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))&((k5_setfam_1(X3,X4)!=X3|m1_setfam_1(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_setfam_1])])])).
fof(c_0_11, plain, (![X2]:(~l1_pre_topc(X2)|m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
fof(c_0_12, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_setfam_1(X4,X3)|r1_tarski(X3,k3_tarski(X4)))&(~r1_tarski(X3,k3_tarski(X4))|m1_setfam_1(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_setfam_1])])])])).
fof(c_0_13, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|r1_tarski(X3,k3_tarski(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t74_zfmisc_1])])).
fof(c_0_14, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
fof(c_0_15, plain, (![X2]:(v1_xboole_0(X2)|(~r2_hidden(k3_tarski(X2),X2)|k4_yellow_0(k2_yellow_1(X2))=k3_tarski(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
fof(c_0_16, plain, (![X2]:k2_yellow_1(X2)=g1_orders_2(X2,k1_yellow_1(X2))), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_17, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|k5_setfam_1(X3,X4)=k3_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
cnf(c_0_18,plain,(k5_setfam_1(X2,X1)=X2|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,plain,(m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,plain,(m1_setfam_1(X1,X2)|~r1_tarski(X2,k3_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,plain,(r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_22, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
cnf(c_0_23,plain,(k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|v1_xboole_0(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24,plain,(k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_25, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_26,plain,(k5_setfam_1(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,plain,(k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=u1_struct_0(X1)|~l1_pre_topc(X1)|~m1_setfam_1(u1_pre_topc(X1),u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_28,plain,(m1_setfam_1(X1,X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29,negated_conjecture,(k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,plain,(k4_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))=k3_tarski(X1)|v1_xboole_0(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,plain,(k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=k3_tarski(u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_26, c_0_19])).
cnf(c_0_33,plain,(k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=u1_struct_0(X1)|~r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:![X7]:((((r2_hidden(u1_struct_0(X4),u1_pre_topc(X4))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&((((~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|(~r1_tarski(X5,u1_pre_topc(X4))|r2_hidden(k5_setfam_1(u1_struct_0(X4),X5),u1_pre_topc(X4))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&(((~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|((~r2_hidden(X6,u1_pre_topc(X4))|~r2_hidden(X7,u1_pre_topc(X4)))|r2_hidden(k9_subset_1(u1_struct_0(X4),X6,X7),u1_pre_topc(X4)))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))))&((((((m1_subset_1(esk3_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk2_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk4_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk2_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk3_1(X4),u1_pre_topc(X4))|m1_subset_1(esk2_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk4_1(X4),u1_pre_topc(X4))|m1_subset_1(esk2_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk3_1(X4),esk4_1(X4)),u1_pre_topc(X4))|m1_subset_1(esk2_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&((((((m1_subset_1(esk3_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk2_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk4_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk2_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk3_1(X4),u1_pre_topc(X4))|r1_tarski(esk2_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk4_1(X4),u1_pre_topc(X4))|r1_tarski(esk2_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk3_1(X4),esk4_1(X4)),u1_pre_topc(X4))|r1_tarski(esk2_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&(((((m1_subset_1(esk3_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk2_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk4_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk2_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk3_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk2_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk4_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk2_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk3_1(X4),esk4_1(X4)),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk2_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])])])).
cnf(c_0_35,negated_conjecture,(k4_yellow_0(g1_orders_2(u1_pre_topc(esk1_0),k1_yellow_1(u1_pre_topc(esk1_0))))!=u1_struct_0(esk1_0)), inference(rw,[status(thm)],[c_0_29, c_0_24])).
cnf(c_0_36,plain,(k4_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37,plain,(k3_tarski(u1_pre_topc(X1))=u1_struct_0(X1)|~r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38,plain,(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39,negated_conjecture,(k3_tarski(u1_pre_topc(esk1_0))!=u1_struct_0(esk1_0)|~r2_hidden(k3_tarski(u1_pre_topc(esk1_0)),u1_pre_topc(esk1_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40,plain,(k3_tarski(u1_pre_topc(X1))=u1_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_41,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_42,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_43,negated_conjecture,(~r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_44,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_38]), c_0_41]), c_0_42])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 24
# Proof object formula steps           : 21
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 31
# Processed clauses                    : 60
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 60
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 16
# ...of the previous two non-trivial   : 13
# Contextual simplify-reflections      : 1
# Paramodulations                      : 16
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 74
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 747
# Rec. Clause-clause subsumption calls : 63
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3240

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 43328 pages
