# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(X3=k3_relat_1(X1,X2)<=>![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2)))))), file('out/t25_relat_1__4_premises__rerun.E_input', d8_relat_1)).
fof(dt_k3_relat_1, axiom, (![X1]:![X2]:v1_relat_1(k3_relat_1(X1,X2))), file('out/t25_relat_1__4_premises__rerun.E_input', dt_k3_relat_1)).
fof(d12_xtuple_0, axiom, (![X1]:![X2]:(X2=k9_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t25_relat_1__4_premises__rerun.E_input', d12_xtuple_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t25_relat_1__4_premises__rerun.E_input', d3_tarski)).
fof(t25_relat_1, conjecture, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k9_xtuple_0(X1))))), file('out/t25_relat_1__4_premises__rerun.E_input', t25_relat_1)).
fof(c_0_5, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X10]:![X11]:![X13]:![X16]:((((((r2_hidden(k4_tarski(X10,esk1_5(X7,X8,X9,X10,X11)),X7)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9))&(((r2_hidden(k4_tarski(esk1_5(X7,X8,X9,X10,X11),X11),X8)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9)))&((((~r2_hidden(k4_tarski(X10,X13),X7)|~r2_hidden(k4_tarski(X13,X11),X8))|r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9)))&((((~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9)|(~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),X16),X7)|~r2_hidden(k4_tarski(X16,esk3_3(X7,X8,X9)),X8)))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9))&((((r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk4_3(X7,X8,X9)),X7)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9))&(((r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk3_3(X7,X8,X9)),X8)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])])).
fof(c_0_6, plain, (![X3]:![X4]:v1_relat_1(k3_relat_1(X3,X4))), inference(variable_rename,[status(thm)],[dt_k3_relat_1])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk5_3(X5,X6,X7)),X5))|X6!=k9_xtuple_0(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k9_xtuple_0(X5)))&(((~r2_hidden(esk6_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk6_2(X5,X6),X11),X5))|X6=k9_xtuple_0(X5))&((r2_hidden(esk6_2(X5,X6),X6)|r2_hidden(k4_tarski(esk6_2(X5,X6),esk7_2(X5,X6)),X5))|X6=k9_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_xtuple_0])])])])])])])).
cnf(c_0_8,plain,(r2_hidden(k4_tarski(X4,esk1_5(X2,X3,X1,X4,X5)),X2)|~v1_relat_1(X1)|X1!=k3_relat_1(X2,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(v1_relat_1(k3_relat_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(r2_hidden(X3,X1)|X1!=k9_xtuple_0(X2)|~r2_hidden(k4_tarski(X3,X4),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(r2_hidden(k4_tarski(X1,esk1_5(X2,X3,k3_relat_1(X2,X3),X1,X4)),X2)|~r2_hidden(k4_tarski(X1,X4),k3_relat_1(X2,X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_8]), c_0_9])])).
cnf(c_0_12,plain,(r2_hidden(k4_tarski(X3,esk5_3(X2,X1,X3)),X2)|X1!=k9_xtuple_0(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(r2_hidden(X1,X2)|X2!=k9_xtuple_0(X3)|~r2_hidden(k4_tarski(X1,X4),k3_relat_1(X3,X5))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14,plain,(r2_hidden(k4_tarski(X1,esk5_3(X2,k9_xtuple_0(X2),X1)),X2)|~r2_hidden(X1,k9_xtuple_0(X2))), inference(er,[status(thm)],[c_0_12])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk8_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk8_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_16,plain,(r2_hidden(X1,X2)|X2!=k9_xtuple_0(X3)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X3,X4)))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17,plain,(r1_tarski(X1,X2)|r2_hidden(esk8_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k9_xtuple_0(X1)))))), inference(assume_negation,[status(cth)],[t25_relat_1])).
cnf(c_0_19,plain,(r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),X3)|r2_hidden(esk8_2(k9_xtuple_0(k3_relat_1(X1,X2)),X3),X4)|X4!=k9_xtuple_0(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_20, negated_conjecture, ((v1_relat_1(esk9_0)&(v1_relat_1(esk10_0)&~r1_tarski(k9_xtuple_0(k3_relat_1(esk9_0,esk10_0)),k9_xtuple_0(esk9_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])).
cnf(c_0_21,plain,(r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),X3)|r2_hidden(esk8_2(k9_xtuple_0(k3_relat_1(X1,X2)),X3),k9_xtuple_0(X1))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(~r1_tarski(k9_xtuple_0(k3_relat_1(esk9_0,esk10_0)),k9_xtuple_0(esk9_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k9_xtuple_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 26
# Proof object clause steps            : 15
# Proof object formula steps           : 11
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 104
# ...of these trivial                  : 1
# ...subsumed                          : 5
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 52
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 12
# Generated clauses                    : 893
# ...of the previous two non-trivial   : 802
# Contextual simplify-reflections      : 2
# Paramodulations                      : 809
# Factorizations                       : 14
# Equation resolutions                 : 70
# Current number of processed clauses  : 64
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 56
# Current number of unprocessed clauses: 476
# ...number of literals in the above   : 2709
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 1561
# Rec. Clause-clause subsumption calls : 582
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 132
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 22751

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.004 s
# Total time               : 0.032 s
# Maximum resident set size: 47348 pages
