# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t25_tex_2, conjecture, (![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>((g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))&(X3=X4&v2_tex_2(X3,X1)))=>v2_tex_2(X4,X2))))))), file('out/t25_tex_2__8_premises__rerun.E_input', t25_tex_2)).
fof(t65_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X1,X2)<=>m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))), file('out/t25_tex_2__8_premises__rerun.E_input', t65_pre_topc)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t25_tex_2__8_premises__rerun.E_input', dt_m1_pre_topc)).
fof(t2_tsep_1, axiom, (![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1))), file('out/t25_tex_2__8_premises__rerun.E_input', t2_tsep_1)).
fof(t59_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=>m1_pre_topc(X2,X1)))), file('out/t25_tex_2__8_premises__rerun.E_input', t59_pre_topc)).
fof(t35_borsuk_1, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>r1_tarski(u1_struct_0(X2),u1_struct_0(X1))))), file('out/t25_tex_2__8_premises__rerun.E_input', t35_borsuk_1)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t25_tex_2__8_premises__rerun.E_input', d10_xboole_0)).
fof(d5_tex_2, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((r1_tarski(X3,X2)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&k9_subset_1(u1_struct_0(X1),X2,X4)=X3)))))))))), file('out/t25_tex_2__8_premises__rerun.E_input', d5_tex_2)).
fof(t33_tops_2, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_pre_topc(X3,X1)=>(v3_pre_topc(X2,X1)=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))=>(X4=X2=>v3_pre_topc(X4,X3)))))))), file('out/t25_tex_2__8_premises__rerun.E_input', t33_tops_2)).
fof(c_0_9, negated_conjecture, (~(![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>((g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))&(X3=X4&v2_tex_2(X3,X1)))=>v2_tex_2(X4,X2)))))))), inference(assume_negation,[status(cth)],[t25_tex_2])).
fof(c_0_10, plain, (![X3]:![X4]:((((~m1_pre_topc(X3,X4)|m1_pre_topc(X3,g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4))))|~l1_pre_topc(X4))|~l1_pre_topc(X3))&(((~m1_pre_topc(X3,g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4)))|m1_pre_topc(X3,X4))|~l1_pre_topc(X4))|~l1_pre_topc(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_pre_topc])])])])])])).
fof(c_0_11, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
fof(c_0_12, plain, (![X2]:(~l1_pre_topc(X2)|m1_pre_topc(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
fof(c_0_13, negated_conjecture, ((l1_pre_topc(esk1_0)&(l1_pre_topc(esk2_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&((g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))&(esk3_0=esk4_0&v2_tex_2(esk3_0,esk1_0)))&~v2_tex_2(esk4_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
cnf(c_0_14,plain,(m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|m1_pre_topc(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t59_pre_topc])])])])])).
cnf(c_0_19,plain,(m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20,negated_conjecture,(m1_pre_topc(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21,negated_conjecture,(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_23, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|r1_tarski(u1_struct_0(X4),u1_struct_0(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_borsuk_1])])])])])).
cnf(c_0_24,plain,(m1_pre_topc(X1,X2)|~m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,negated_conjecture,(m1_pre_topc(esk2_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_17])])).
cnf(c_0_26,negated_conjecture,(m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_22])).
fof(c_0_27, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_28,plain,(r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29,negated_conjecture,(m1_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_22])])).
cnf(c_0_30,negated_conjecture,(m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X1,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_21]), c_0_17])])).
cnf(c_0_31,negated_conjecture,(m1_pre_topc(esk1_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_26]), c_0_22])])).
fof(c_0_32, plain, (![X5]:![X6]:![X7]:![X10]:(((((((m1_subset_1(esk5_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(X5)))|~r1_tarski(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~v2_tex_2(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))&((((((v3_pre_topc(esk5_3(X5,X6,X7),X5)|~r1_tarski(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~v2_tex_2(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))&(((((k9_subset_1(u1_struct_0(X5),X6,esk5_3(X5,X6,X7))=X7|~r1_tarski(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~v2_tex_2(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))))&((((m1_subset_1(esk6_2(X5,X6),k1_zfmisc_1(u1_struct_0(X5)))|v2_tex_2(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))&((((r1_tarski(esk6_2(X5,X6),X6)|v2_tex_2(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))&((((~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X5)))|(~v3_pre_topc(X10,X5)|k9_subset_1(u1_struct_0(X5),X6,X10)!=esk6_2(X5,X6)))|v2_tex_2(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_tex_2])])])])])])])).
cnf(c_0_33,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34,negated_conjecture,(esk3_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35,negated_conjecture,(~v2_tex_2(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37,negated_conjecture,(r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_22])])).
cnf(c_0_38,negated_conjecture,(m1_pre_topc(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39,plain,(v2_tex_2(X2,X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41,negated_conjecture,(~v2_tex_2(esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_35, c_0_34])).
cnf(c_0_42,negated_conjecture,(u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,negated_conjecture,(r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_38]), c_0_17])])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk6_2(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_17])]), c_0_41])).
cnf(c_0_45,negated_conjecture,(u1_struct_0(esk2_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43])])).
fof(c_0_46, plain, (![X5]:![X6]:![X7]:![X8]:(~l1_pre_topc(X5)|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(~m1_pre_topc(X7,X5)|(~v3_pre_topc(X6,X5)|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(X8!=X6|v3_pre_topc(X8,X7)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_tops_2])])])])])).
cnf(c_0_47,plain,(m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48,negated_conjecture,(m1_subset_1(esk6_2(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49,plain,(v2_tex_2(X2,X1)|r1_tarski(esk6_2(X1,X2),X2)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_50,plain,(v3_pre_topc(esk5_3(X1,X2,X3),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_51,plain,(v3_pre_topc(X1,X2)|X1!=X3|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52,negated_conjecture,(m1_subset_1(esk5_3(esk1_0,X1,esk6_2(esk2_0,esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(esk6_2(esk2_0,esk3_0),X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_22])])).
cnf(c_0_53,negated_conjecture,(r1_tarski(esk6_2(esk2_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_40]), c_0_17])]), c_0_41])).
cnf(c_0_54,negated_conjecture,(v2_tex_2(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_55,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_56,negated_conjecture,(v3_pre_topc(esk5_3(esk1_0,X1,esk6_2(esk2_0,esk3_0)),esk1_0)|~r1_tarski(esk6_2(esk2_0,esk3_0),X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_48]), c_0_22])])).
cnf(c_0_57,plain,(v3_pre_topc(X1,X2)|~v3_pre_topc(X1,X3)|~m1_pre_topc(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X3)), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_58,negated_conjecture,(m1_subset_1(esk5_3(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_59,negated_conjecture,(v3_pre_topc(esk5_3(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_60,plain,(v2_tex_2(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|k9_subset_1(u1_struct_0(X1),X2,X3)!=esk6_2(X1,X2)|~v3_pre_topc(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_61,negated_conjecture,(v3_pre_topc(esk5_3(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0)),X1)|~m1_pre_topc(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0)),k1_zfmisc_1(u1_struct_0(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_22])])).
cnf(c_0_62,plain,(k9_subset_1(u1_struct_0(X1),X2,esk5_3(X1,X2,X3))=X3|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_63,negated_conjecture,(v2_tex_2(X1,esk2_0)|k9_subset_1(u1_struct_0(esk1_0),X1,X2)!=esk6_2(esk2_0,X1)|~v3_pre_topc(X2,esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_45]), c_0_17])])).
cnf(c_0_64,negated_conjecture,(v3_pre_topc(esk5_3(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_45]), c_0_29]), c_0_58])])).
cnf(c_0_65,negated_conjecture,(k9_subset_1(u1_struct_0(esk1_0),X1,esk5_3(esk1_0,X1,esk6_2(esk2_0,esk3_0)))=esk6_2(esk2_0,esk3_0)|~r1_tarski(esk6_2(esk2_0,esk3_0),X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_48]), c_0_22])])).
cnf(c_0_66,negated_conjecture,(v2_tex_2(X1,esk2_0)|k9_subset_1(u1_struct_0(esk1_0),X1,esk5_3(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0)))!=esk6_2(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_58]), c_0_64])])).
cnf(c_0_67,negated_conjecture,(k9_subset_1(u1_struct_0(esk1_0),esk3_0,esk5_3(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0)))=esk6_2(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_68,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_55])]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 50
# Proof object formula steps           : 19
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 54
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 132
# ...of these trivial                  : 1
# ...subsumed                          : 9
# ...remaining for further processing  : 122
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 9
# Generated clauses                    : 215
# ...of the previous two non-trivial   : 200
# Contextual simplify-reflections      : 2
# Paramodulations                      : 212
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 88
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 62
# Current number of unprocessed clauses: 99
# ...number of literals in the above   : 396
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 276
# Rec. Clause-clause subsumption calls : 142
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 128
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8235

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43152 pages
