# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k2_subset_1, axiom, (![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), file('out/t25_compts_1__11_premises__rerun.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, (![X1]:k2_subset_1(X1)=X1), file('out/t25_compts_1__11_premises__rerun.E_input', d4_subset_1)).
fof(t7_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3))))), file('out/t25_compts_1__11_premises__rerun.E_input', t7_subset_1)).
fof(t25_compts_1, conjecture, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((v1_compts_1(X1)&(v8_pre_topc(X2)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&v5_pre_topc(X3,X1,X2))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X4,X1)=>v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)))))))), file('out/t25_compts_1__11_premises__rerun.E_input', t25_compts_1)).
fof(l3_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1)))), file('out/t25_compts_1__11_premises__rerun.E_input', l3_subset_1)).
fof(t10_compts_1, axiom, (![X1]:(l1_pre_topc(X1)=>(v1_compts_1(X1)<=>v2_compts_1(k2_struct_0(X1),X1)))), file('out/t25_compts_1__11_premises__rerun.E_input', t10_compts_1)).
fof(d3_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1))), file('out/t25_compts_1__11_premises__rerun.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t25_compts_1__11_premises__rerun.E_input', dt_l1_pre_topc)).
fof(t18_compts_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v2_compts_1(X2,X1)&(r1_tarski(X3,X2)&v4_pre_topc(X3,X1)))=>v2_compts_1(X3,X1)))))), file('out/t25_compts_1__11_premises__rerun.E_input', t18_compts_1)).
fof(t24_compts_1, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:((~(v2_struct_0(X3))&l1_pre_topc(X3))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))))=>((v5_pre_topc(X4,X1,X3)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X3),X4)=k2_struct_0(X3)&v2_compts_1(X2,X1)))=>v2_compts_1(k7_relset_1(u1_struct_0(X1),u1_struct_0(X3),X4,X2),X3))))))), file('out/t25_compts_1__11_premises__rerun.E_input', t24_compts_1)).
fof(dt_k7_relset_1, axiom, (![X1]:![X2]:![X3]:![X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2)))), file('out/t25_compts_1__11_premises__rerun.E_input', dt_k7_relset_1)).
fof(t16_compts_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v8_pre_topc(X1)&v2_compts_1(X2,X1))=>v4_pre_topc(X2,X1))))), file('out/t25_compts_1__11_premises__rerun.E_input', t16_compts_1)).
fof(c_0_12, plain, (![X2]:m1_subset_1(k2_subset_1(X2),k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_13, plain, (![X2]:k2_subset_1(X2)=X2), inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:((((m1_subset_1(esk1_3(X5,X6,X7),X5)|r1_tarski(X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(X5)))|~m1_subset_1(X6,k1_zfmisc_1(X5)))&((((r2_hidden(esk1_3(X5,X6,X7),X6)|r1_tarski(X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(X5)))|~m1_subset_1(X6,k1_zfmisc_1(X5)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|r1_tarski(X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(X5)))|~m1_subset_1(X6,k1_zfmisc_1(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])])])).
cnf(c_0_15,plain,(m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, (~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((v1_compts_1(X1)&(v8_pre_topc(X2)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&v5_pre_topc(X3,X1,X2))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X4,X1)=>v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2))))))))), inference(assume_negation,[status(cth)],[t25_compts_1])).
cnf(c_0_18,plain,(r1_tarski(X1,X3)|r2_hidden(esk1_3(X2,X1,X3),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_20, negated_conjecture, (((v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0))&((~v2_struct_0(esk3_0)&(v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0)))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0))))))&((v1_compts_1(esk2_0)&(v8_pre_topc(esk3_0)&(k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0)=k2_struct_0(esk3_0)&v5_pre_topc(esk4_0,esk2_0,esk3_0))))&(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(v4_pre_topc(esk5_0,esk2_0)&~v4_pre_topc(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,esk5_0),esk3_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X5,k1_zfmisc_1(X4))|(~r2_hidden(X6,X5)|r2_hidden(X6,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])])])).
cnf(c_0_22,plain,(r1_tarski(X1,X2)|r2_hidden(esk1_3(X2,X1,X2),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25,negated_conjecture,(r1_tarski(esk5_0,u1_struct_0(esk2_0))|r2_hidden(esk1_3(u1_struct_0(esk2_0),esk5_0,u1_struct_0(esk2_0)),esk5_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26,negated_conjecture,(r1_tarski(esk5_0,u1_struct_0(esk2_0))|r2_hidden(esk1_3(u1_struct_0(esk2_0),esk5_0,u1_struct_0(esk2_0)),X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_27, plain, (![X2]:(((~v1_compts_1(X2)|v2_compts_1(k2_struct_0(X2),X2))|~l1_pre_topc(X2))&((~v2_compts_1(k2_struct_0(X2),X2)|v1_compts_1(X2))|~l1_pre_topc(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_compts_1])])])).
fof(c_0_28, plain, (![X2]:(~l1_struct_0(X2)|k2_struct_0(X2)=u1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_29, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_30, plain, (![X4]:![X5]:![X6]:((~v2_pre_topc(X4)|~l1_pre_topc(X4))|(~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|((~v2_compts_1(X5,X4)|(~r1_tarski(X6,X5)|~v4_pre_topc(X6,X4)))|v2_compts_1(X6,X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_compts_1])])])])])).
cnf(c_0_31,plain,(r1_tarski(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(esk1_3(X2,X1,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_32,negated_conjecture,(r1_tarski(esk5_0,u1_struct_0(esk2_0))|r2_hidden(esk1_3(u1_struct_0(esk2_0),esk5_0,u1_struct_0(esk2_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_26, c_0_23])).
cnf(c_0_33,plain,(v2_compts_1(k2_struct_0(X1),X1)|~l1_pre_topc(X1)|~v1_compts_1(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34,negated_conjecture,(v1_compts_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_35,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_38, plain, (![X5]:![X6]:![X7]:![X8]:(~l1_pre_topc(X5)|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|((v2_struct_0(X7)|~l1_pre_topc(X7))|((~v1_funct_1(X8)|(~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X7))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X7))))))|((~v5_pre_topc(X8,X5,X7)|(k2_relset_1(u1_struct_0(X5),u1_struct_0(X7),X8)!=k2_struct_0(X7)|~v2_compts_1(X6,X5)))|v2_compts_1(k7_relset_1(u1_struct_0(X5),u1_struct_0(X7),X8,X6),X7))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_compts_1])])])])])])).
cnf(c_0_39,plain,(v2_compts_1(X1,X2)|~v4_pre_topc(X1,X2)|~r1_tarski(X1,X3)|~v2_compts_1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40,negated_conjecture,(r1_tarski(esk5_0,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_19]), c_0_23])])).
cnf(c_0_41,negated_conjecture,(v2_compts_1(k2_struct_0(esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_42,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,plain,(v2_compts_1(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|v2_struct_0(X2)|~v2_compts_1(X4,X1)|k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44,negated_conjecture,(k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0)=k2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_45,negated_conjecture,(v5_pre_topc(esk4_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_46,negated_conjecture,(v1_funct_2(esk4_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_47,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_48,negated_conjecture,(l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_49,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0))))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_50,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_51,negated_conjecture,(v2_compts_1(esk5_0,X1)|~v2_compts_1(u1_struct_0(esk2_0),X1)|~v4_pre_topc(esk5_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_52,negated_conjecture,(v4_pre_topc(esk5_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_53,negated_conjecture,(v2_compts_1(u1_struct_0(esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_35])])).
cnf(c_0_54,negated_conjecture,(v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_55, plain, (![X5]:![X6]:![X7]:![X8]:(~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|m1_subset_1(k7_relset_1(X5,X6,X7,X8),k1_zfmisc_1(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])])])).
fof(c_0_56, plain, (![X3]:![X4]:((~v2_pre_topc(X3)|~l1_pre_topc(X3))|(~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|((~v8_pre_topc(X3)|~v2_compts_1(X4,X3))|v4_pre_topc(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_compts_1])])])])])).
cnf(c_0_57,negated_conjecture,(v2_compts_1(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,X1),esk3_0)|~v2_compts_1(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_35]), c_0_49])]), c_0_50])).
cnf(c_0_58,negated_conjecture,(v2_compts_1(esk5_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_35]), c_0_54]), c_0_19]), c_0_23])])).
cnf(c_0_59,plain,(m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60,plain,(v4_pre_topc(X1,X2)|~v2_compts_1(X1,X2)|~v8_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_61,negated_conjecture,(v2_compts_1(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,esk5_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_23])])).
cnf(c_0_62,negated_conjecture,(v8_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_63,negated_conjecture,(v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_64,negated_conjecture,(m1_subset_1(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,X1),k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(spm,[status(thm)],[c_0_59, c_0_49])).
cnf(c_0_65,negated_conjecture,(~v4_pre_topc(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,esk5_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_66,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_48]), c_0_63]), c_0_64])]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 67
# Proof object clause steps            : 42
# Proof object formula steps           : 25
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 28
# Processed clauses                    : 98
# ...of these trivial                  : 2
# ...subsumed                          : 3
# ...remaining for further processing  : 93
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 73
# ...of the previous two non-trivial   : 66
# Contextual simplify-reflections      : 0
# Paramodulations                      : 73
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 32
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 33
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 252
# Rec. Clause-clause subsumption calls : 72
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4141

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 47680 pages
