# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t25_yellow_0, conjecture, (![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k12_lattice3(X1,X2,X3)<=>r3_orders_2(X1,X2,X3)))))), file('out/t25_yellow_0__6_premises__rerun.E_input', t25_yellow_0)).
fof(redefinition_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3))), file('out/t25_yellow_0__6_premises__rerun.E_input', redefinition_k12_lattice3)).
fof(l28_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))))))), file('out/t25_yellow_0__6_premises__rerun.E_input', l28_lattice3)).
fof(cc2_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t25_yellow_0__6_premises__rerun.E_input', cc2_lattice3)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t25_yellow_0__6_premises__rerun.E_input', redefinition_r3_orders_2)).
fof(reflexivity_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2))), file('out/t25_yellow_0__6_premises__rerun.E_input', reflexivity_r3_orders_2)).
fof(t15_lattice3, axiom, (![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2))))), file('out/t25_yellow_0__6_premises__rerun.E_input', t15_lattice3)).
fof(c_0_7, negated_conjecture, (~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k12_lattice3(X1,X2,X3)<=>r3_orders_2(X1,X2,X3))))))), inference(assume_negation,[status(cth)],[t25_yellow_0])).
fof(c_0_8, negated_conjecture, (((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((esk2_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r3_orders_2(esk1_0,esk2_0,esk3_0))&(esk2_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r3_orders_2(esk1_0,esk2_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k12_lattice3(X4,X5,X6)=k11_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
fof(c_0_10, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((r1_orders_2(X6,X9,X7)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((r1_orders_2(X6,X9,X8)|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((((((~m1_subset_1(X10,u1_struct_0(X6))|((~r1_orders_2(X6,X10,X7)|~r1_orders_2(X6,X10,X8))|r1_orders_2(X6,X10,X9)))|X9!=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))))&((((((((m1_subset_1(esk4_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((((((r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))&(((((((~r1_orders_2(X6,esk4_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8))|~r1_orders_2(X6,X9,X7))|X9=k11_lattice3(X6,X7,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~m1_subset_1(X7,u1_struct_0(X6)))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l28_lattice3])])])])])])])])).
fof(c_0_11, plain, (![X2]:(~l1_orders_2(X2)|(~v2_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_13,negated_conjecture,(r3_orders_2(esk1_0,esk2_0,esk3_0)|esk2_0=k12_lattice3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20,plain,(v2_struct_0(X1)|r1_orders_2(X1,X4,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X4!=k11_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21,plain,(~v2_struct_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,esk3_0)=esk2_0|r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_24,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|r3_orders_2(X4,X5,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])])])).
cnf(c_0_26,plain,(r1_orders_2(X1,X2,X3)|X2!=k11_lattice3(X1,X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,esk3_0)=esk2_0|r1_orders_2(esk1_0,esk2_0,esk3_0)|v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_16]), c_0_15]), c_0_17]), c_0_24])])).
cnf(c_0_28,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29,plain,(v2_struct_0(X1)|X4=k11_lattice3(X1,X2,X3)|r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_30,plain,(r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31,negated_conjecture,(~r3_orders_2(esk1_0,esk2_0,esk3_0)|esk2_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_32,plain,(v2_struct_0(X2)|r3_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r1_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_33,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,X1,esk3_0)|v2_struct_0(esk1_0)|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19])])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|k11_lattice3(X4,X5,X6)=k11_lattice3(X4,X6,X5))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_lattice3])])])])])).
cnf(c_0_35,plain,(X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk4_4(X2,X3,X4,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_28, c_0_21])).
cnf(c_0_36,plain,(X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk4_4(X2,X3,X4,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_29, c_0_21])).
cnf(c_0_37,negated_conjecture,(v2_struct_0(esk1_0)|r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_15]), c_0_17]), c_0_24])])).
cnf(c_0_38,negated_conjecture,(k11_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0|~r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_39,negated_conjecture,(v2_struct_0(esk1_0)|r3_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_24]), c_0_17])])).
cnf(c_0_40,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk3_0)|v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_16])).
cnf(c_0_41,plain,(k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42,plain,(k11_lattice3(X1,X2,X3)=X3|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43,negated_conjecture,(r1_orders_2(esk1_0,X1,X1)|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_37]), c_0_17]), c_0_24])])).
cnf(c_0_44,negated_conjecture,(v2_struct_0(esk1_0)|k11_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_16]), c_0_15])]), c_0_40])).
cnf(c_0_45,plain,(k11_lattice3(X1,X2,X3)=X2|~r1_orders_2(X1,X2,X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk2_0)|v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_43, c_0_16])).
cnf(c_0_47,negated_conjecture,(v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_16]), c_0_15]), c_0_17]), c_0_18]), c_0_19])]), c_0_40]), c_0_46])).
cnf(c_0_48,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_47]), c_0_17]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 49
# Proof object clause steps            : 34
# Proof object formula steps           : 15
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 75
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 61
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 9
# Generated clauses                    : 62
# ...of the previous two non-trivial   : 40
# Contextual simplify-reflections      : 18
# Paramodulations                      : 53
# Factorizations                       : 0
# Equation resolutions                 : 9
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 56
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 973
# Rec. Clause-clause subsumption calls : 120
# Non-unit clause-clause subsumptions  : 34
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4485

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 43188 pages
