# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t13_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>(k3_waybel_0(X1,X3)=k3_waybel_0(X2,X4)&k4_waybel_0(X1,X3)=k4_waybel_0(X2,X4))))))))), file('out/t25_waybel_0__3_premises__rerun.E_input', t13_waybel_0)).
fof(t25_waybel_0, conjecture, (![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>((v12_waybel_0(X3,X1)=>v12_waybel_0(X4,X2))&(v13_waybel_0(X3,X1)=>v13_waybel_0(X4,X2)))))))))), file('out/t25_waybel_0__3_premises__rerun.E_input', t25_waybel_0)).
fof(t24_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>r1_tarski(k4_waybel_0(X1,X2),X2))))), file('out/t25_waybel_0__3_premises__rerun.E_input', t24_waybel_0)).
fof(t23_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>r1_tarski(k3_waybel_0(X1,X2),X2))))), file('out/t25_waybel_0__3_premises__rerun.E_input', t23_waybel_0)).
fof(c_0_4, plain, (![X5]:![X6]:![X7]:![X8]:(((((((k3_waybel_0(X5,X7)=k3_waybel_0(X6,X8)|X7!=X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|g1_orders_2(u1_struct_0(X5),u1_orders_2(X5))!=g1_orders_2(u1_struct_0(X6),u1_orders_2(X6)))|~l1_orders_2(X6))|~l1_orders_2(X5))&((((((k4_waybel_0(X5,X7)=k4_waybel_0(X6,X8)|X7!=X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|g1_orders_2(u1_struct_0(X5),u1_orders_2(X5))!=g1_orders_2(u1_struct_0(X6),u1_orders_2(X6)))|~l1_orders_2(X6))|~l1_orders_2(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_waybel_0])])])])])])).
fof(c_0_5, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>((v12_waybel_0(X3,X1)=>v12_waybel_0(X4,X2))&(v13_waybel_0(X3,X1)=>v13_waybel_0(X4,X2))))))))))), inference(assume_negation,[status(cth)],[t25_waybel_0])).
cnf(c_0_6,plain,(k4_waybel_0(X1,X3)=k4_waybel_0(X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, negated_conjecture, ((l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(esk3_0=esk4_0&(((v13_waybel_0(esk3_0,esk1_0)|v12_waybel_0(esk3_0,esk1_0))&(~v13_waybel_0(esk4_0,esk2_0)|v12_waybel_0(esk3_0,esk1_0)))&((v13_waybel_0(esk3_0,esk1_0)|~v12_waybel_0(esk4_0,esk2_0))&(~v13_waybel_0(esk4_0,esk2_0)|~v12_waybel_0(esk4_0,esk2_0))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_8,plain,(k3_waybel_0(X1,X3)=k3_waybel_0(X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,plain,(k4_waybel_0(X1,X2)=k4_waybel_0(X3,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X3)|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(k3_waybel_0(X1,X2)=k3_waybel_0(X3,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X3)|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_8])).
fof(c_0_13, plain, (![X3]:![X4]:((((~v13_waybel_0(X4,X3)|r1_tarski(k4_waybel_0(X3,X4),X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_orders_2(X3))&(((~r1_tarski(k4_waybel_0(X3,X4),X4)|v13_waybel_0(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_waybel_0])])])])])])).
cnf(c_0_14,negated_conjecture,(k4_waybel_0(X1,X2)=k4_waybel_0(esk2_0,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])])).
cnf(c_0_15,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_16, plain, (![X3]:![X4]:((((~v12_waybel_0(X4,X3)|r1_tarski(k3_waybel_0(X3,X4),X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_orders_2(X3))&(((~r1_tarski(k3_waybel_0(X3,X4),X4)|v12_waybel_0(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_waybel_0])])])])])])).
cnf(c_0_17,negated_conjecture,(k3_waybel_0(X1,X2)=k3_waybel_0(esk2_0,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_10]), c_0_11])])).
cnf(c_0_18,plain,(v13_waybel_0(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(k4_waybel_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(k4_waybel_0(esk2_0,X1)=k4_waybel_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15])])).
cnf(c_0_20,plain,(v12_waybel_0(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(k3_waybel_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,negated_conjecture,(k3_waybel_0(esk2_0,X1)=k3_waybel_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_17]), c_0_15])])).
cnf(c_0_22,negated_conjecture,(v13_waybel_0(X1,esk2_0)|~r1_tarski(k4_waybel_0(esk1_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_11])])).
cnf(c_0_23,plain,(r1_tarski(k4_waybel_0(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v13_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25,negated_conjecture,(esk3_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_26,negated_conjecture,(v12_waybel_0(X1,esk2_0)|~r1_tarski(k3_waybel_0(esk1_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_11])])).
cnf(c_0_27,plain,(r1_tarski(k3_waybel_0(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v12_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28,negated_conjecture,(v12_waybel_0(esk3_0,esk1_0)|~v13_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_29,negated_conjecture,(v13_waybel_0(X1,esk2_0)|~v13_waybel_0(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_15])])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32,negated_conjecture,(~v12_waybel_0(esk4_0,esk2_0)|~v13_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_33,negated_conjecture,(v13_waybel_0(esk3_0,esk1_0)|~v12_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34,negated_conjecture,(v12_waybel_0(X1,esk2_0)|~v12_waybel_0(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_15])])).
cnf(c_0_35,negated_conjecture,(v12_waybel_0(esk3_0,esk1_0)|~v13_waybel_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_28, c_0_25])).
cnf(c_0_36,negated_conjecture,(v13_waybel_0(esk3_0,esk2_0)|~v13_waybel_0(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_37,negated_conjecture,(v12_waybel_0(esk3_0,esk1_0)|v13_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_38,negated_conjecture,(~v13_waybel_0(esk3_0,esk2_0)|~v12_waybel_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_25]), c_0_25])).
cnf(c_0_39,negated_conjecture,(v13_waybel_0(esk3_0,esk1_0)|~v12_waybel_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_33, c_0_25])).
cnf(c_0_40,negated_conjecture,(v12_waybel_0(esk3_0,esk2_0)|~v12_waybel_0(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_30]), c_0_31])])).
cnf(c_0_41,negated_conjecture,(v12_waybel_0(esk3_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_42,negated_conjecture,(~v12_waybel_0(esk3_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_36]), c_0_39])).
cnf(c_0_43,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])]), c_0_42]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 35
# Proof object formula steps           : 9
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 51
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 49
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 27
# ...of the previous two non-trivial   : 21
# Contextual simplify-reflections      : 2
# Paramodulations                      : 21
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 19
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 4
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 200
# Rec. Clause-clause subsumption calls : 17
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2052

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45420 pages
