# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, (![X1]:![X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), file('out/t26_wellord1__12_premises__rerun.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t26_wellord1__12_premises__rerun.E_input', t70_enumset1)).
fof(t127_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(k3_xboole_0(X1,X2),X3))), file('out/t26_wellord1__12_premises__rerun.E_input', t127_relat_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t26_wellord1__12_premises__rerun.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), file('out/t26_wellord1__12_premises__rerun.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), file('out/t26_wellord1__12_premises__rerun.E_input', t73_enumset1)).
fof(t26_wellord1, conjecture, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k2_wellord1(k2_wellord1(X3,X1),X2)=k2_wellord1(X3,k3_xboole_0(X1,X2)))), file('out/t26_wellord1__12_premises__rerun.E_input', t26_wellord1)).
fof(t100_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k7_relat_1(k7_relat_1(X3,X1),X2)=k7_relat_1(X3,k3_xboole_0(X1,X2)))), file('out/t26_wellord1__12_premises__rerun.E_input', t100_relat_1)).
fof(t17_wellord1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k2_wellord1(X2,X1)=k7_relat_1(k8_relat_1(X1,X2),X1))), file('out/t26_wellord1__12_premises__rerun.E_input', t17_wellord1)).
fof(dt_k8_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>v1_relat_1(k8_relat_1(X1,X2)))), file('out/t26_wellord1__12_premises__rerun.E_input', dt_k8_relat_1)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t26_wellord1__12_premises__rerun.E_input', commutativity_k2_tarski)).
fof(t140_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k7_relat_1(k8_relat_1(X1,X3),X2)=k8_relat_1(X1,k7_relat_1(X3,X2)))), file('out/t26_wellord1__12_premises__rerun.E_input', t140_relat_1)).
fof(dt_k7_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2)))), file('out/t26_wellord1__12_premises__rerun.E_input', dt_k7_relat_1)).
fof(c_0_13, plain, (![X3]:![X4]:k1_setfam_1(k2_tarski(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_14, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k8_relat_1(X4,k8_relat_1(X5,X6))=k8_relat_1(k3_xboole_0(X4,X5),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t127_relat_1])])).
cnf(c_0_16,plain,(k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:k3_enumset1(X5,X5,X6,X7,X8)=k2_enumset1(X5,X6,X7,X8)), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_20, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k4_enumset1(X6,X6,X7,X8,X9,X10)=k3_enumset1(X6,X7,X8,X9,X10)), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_21, negated_conjecture, (~(![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k2_wellord1(k2_wellord1(X3,X1),X2)=k2_wellord1(X3,k3_xboole_0(X1,X2))))), inference(assume_negation,[status(cth)],[t26_wellord1])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k7_relat_1(k7_relat_1(X6,X4),X5)=k7_relat_1(X6,k3_xboole_0(X4,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t100_relat_1])])).
fof(c_0_23, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k2_wellord1(X4,X3)=k7_relat_1(k8_relat_1(X3,X4),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_wellord1])])).
cnf(c_0_24,plain,(k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(k3_xboole_0(X1,X2),X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_26,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, (![X3]:![X4]:(~v1_relat_1(X4)|v1_relat_1(k8_relat_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relat_1])])).
fof(c_0_30, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_31, negated_conjecture, ((v1_relat_1(esk3_0)&k2_wellord1(k2_wellord1(esk3_0,esk1_0),esk2_0)!=k2_wellord1(esk3_0,k3_xboole_0(esk1_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
cnf(c_0_32,plain,(k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33,plain,(k2_wellord1(X1,X2)=k7_relat_1(k8_relat_1(X2,X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34,plain,(k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)),X3)|~v1_relat_1(X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_35,plain,(v1_relat_1(k8_relat_1(X1,X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|k7_relat_1(k8_relat_1(X4,X6),X5)=k8_relat_1(X4,k7_relat_1(X6,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t140_relat_1])])).
cnf(c_0_37,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38,negated_conjecture,(k2_wellord1(k2_wellord1(esk3_0,esk1_0),esk2_0)!=k2_wellord1(esk3_0,k3_xboole_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39,plain,(k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,k1_setfam_1(k4_enumset1(X2,X2,X2,X2,X2,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_25]), c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_40,plain,(k7_relat_1(k8_relat_1(X1,k8_relat_1(X2,X3)),k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))=k2_wellord1(X3,k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41,plain,(v1_relat_1(k8_relat_1(X1,k8_relat_1(X2,X3)))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_35, c_0_34])).
cnf(c_0_42,plain,(k7_relat_1(k8_relat_1(X1,X2),X3)=k8_relat_1(X1,k7_relat_1(X2,X3))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_43, plain, (![X3]:![X4]:(~v1_relat_1(X3)|v1_relat_1(k7_relat_1(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])])])).
cnf(c_0_44,plain,(k4_enumset1(X1,X1,X1,X1,X1,X2)=k4_enumset1(X2,X2,X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_17]), c_0_17]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_28]), c_0_28])).
cnf(c_0_45,negated_conjecture,(k2_wellord1(k2_wellord1(esk3_0,esk1_0),esk2_0)!=k2_wellord1(esk3_0,k1_setfam_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_25]), c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_46,plain,(k2_wellord1(X1,k1_setfam_1(k4_enumset1(X2,X2,X2,X2,X2,X3)))=k7_relat_1(k7_relat_1(k8_relat_1(X2,k8_relat_1(X3,X1)),X2),X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_47,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_48,plain,(k8_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)),k7_relat_1(X3,X4))=k7_relat_1(k8_relat_1(X1,k8_relat_1(X2,X3)),X4)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_42, c_0_34])).
cnf(c_0_49,plain,(v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50,plain,(k8_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)),X3)=k8_relat_1(X2,k8_relat_1(X1,X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_34, c_0_44])).
cnf(c_0_51,negated_conjecture,(k2_wellord1(k2_wellord1(esk3_0,esk1_0),esk2_0)!=k7_relat_1(k7_relat_1(k8_relat_1(esk1_0,k8_relat_1(esk2_0,esk3_0)),esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_52,plain,(k7_relat_1(k8_relat_1(X1,k8_relat_1(X2,X3)),X4)=k8_relat_1(X1,k8_relat_1(X2,k7_relat_1(X3,X4)))|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_48]), c_0_49])).
cnf(c_0_53,plain,(k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(X2,k8_relat_1(X1,X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_34, c_0_50])).
cnf(c_0_54,negated_conjecture,(k7_relat_1(k7_relat_1(k8_relat_1(esk1_0,k8_relat_1(esk2_0,esk3_0)),esk1_0),esk2_0)!=k7_relat_1(k8_relat_1(esk2_0,k2_wellord1(esk3_0,esk1_0)),esk2_0)|~v1_relat_1(k2_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_33])).
cnf(c_0_55,plain,(k7_relat_1(k8_relat_1(X1,k8_relat_1(X2,X3)),X4)=k8_relat_1(X2,k8_relat_1(X1,k7_relat_1(X3,X4)))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_56,negated_conjecture,(k7_relat_1(k8_relat_1(esk2_0,k2_wellord1(esk3_0,esk1_0)),esk2_0)!=k7_relat_1(k8_relat_1(esk2_0,k8_relat_1(esk1_0,k7_relat_1(esk3_0,esk1_0))),esk2_0)|~v1_relat_1(k2_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_47])])).
cnf(c_0_57,plain,(k2_wellord1(X1,X2)=k8_relat_1(X2,k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_42])).
cnf(c_0_58,negated_conjecture,(~v1_relat_1(k8_relat_1(esk1_0,k7_relat_1(esk3_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_47])])).
cnf(c_0_59,plain,(v1_relat_1(k8_relat_1(X1,k7_relat_1(X2,X3)))|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_42]), c_0_35])).
cnf(c_0_60,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_47])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 61
# Proof object clause steps            : 34
# Proof object formula steps           : 27
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 9
# Processed clauses                    : 675
# ...of these trivial                  : 0
# ...subsumed                          : 460
# ...remaining for further processing  : 215
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 51
# Backward-rewritten                   : 0
# Generated clauses                    : 5273
# ...of the previous two non-trivial   : 5272
# Contextual simplify-reflections      : 307
# Paramodulations                      : 5273
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 155
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 148
# Current number of unprocessed clauses: 3792
# ...number of literals in the above   : 10899
# Current number of archived formulas  : 0
# Current number of archived clauses   : 65
# Clause-clause subsumption calls (NU) : 10420
# Rec. Clause-clause subsumption calls : 10420
# Non-unit clause-clause subsumptions  : 817
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 71646

# -------------------------------------------------
# User time                : 0.056 s
# System time              : 0.000 s
# Total time               : 0.056 s
# Maximum resident set size: 47128 pages
