# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t26_waybel_9, conjecture, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_waybel_9(X1))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(v5_pre_topc(k4_waybel_1(X1,X3),X1,X1)=>r2_hidden(k12_lattice3(X1,X3,k11_yellow_6(X1,X2)),k10_yellow_6(X1,k3_waybel_2(X1,X3,X2)))))))), file('out/t26_waybel_9__12_premises__rerun.E_input', t26_waybel_9)).
fof(dt_l1_waybel_9, axiom, (![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1)))), file('out/t26_waybel_9__12_premises__rerun.E_input', dt_l1_waybel_9)).
fof(cc2_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t26_waybel_9__12_premises__rerun.E_input', cc2_lattice3)).
fof(dt_k11_yellow_6, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&l1_pre_topc(X1))))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1))))))=>m1_subset_1(k11_yellow_6(X1,X2),u1_struct_0(X1)))), file('out/t26_waybel_9__12_premises__rerun.E_input', dt_k11_yellow_6)).
fof(dt_k4_waybel_1, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(v1_funct_1(k4_waybel_1(X1,X2))&(v1_funct_2(k4_waybel_1(X1,X2),u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(k4_waybel_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))))), file('out/t26_waybel_9__12_premises__rerun.E_input', dt_k4_waybel_1)).
fof(t18_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k6_waybel_9(X1,X1,k4_waybel_1(X1,X3),X2)=k3_waybel_2(X1,X3,X2))))), file('out/t26_waybel_9__12_premises__rerun.E_input', t18_waybel_9)).
fof(d18_waybel_1, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))=>(X3=k4_waybel_1(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),X3,X4)=k11_lattice3(X1,X2,X4))))))), file('out/t26_waybel_9__12_premises__rerun.E_input', d18_waybel_1)).
fof(t25_waybel_9, axiom, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_waybel_9(X1))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))=>(v5_pre_topc(X3,X1,X1)=>r2_hidden(k2_yellow_6(u1_struct_0(X1),X1,X3,k11_yellow_6(X1,X2)),k10_yellow_6(X1,k6_waybel_9(X1,X1,X3,X2)))))))), file('out/t26_waybel_9__12_premises__rerun.E_input', t25_waybel_9)).
fof(redefinition_k2_yellow_6, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((~(v2_struct_0(X2))&l1_orders_2(X2))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))))&m1_subset_1(X4,X1))))=>k2_yellow_6(X1,X2,X3,X4)=k1_funct_1(X3,X4))), file('out/t26_waybel_9__12_premises__rerun.E_input', redefinition_k2_yellow_6)).
fof(redefinition_k12_lattice3, axiom, (![X1]:![X2]:![X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3))), file('out/t26_waybel_9__12_premises__rerun.E_input', redefinition_k12_lattice3)).
fof(redefinition_k3_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4))), file('out/t26_waybel_9__12_premises__rerun.E_input', redefinition_k3_funct_2)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t26_waybel_9__12_premises__rerun.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t26_waybel_9__12_premises__rerun.E_input', dt_l1_pre_topc)).
fof(c_0_13, negated_conjecture, (~(![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_waybel_9(X1))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(v5_pre_topc(k4_waybel_1(X1,X3),X1,X1)=>r2_hidden(k12_lattice3(X1,X3,k11_yellow_6(X1,X2)),k10_yellow_6(X1,k3_waybel_2(X1,X3,X2))))))))), inference(assume_negation,[status(cth)],[t26_waybel_9])).
fof(c_0_14, plain, (![X2]:((l1_pre_topc(X2)|~l1_waybel_9(X2))&(l1_orders_2(X2)|~l1_waybel_9(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
fof(c_0_15, negated_conjecture, (((v2_pre_topc(esk1_0)&(v8_pre_topc(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&l1_waybel_9(esk1_0))))))))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&(v3_yellow_6(esk2_0,esk1_0)&l1_waybel_0(esk2_0,esk1_0)))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(v5_pre_topc(k4_waybel_1(esk1_0,esk3_0),esk1_0,esk1_0)&~r2_hidden(k12_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])])])).
fof(c_0_16, plain, (![X2]:(~l1_orders_2(X2)|(~v2_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
cnf(c_0_17,plain,(l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18,negated_conjecture,(l1_waybel_9(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, (![X3]:![X4]:(((v2_struct_0(X3)|(~v2_pre_topc(X3)|(~v8_pre_topc(X3)|~l1_pre_topc(X3))))|(v2_struct_0(X4)|(~v4_orders_2(X4)|(~v7_waybel_0(X4)|(~v3_yellow_6(X4,X3)|~l1_waybel_0(X4,X3))))))|m1_subset_1(k11_yellow_6(X3,X4),u1_struct_0(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k11_yellow_6])])])).
cnf(c_0_20,plain,(~v2_struct_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,negated_conjecture,(l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_23, plain, (![X3]:![X4]:((v1_funct_1(k4_waybel_1(X3,X4))|((v2_struct_0(X3)|~l1_orders_2(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&((v1_funct_2(k4_waybel_1(X3,X4),u1_struct_0(X3),u1_struct_0(X3))|((v2_struct_0(X3)|~l1_orders_2(X3))|~m1_subset_1(X4,u1_struct_0(X3))))&(m1_subset_1(k4_waybel_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3))))|((v2_struct_0(X3)|~l1_orders_2(X3))|~m1_subset_1(X4,u1_struct_0(X3))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_waybel_1])])])])).
fof(c_0_24, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|~l1_orders_2(X4))|((v2_struct_0(X5)|~l1_waybel_0(X5,X4))|(~m1_subset_1(X6,u1_struct_0(X4))|k6_waybel_9(X4,X4,k4_waybel_1(X4,X6),X5)=k3_waybel_2(X4,X6,X5))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_9])])])])])])).
fof(c_0_25, plain, (![X5]:![X6]:![X7]:![X8]:(((((X7!=k4_waybel_1(X5,X6)|(~m1_subset_1(X8,u1_struct_0(X5))|k3_funct_2(u1_struct_0(X5),u1_struct_0(X5),X7,X8)=k11_lattice3(X5,X6,X8)))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X5))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X5)))))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l1_orders_2(X5)))&(((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|X7=k4_waybel_1(X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X5))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X5)))))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l1_orders_2(X5)))&((((k3_funct_2(u1_struct_0(X5),u1_struct_0(X5),X7,esk4_3(X5,X6,X7))!=k11_lattice3(X5,X6,esk4_3(X5,X6,X7))|X7=k4_waybel_1(X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X5))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X5)))))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_waybel_1])])])])])])])])).
cnf(c_0_26,plain,(m1_subset_1(k11_yellow_6(X1,X2),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~v3_yellow_6(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,negated_conjecture,(v3_yellow_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29,negated_conjecture,(v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30,negated_conjecture,(v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31,negated_conjecture,(v8_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_32,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_33,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_34,negated_conjecture,(~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
fof(c_0_35, plain, (![X4]:![X5]:![X6]:((~v2_pre_topc(X4)|(~v8_pre_topc(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|(~v1_lattice3(X4)|(~v2_lattice3(X4)|~l1_waybel_9(X4))))))))|((v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|(~v3_yellow_6(X5,X4)|~l1_waybel_0(X5,X4)))))|((~v1_funct_1(X6)|(~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X4))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X4))))))|(~v5_pre_topc(X6,X4,X4)|r2_hidden(k2_yellow_6(u1_struct_0(X4),X4,X6,k11_yellow_6(X4,X5)),k10_yellow_6(X4,k6_waybel_9(X4,X4,X6,X5)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_waybel_9])])])])])])).
cnf(c_0_36,plain,(v2_struct_0(X2)|v1_funct_2(k4_waybel_1(X2,X1),u1_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_38,plain,(v2_struct_0(X2)|v1_funct_1(k4_waybel_1(X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_39,plain,(v2_struct_0(X2)|m1_subset_1(k4_waybel_1(X2,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40,plain,(k6_waybel_9(X1,X1,k4_waybel_1(X1,X2),X3)=k3_waybel_2(X1,X2,X3)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_41, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((v2_struct_0(X6)|~l1_orders_2(X6))|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,u1_struct_0(X6))))))|~m1_subset_1(X8,X5))))|k2_yellow_6(X5,X6,X7,X8)=k1_funct_1(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k2_yellow_6])])])).
cnf(c_0_42,plain,(v2_struct_0(X1)|k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),X3,X4)=k11_lattice3(X1,X2,X4)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k4_waybel_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_43,negated_conjecture,(m1_subset_1(k11_yellow_6(esk1_0,esk2_0),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34])).
cnf(c_0_44,plain,(l1_pre_topc(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_45, plain, (![X4]:![X5]:![X6]:(((~v5_orders_2(X4)|(~v2_lattice3(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k12_lattice3(X4,X5,X6)=k11_lattice3(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_46,plain,(r2_hidden(k2_yellow_6(u1_struct_0(X1),X1,X2,k11_yellow_6(X1,X3)),k10_yellow_6(X1,k6_waybel_9(X1,X1,X2,X3)))|v2_struct_0(X3)|~v5_pre_topc(X2,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_waybel_0(X3,X1)|~v3_yellow_6(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_waybel_9(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47,negated_conjecture,(v5_pre_topc(k4_waybel_1(esk1_0,esk3_0),esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_48,negated_conjecture,(v1_funct_2(k4_waybel_1(esk1_0,esk3_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_22])]), c_0_34])).
cnf(c_0_49,negated_conjecture,(v1_funct_1(k4_waybel_1(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_37]), c_0_22])]), c_0_34])).
cnf(c_0_50,negated_conjecture,(m1_subset_1(k4_waybel_1(esk1_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_37]), c_0_22])]), c_0_34])).
cnf(c_0_51,negated_conjecture,(v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_52,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_53,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_54,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_55,negated_conjecture,(k6_waybel_9(esk1_0,esk1_0,k4_waybel_1(esk1_0,X1),esk2_0)=k3_waybel_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_28]), c_0_22])]), c_0_33]), c_0_34])).
cnf(c_0_56,plain,(k2_yellow_6(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v2_struct_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~v1_funct_2(X3,X1,u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_57, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X8,X5)))|k3_funct_2(X5,X6,X7,X8)=k1_funct_1(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_58,plain,(k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),k4_waybel_1(X1,X2),X3)=k11_lattice3(X1,X2,X3)|v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_42]), c_0_39]), c_0_38]), c_0_36])).
cnf(c_0_59,negated_conjecture,(m1_subset_1(k11_yellow_6(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_18])])).
cnf(c_0_60,plain,(k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_61,negated_conjecture,(r2_hidden(k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,X1)),k10_yellow_6(esk1_0,k6_waybel_9(esk1_0,esk1_0,k4_waybel_1(esk1_0,esk3_0),X1)))|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v3_yellow_6(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_18]), c_0_21]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_31]), c_0_32])])).
cnf(c_0_62,negated_conjecture,(k6_waybel_9(esk1_0,esk1_0,k4_waybel_1(esk1_0,esk3_0),esk2_0)=k3_waybel_2(esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_37])).
cnf(c_0_63,negated_conjecture,(k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),X1)=k1_funct_1(k4_waybel_1(esk1_0,esk3_0),X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_48]), c_0_49]), c_0_22]), c_0_50])]), c_0_34])).
cnf(c_0_64,plain,(k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65,negated_conjecture,(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k4_waybel_1(esk1_0,X1),k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,X1,k11_yellow_6(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_22])]), c_0_34])).
cnf(c_0_66,negated_conjecture,(k12_lattice3(esk1_0,X1,k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,X1,k11_yellow_6(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_59]), c_0_22]), c_0_21]), c_0_52])])).
cnf(c_0_67,negated_conjecture,(r2_hidden(k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_27]), c_0_62]), c_0_28]), c_0_29]), c_0_30])]), c_0_33])).
cnf(c_0_68,negated_conjecture,(k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))=k1_funct_1(k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_63, c_0_59])).
cnf(c_0_69,negated_conjecture,(k1_funct_1(k4_waybel_1(esk1_0,esk3_0),X1)=k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k4_waybel_1(esk1_0,esk3_0),X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_70,negated_conjecture,(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_65, c_0_37])).
cnf(c_0_71,negated_conjecture,(~r2_hidden(k12_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_72,negated_conjecture,(k12_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_66, c_0_37])).
fof(c_0_73, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_74,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k1_funct_1(k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_75,negated_conjecture,(k1_funct_1(k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_59]), c_0_70])).
cnf(c_0_76,negated_conjecture,(~r2_hidden(k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(rw,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_78,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
fof(c_0_79, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_80,negated_conjecture,(~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_34])).
cnf(c_0_81,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_82,negated_conjecture,(~l1_pre_topc(esk1_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_83,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_44]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 84
# Proof object clause steps            : 57
# Proof object formula steps           : 27
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 33
# Processed clauses                    : 123
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 123
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 15
# Generated clauses                    : 95
# ...of the previous two non-trivial   : 95
# Contextual simplify-reflections      : 3
# Paramodulations                      : 94
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 74
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 30
# ...number of literals in the above   : 90
# Current number of archived formulas  : 0
# Current number of archived clauses   : 48
# Clause-clause subsumption calls (NU) : 1036
# Rec. Clause-clause subsumption calls : 253
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 50
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7246

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45368 pages
