# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t26_tops_2, conjecture, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)=>v3_pre_topc(k5_setfam_1(u1_struct_0(X1),X2),X1))))), file('out/t26_tops_2__7_premises__rerun.E_input', t26_tops_2)).
fof(d1_tops_2, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,X2)=>v3_pre_topc(X3,X1))))))), file('out/t26_tops_2__7_premises__rerun.E_input', d1_tops_2)).
fof(d1_zfmisc_1, axiom, (![X1]:![X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1)))), file('out/t26_tops_2__7_premises__rerun.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t26_tops_2__7_premises__rerun.E_input', d3_tarski)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t26_tops_2__7_premises__rerun.E_input', t3_subset)).
fof(d5_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1)))))), file('out/t26_tops_2__7_premises__rerun.E_input', d5_pre_topc)).
fof(dt_k5_setfam_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>m1_subset_1(k5_setfam_1(X1,X2),k1_zfmisc_1(X1)))), file('out/t26_tops_2__7_premises__rerun.E_input', dt_k5_setfam_1)).
fof(d1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1)))))))))), file('out/t26_tops_2__7_premises__rerun.E_input', d1_pre_topc)).
fof(c_0_8, negated_conjecture, (~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)=>v3_pre_topc(k5_setfam_1(u1_struct_0(X1),X2),X1)))))), inference(assume_negation,[status(cth)],[t26_tops_2])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:((((~v1_tops_2(X5,X4)|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(~r2_hidden(X6,X5)|v3_pre_topc(X6,X4))))|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~l1_pre_topc(X4))&((((m1_subset_1(esk5_2(X4,X5),k1_zfmisc_1(u1_struct_0(X4)))|v1_tops_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~l1_pre_topc(X4))&((((r2_hidden(esk5_2(X4,X5),X5)|v1_tops_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~l1_pre_topc(X4))&(((~v3_pre_topc(esk5_2(X4,X5),X4)|v1_tops_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tops_2])])])])])])])).
fof(c_0_10, negated_conjecture, (((v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk3_0))))&(v1_tops_2(esk4_0,esk3_0)&~v3_pre_topc(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),esk3_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|r1_tarski(X6,X4))|X5!=k1_zfmisc_1(X4))&((~r1_tarski(X6,X4)|r2_hidden(X6,X5))|X5!=k1_zfmisc_1(X4)))&(((~r2_hidden(esk1_2(X4,X5),X5)|~r1_tarski(esk1_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))&((r2_hidden(esk1_2(X4,X5),X5)|r1_tarski(esk1_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk2_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk2_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_13, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_14, plain, (![X3]:![X4]:((((~v3_pre_topc(X4,X3)|r2_hidden(X4,u1_pre_topc(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3))&(((~r2_hidden(X4,u1_pre_topc(X3))|v3_pre_topc(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])])])).
cnf(c_0_15,plain,(v3_pre_topc(X3,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~r2_hidden(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,negated_conjecture,(v1_tops_2(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,negated_conjecture,(l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk3_0))))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,plain,(r1_tarski(X3,X2)|X1!=k1_zfmisc_1(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,plain,(r2_hidden(X2,u1_pre_topc(X1))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,negated_conjecture,(v3_pre_topc(X1,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_24,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,plain,(r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(r2_hidden(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,negated_conjecture,(~v3_pre_topc(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_28,plain,(v3_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,u1_pre_topc(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_29, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|m1_subset_1(k5_setfam_1(X3,X4),k1_zfmisc_1(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_setfam_1])])).
cnf(c_0_30,negated_conjecture,(r2_hidden(X1,u1_pre_topc(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_17])])).
cnf(c_0_31,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32,negated_conjecture,(r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_26, c_0_18])).
cnf(c_0_33,negated_conjecture,(~m1_subset_1(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),u1_pre_topc(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17])])).
cnf(c_0_34,plain,(m1_subset_1(k5_setfam_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, (![X4]:![X5]:![X6]:![X7]:((((r2_hidden(u1_struct_0(X4),u1_pre_topc(X4))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&((((~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|(~r1_tarski(X5,u1_pre_topc(X4))|r2_hidden(k5_setfam_1(u1_struct_0(X4),X5),u1_pre_topc(X4))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&(((~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|((~r2_hidden(X6,u1_pre_topc(X4))|~r2_hidden(X7,u1_pre_topc(X4)))|r2_hidden(k9_subset_1(u1_struct_0(X4),X6,X7),u1_pre_topc(X4)))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))))&((((((m1_subset_1(esk7_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk6_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk8_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk6_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk7_1(X4),u1_pre_topc(X4))|m1_subset_1(esk6_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk8_1(X4),u1_pre_topc(X4))|m1_subset_1(esk6_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk7_1(X4),esk8_1(X4)),u1_pre_topc(X4))|m1_subset_1(esk6_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&((((((m1_subset_1(esk7_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk6_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk8_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk6_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk7_1(X4),u1_pre_topc(X4))|r1_tarski(esk6_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk8_1(X4),u1_pre_topc(X4))|r1_tarski(esk6_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk7_1(X4),esk8_1(X4)),u1_pre_topc(X4))|r1_tarski(esk6_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&(((((m1_subset_1(esk7_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk6_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk8_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk6_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk7_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk6_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk8_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk6_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk7_1(X4),esk8_1(X4)),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk6_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])])])).
cnf(c_0_36,plain,(r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37,negated_conjecture,(r2_hidden(X1,u1_pre_topc(esk3_0))|~r2_hidden(X1,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_38,negated_conjecture,(~r2_hidden(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),u1_pre_topc(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_18])])).
cnf(c_0_39,plain,(r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r1_tarski(X2,u1_pre_topc(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40,negated_conjecture,(v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_41,negated_conjecture,(r1_tarski(X1,u1_pre_topc(esk3_0))|~r2_hidden(esk2_2(X1,u1_pre_topc(esk3_0)),esk4_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42,plain,(r1_tarski(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_43,negated_conjecture,(~r1_tarski(esk4_0,u1_pre_topc(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_17]), c_0_40]), c_0_18])])).
cnf(c_0_44,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 28
# Proof object formula steps           : 17
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 104
# ...of these trivial                  : 1
# ...subsumed                          : 5
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 80
# ...of the previous two non-trivial   : 60
# Contextual simplify-reflections      : 1
# Paramodulations                      : 78
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 54
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 143
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 845
# Rec. Clause-clause subsumption calls : 106
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5220

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45492 pages
