# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_mcart_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1))))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', t23_mcart_1)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', t2_subset)).
fof(t26_mcart_1, conjecture, (![X1]:![X2]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~(![X3]:(m1_subset_1(X3,k2_zfmisc_1(X1,X2))=>(~(X3=k1_mcart_1(X3))&~(X3=k2_mcart_1(X3))))))))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', t26_mcart_1)).
fof(t20_mcart_1, axiom, (![X1]:(?[X2]:?[X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1))))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', t20_mcart_1)).
fof(fc6_relat_1, axiom, (![X1]:![X2]:v1_relat_1(k2_zfmisc_1(X1,X2))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', fc6_relat_1)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', d2_subset_1)).
fof(fc1_zfmisc_1, axiom, (![X1]:![X2]:~(v1_xboole_0(k4_tarski(X1,X2)))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', fc1_zfmisc_1)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', l13_xboole_0)).
fof(fc10_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))=>~(v1_xboole_0(k2_zfmisc_1(X1,X2))))), file('out/t26_mcart_1__9_premises__rerun__rerun.E_input', fc10_subset_1)).
fof(c_0_9, plain, (![X3]:![X4]:(~v1_relat_1(X4)|(~r2_hidden(X3,X4)|X3=k4_tarski(k1_mcart_1(X3),k2_mcart_1(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_mcart_1])])).
fof(c_0_10, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~(![X3]:(m1_subset_1(X3,k2_zfmisc_1(X1,X2))=>(~(X3=k1_mcart_1(X3))&~(X3=k2_mcart_1(X3)))))))))), inference(assume_negation,[status(cth)],[t26_mcart_1])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:((X4!=k1_mcart_1(X4)|X4!=k4_tarski(X5,X6))&(X4!=k2_mcart_1(X4)|X4!=k4_tarski(X5,X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])])])).
cnf(c_0_13,plain,(X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1))|~r2_hidden(X1,X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, negated_conjecture, ((esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(m1_subset_1(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))&(esk3_0=k1_mcart_1(esk3_0)|esk3_0=k2_mcart_1(esk3_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])).
fof(c_0_16, plain, (![X3]:![X4]:v1_relat_1(k2_zfmisc_1(X3,X4))), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_17, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
cnf(c_0_18,plain,(X1!=k4_tarski(X2,X3)|X1!=k2_mcart_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1))=X1|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20,negated_conjecture,(m1_subset_1(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(X1!=k4_tarski(X2,X3)|X1!=k1_mcart_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,plain,(v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,plain,(k2_mcart_1(k4_tarski(X1,X2))!=k4_tarski(X1,X2)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_25,negated_conjecture,(k4_tarski(k1_mcart_1(esk3_0),k2_mcart_1(esk3_0))=esk3_0|v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_26,plain,(k1_mcart_1(k4_tarski(X1,X2))!=k4_tarski(X1,X2)), inference(er,[status(thm)],[c_0_22])).
fof(c_0_27, plain, (![X3]:![X4]:~v1_xboole_0(k4_tarski(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_zfmisc_1])])).
cnf(c_0_28,negated_conjecture,(v1_xboole_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_23, c_0_20])).
cnf(c_0_29,negated_conjecture,(v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|k2_mcart_1(esk3_0)!=esk3_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30,negated_conjecture,(v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|k1_mcart_1(esk3_0)!=esk3_0), inference(spm,[status(thm)],[c_0_26, c_0_25])).
fof(c_0_31, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_32,plain,(~v1_xboole_0(k4_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,negated_conjecture,(v1_xboole_0(esk3_0)|k2_mcart_1(esk3_0)!=esk3_0), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34,negated_conjecture,(esk3_0=k2_mcart_1(esk3_0)|esk3_0=k1_mcart_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_35,negated_conjecture,(v1_xboole_0(esk3_0)|k1_mcart_1(esk3_0)!=esk3_0), inference(spm,[status(thm)],[c_0_28, c_0_30])).
cnf(c_0_36,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37,negated_conjecture,(v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_25])).
cnf(c_0_38,negated_conjecture,(v1_xboole_0(esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])).
fof(c_0_39, plain, (![X3]:![X4]:((v1_xboole_0(X3)|v1_xboole_0(X4))|~v1_xboole_0(k2_zfmisc_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc10_subset_1])])])).
cnf(c_0_40,negated_conjecture,(k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41,negated_conjecture,(k1_xboole_0=esk3_0), inference(spm,[status(thm)],[c_0_36, c_0_38])).
cnf(c_0_42,plain,(v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_43,negated_conjecture,(k2_zfmisc_1(esk1_0,esk2_0)=esk3_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_38])]), c_0_41])).
cnf(c_0_44,negated_conjecture,(esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_45,plain,(X1=esk3_0|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_36, c_0_41])).
cnf(c_0_46,negated_conjecture,(v1_xboole_0(esk1_0)|v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_38])])).
cnf(c_0_47,negated_conjecture,(esk2_0!=esk3_0), inference(rw,[status(thm)],[c_0_44, c_0_41])).
cnf(c_0_48,negated_conjecture,(esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_49,negated_conjecture,(v1_xboole_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_50,negated_conjecture,(esk1_0!=esk3_0), inference(rw,[status(thm)],[c_0_48, c_0_41])).
cnf(c_0_51,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_49]), c_0_50]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 33
# Proof object formula steps           : 19
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 71
# ...of these trivial                  : 4
# ...subsumed                          : 5
# ...remaining for further processing  : 62
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 22
# Generated clauses                    : 53
# ...of the previous two non-trivial   : 53
# Contextual simplify-reflections      : 7
# Paramodulations                      : 51
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 22
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 49
# Rec. Clause-clause subsumption calls : 48
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 36
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1395

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45248 pages
