# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t27_yellow_1, conjecture, (![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), file('out/t27_yellow_1__11_premises__rerun.E_input', t27_yellow_1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t27_yellow_1__11_premises__rerun.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t27_yellow_1__11_premises__rerun.E_input', t70_enumset1)).
fof(t26_yellow_1, axiom, (![X1]:((v1_relat_1(X1)&(v4_relat_1(X1,k1_xboole_0)&(v1_funct_1(X1)&(v1_partfun1(X1,k1_xboole_0)&v1_yellow_1(X1)))))=>k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), file('out/t27_yellow_1__11_premises__rerun.E_input', t26_yellow_1)).
fof(fc11_funct_2, axiom, (![X1]:(v1_relat_1(k6_relat_1(X1))&(v4_relat_1(k6_relat_1(X1),X1)&(v1_funct_1(k6_relat_1(X1))&v1_partfun1(k6_relat_1(X1),X1))))), file('out/t27_yellow_1__11_premises__rerun.E_input', fc11_funct_2)).
fof(redefinition_k6_partfun1, axiom, (![X1]:k6_partfun1(X1)=k6_relat_1(X1)), file('out/t27_yellow_1__11_premises__rerun.E_input', redefinition_k6_partfun1)).
fof(d5_yellow_1, axiom, (![X1]:![X2]:(l1_orders_2(X2)=>k6_yellow_1(X1,X2)=k5_yellow_1(X1,k7_funcop_1(X1,X2)))), file('out/t27_yellow_1__11_premises__rerun.E_input', d5_yellow_1)).
fof(redefinition_k7_funcop_1, axiom, (![X1]:![X2]:k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2)), file('out/t27_yellow_1__11_premises__rerun.E_input', redefinition_k7_funcop_1)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t27_yellow_1__11_premises__rerun.E_input', l13_xboole_0)).
fof(fc2_funcop_1, axiom, (![X1]:v1_xboole_0(k2_funcop_1(k1_xboole_0,X1))), file('out/t27_yellow_1__11_premises__rerun.E_input', fc2_funcop_1)).
fof(t81_relat_1, axiom, (k6_relat_1(k1_xboole_0)=k1_xboole_0), file('out/t27_yellow_1__11_premises__rerun.E_input', t81_relat_1)).
fof(fc10_yellow_1, axiom, (![X1]:![X2]:(l1_orders_2(X2)=>v1_yellow_1(k2_funcop_1(X1,X2)))), file('out/t27_yellow_1__11_premises__rerun.E_input', fc10_yellow_1)).
fof(c_0_12, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t27_yellow_1])).
fof(c_0_13, negated_conjecture, ((l1_orders_2(esk1_0)&k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_14, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_15, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, (![X2]:((~v1_relat_1(X2)|(~v4_relat_1(X2,k1_xboole_0)|(~v1_funct_1(X2)|(~v1_partfun1(X2,k1_xboole_0)|~v1_yellow_1(X2)))))|k5_yellow_1(k1_xboole_0,X2)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_yellow_1])])).
fof(c_0_17, plain, (![X2]:![X2]:![X2]:![X2]:(v1_relat_1(k6_relat_1(X2))&(v4_relat_1(k6_relat_1(X2),X2)&(v1_funct_1(k6_relat_1(X2))&v1_partfun1(k6_relat_1(X2),X2))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc11_funct_2])])])).
fof(c_0_18, plain, (![X2]:k6_partfun1(X2)=k6_relat_1(X2)), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
cnf(c_0_19,negated_conjecture,(k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))|~v1_yellow_1(X1)|~v1_partfun1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, (![X3]:![X4]:(~l1_orders_2(X4)|k6_yellow_1(X3,X4)=k5_yellow_1(X3,k7_funcop_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_1])])).
fof(c_0_24, plain, (![X3]:![X4]:k7_funcop_1(X3,X4)=k2_funcop_1(X3,X4)), inference(variable_rename,[status(thm)],[redefinition_k7_funcop_1])).
fof(c_0_25, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
fof(c_0_26, plain, (![X2]:v1_xboole_0(k2_funcop_1(k1_xboole_0,X2))), inference(variable_rename,[status(thm)],[fc2_funcop_1])).
cnf(c_0_27,plain,(v1_partfun1(k6_relat_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28,plain,(k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29,plain,(k6_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t81_relat_1])).
cnf(c_0_30,plain,(v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31,plain,(v4_relat_1(k6_relat_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32,plain,(v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_33,negated_conjecture,(k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0),k6_partfun1(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_20]), c_0_21]), c_0_21])).
cnf(c_0_34,plain,(k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0),k6_partfun1(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_yellow_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_partfun1(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_20]), c_0_20]), c_0_21]), c_0_21])).
cnf(c_0_35,plain,(k6_yellow_1(X1,X2)=k5_yellow_1(X1,k7_funcop_1(X1,X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36,plain,(k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38,plain,(v1_xboole_0(k2_funcop_1(k1_xboole_0,X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39,plain,(v1_partfun1(k6_partfun1(X1),X1)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_40,plain,(k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_41,plain,(v1_funct_1(k6_partfun1(X1))), inference(rw,[status(thm)],[c_0_30, c_0_28])).
cnf(c_0_42,plain,(v4_relat_1(k6_partfun1(X1),X1)), inference(rw,[status(thm)],[c_0_31, c_0_28])).
cnf(c_0_43,plain,(v1_relat_1(k6_partfun1(X1))), inference(rw,[status(thm)],[c_0_32, c_0_28])).
cnf(c_0_44,negated_conjecture,(k5_yellow_1(k1_xboole_0,X1)!=k6_yellow_1(k1_xboole_0,esk1_0)|~v1_yellow_1(X1)|~v1_partfun1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_45,plain,(k6_yellow_1(X1,X2)=k5_yellow_1(X1,k2_funcop_1(X1,X2))|~l1_orders_2(X2)), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_46,plain,(k2_funcop_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47,plain,(v1_partfun1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48,plain,(v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_40])).
cnf(c_0_49,plain,(v4_relat_1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_42, c_0_40])).
cnf(c_0_50,plain,(v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_43, c_0_40])).
fof(c_0_51, plain, (![X3]:![X4]:(~l1_orders_2(X4)|v1_yellow_1(k2_funcop_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc10_yellow_1])])).
cnf(c_0_52,negated_conjecture,(k6_yellow_1(k1_xboole_0,X1)!=k6_yellow_1(k1_xboole_0,esk1_0)|~v1_yellow_1(k1_xboole_0)|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_46]), c_0_48]), c_0_46]), c_0_49]), c_0_46]), c_0_50])])).
cnf(c_0_53,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,plain,(v1_yellow_1(k2_funcop_1(X1,X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_55,negated_conjecture,(~v1_yellow_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_52]), c_0_53])])).
cnf(c_0_56,plain,(~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_46]), c_0_55])).
cnf(c_0_57,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_53, c_0_56]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 34
# Proof object formula steps           : 24
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 12
# Processed clauses                    : 35
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 34
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 14
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 0
# Paramodulations                      : 12
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 20
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 6
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 16
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 29
# Rec. Clause-clause subsumption calls : 5
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 7
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1048

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45516 pages
