# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2m
# and selection function SelectCQArNTNpEqFirst.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t134_pboole, axiom, (![X1]:((v1_relat_1(X1)&(v4_relat_1(X1,k1_xboole_0)&(v1_funct_1(X1)&v1_partfun1(X1,k1_xboole_0))))=>X1=k1_xboole_0)), file('out/t27_yellow_1__7_premises__rerun.E_input', t134_pboole)).
fof(rc1_yellow_1, axiom, (![X1]:?[X2]:(v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_partfun1(X2,X1)&v1_yellow_1(X2)))))), file('out/t27_yellow_1__7_premises__rerun.E_input', rc1_yellow_1)).
fof(t27_yellow_1, conjecture, (![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), file('out/t27_yellow_1__7_premises__rerun.E_input', t27_yellow_1)).
fof(t26_yellow_1, axiom, (![X1]:((v1_relat_1(X1)&(v4_relat_1(X1,k1_xboole_0)&(v1_funct_1(X1)&(v1_partfun1(X1,k1_xboole_0)&v1_yellow_1(X1)))))=>k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), file('out/t27_yellow_1__7_premises__rerun.E_input', t26_yellow_1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t27_yellow_1__7_premises__rerun.E_input', t69_enumset1)).
fof(d5_yellow_1, axiom, (![X1]:![X2]:(l1_orders_2(X2)=>k6_yellow_1(X1,X2)=k5_yellow_1(X1,k7_funcop_1(X1,X2)))), file('out/t27_yellow_1__7_premises__rerun.E_input', d5_yellow_1)).
fof(redefinition_k7_funcop_1, axiom, (![X1]:![X2]:k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2)), file('out/t27_yellow_1__7_premises__rerun.E_input', redefinition_k7_funcop_1)).
fof(fc20_funcop_1, axiom, (![X1]:![X2]:(v1_relat_1(k2_funcop_1(X1,X2))&(v4_relat_1(k2_funcop_1(X1,X2),X1)&(v1_funct_1(k2_funcop_1(X1,X2))&v1_partfun1(k2_funcop_1(X1,X2),X1))))), file('out/t27_yellow_1__7_premises__rerun.E_input', fc20_funcop_1)).
fof(c_0_8, plain, (![X2]:((~v1_relat_1(X2)|(~v4_relat_1(X2,k1_xboole_0)|(~v1_funct_1(X2)|~v1_partfun1(X2,k1_xboole_0))))|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t134_pboole])])).
fof(c_0_9, plain, (![X3]:(v1_relat_1(esk1_1(X3))&(v4_relat_1(esk1_1(X3),X3)&(v1_funct_1(esk1_1(X3))&(v1_partfun1(esk1_1(X3),X3)&v1_yellow_1(esk1_1(X3))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_yellow_1])])).
fof(c_0_10, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t27_yellow_1])).
fof(c_0_11, plain, (![X2]:((~v1_relat_1(X2)|(~v4_relat_1(X2,k1_xboole_0)|(~v1_funct_1(X2)|(~v1_partfun1(X2,k1_xboole_0)|~v1_yellow_1(X2)))))|k5_yellow_1(k1_xboole_0,X2)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_yellow_1])])).
fof(c_0_12, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_13,plain,(X1=k1_xboole_0|~v1_partfun1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(v1_partfun1(esk1_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(v1_funct_1(esk1_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,plain,(v4_relat_1(esk1_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,plain,(v1_relat_1(esk1_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_18, negated_conjecture, ((l1_orders_2(esk2_0)&k6_yellow_1(k1_xboole_0,esk2_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_19,plain,(k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))|~v1_yellow_1(X1)|~v1_partfun1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,plain,(v1_yellow_1(esk1_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22,plain,(esk1_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17])])).
fof(c_0_23, plain, (![X3]:![X4]:(~l1_orders_2(X4)|k6_yellow_1(X3,X4)=k5_yellow_1(X3,k7_funcop_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_1])])).
fof(c_0_24, plain, (![X3]:![X4]:k7_funcop_1(X3,X4)=k2_funcop_1(X3,X4)), inference(variable_rename,[status(thm)],[redefinition_k7_funcop_1])).
fof(c_0_25, plain, (![X3]:![X4]:![X3]:![X4]:![X3]:![X4]:![X3]:![X4]:(v1_relat_1(k2_funcop_1(X3,X4))&(v4_relat_1(k2_funcop_1(X3,X4),X3)&(v1_funct_1(k2_funcop_1(X3,X4))&v1_partfun1(k2_funcop_1(X3,X4),X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc20_funcop_1])])])).
cnf(c_0_26,negated_conjecture,(k6_yellow_1(k1_xboole_0,esk2_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k2_tarski(k1_xboole_0,k1_xboole_0),k6_partfun1(k2_tarski(k1_xboole_0,k1_xboole_0)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_yellow_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_partfun1(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_20])).
cnf(c_0_28,plain,(v1_yellow_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29,plain,(k6_yellow_1(X1,X2)=k5_yellow_1(X1,k7_funcop_1(X1,X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(v1_partfun1(k2_funcop_1(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,plain,(v1_funct_1(k2_funcop_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,plain,(v4_relat_1(k2_funcop_1(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(v1_relat_1(k2_funcop_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35,negated_conjecture,(k6_yellow_1(k1_xboole_0,esk2_0)!=g1_orders_2(k2_tarski(k1_xboole_0,k1_xboole_0),k6_partfun1(k2_tarski(k1_xboole_0,k1_xboole_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_20]), c_0_20])).
cnf(c_0_36,plain,(g1_orders_2(k2_tarski(k1_xboole_0,k1_xboole_0),k6_partfun1(k2_tarski(k1_xboole_0,k1_xboole_0)))=k5_yellow_1(k1_xboole_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_14]), c_0_22]), c_0_22]), c_0_28]), c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_37,plain,(k6_yellow_1(X1,X2)=k5_yellow_1(X1,k2_funcop_1(X1,X2))|~l1_orders_2(X2)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38,plain,(k2_funcop_1(k1_xboole_0,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_31]), c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_39,negated_conjecture,(k5_yellow_1(k1_xboole_0,k1_xboole_0)!=k6_yellow_1(k1_xboole_0,esk2_0)), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40,plain,(k5_yellow_1(k1_xboole_0,k1_xboole_0)=k6_yellow_1(k1_xboole_0,X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_41,negated_conjecture,(k6_yellow_1(k1_xboole_0,X1)!=k6_yellow_1(k1_xboole_0,esk2_0)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_42,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_43,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_41]), c_0_42])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 27
# Proof object formula steps           : 17
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 14
# Processed clauses                    : 41
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 41
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 31
# ...of the previous two non-trivial   : 21
# Contextual simplify-reflections      : 0
# Paramodulations                      : 30
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 6
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 22
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 15
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1329

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45516 pages
