# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_connsp_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0)))))), file('out/t28_connsp_2__11_premises__rerun.E_input', t28_connsp_2)).
fof(t113_zfmisc_1, axiom, (![X1]:![X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), file('out/t28_connsp_2__11_premises__rerun.E_input', t113_zfmisc_1)).
fof(t152_zfmisc_1, axiom, (![X1]:![X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2)))), file('out/t28_connsp_2__11_premises__rerun.E_input', t152_zfmisc_1)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t28_connsp_2__11_premises__rerun.E_input', t2_subset)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t28_connsp_2__11_premises__rerun.E_input', fc2_struct_0)).
fof(fc4_pre_topc, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>v4_pre_topc(k2_struct_0(X1),X1))), file('out/t28_connsp_2__11_premises__rerun.E_input', fc4_pre_topc)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t28_connsp_2__11_premises__rerun.E_input', dt_l1_pre_topc)).
fof(dt_k2_subset_1, axiom, (![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), file('out/t28_connsp_2__11_premises__rerun.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, (![X1]:k2_subset_1(X1)=X1), file('out/t28_connsp_2__11_premises__rerun.E_input', d4_subset_1)).
fof(d3_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1))), file('out/t28_connsp_2__11_premises__rerun.E_input', d3_struct_0)).
fof(d1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1)))))))))), file('out/t28_connsp_2__11_premises__rerun.E_input', d1_pre_topc)).
fof(d5_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1)))))), file('out/t28_connsp_2__11_premises__rerun.E_input', d5_pre_topc)).
fof(c_0_12, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0))))))), inference(assume_negation,[status(cth)],[t28_connsp_2])).
fof(c_0_13, plain, (![X3]:![X4]:![X3]:![X4]:((k2_zfmisc_1(X3,X4)!=k1_xboole_0|(X3=k1_xboole_0|X4=k1_xboole_0))&((X3!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)&(X4!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])])])).
fof(c_0_14, negated_conjecture, (![X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))&(((((v3_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&(((v4_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&((r2_hidden(esk2_0,X8)|~r2_hidden(X8,esk3_0))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))))&(((~v3_pre_topc(X8,esk1_0)|(~v4_pre_topc(X8,esk1_0)|~r2_hidden(esk2_0,X8)))|r2_hidden(X8,esk3_0))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0)))))&esk3_0=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])])])).
fof(c_0_15, plain, (![X3]:![X4]:~r2_hidden(X3,k2_zfmisc_1(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_16,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_18,negated_conjecture,(r2_hidden(X1,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk2_0,X1)|~v4_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,negated_conjecture,(esk3_0=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_16])).
fof(c_0_22, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_23,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25,negated_conjecture,(r2_hidden(X1,k1_xboole_0)|~v4_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26,plain,(~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_27, plain, (![X2]:((~v2_pre_topc(X2)|~l1_pre_topc(X2))|v4_pre_topc(k2_struct_0(X2),X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_pre_topc])])).
cnf(c_0_28,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_31, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_32, plain, (![X2]:m1_subset_1(k2_subset_1(X2),k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_33, plain, (![X2]:k2_subset_1(X2)=X2), inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_34,negated_conjecture,(~v4_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35,plain,(v4_pre_topc(k2_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_38, plain, (![X2]:(~l1_struct_0(X2)|k2_struct_0(X2)=u1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_39,negated_conjecture,(r2_hidden(esk2_0,u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_40,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41,plain,(m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42,plain,(k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_43, plain, (![X4]:![X5]:![X6]:![X7]:((((r2_hidden(u1_struct_0(X4),u1_pre_topc(X4))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&((((~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|(~r1_tarski(X5,u1_pre_topc(X4))|r2_hidden(k5_setfam_1(u1_struct_0(X4),X5),u1_pre_topc(X4))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&(((~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|((~r2_hidden(X6,u1_pre_topc(X4))|~r2_hidden(X7,u1_pre_topc(X4)))|r2_hidden(k9_subset_1(u1_struct_0(X4),X6,X7),u1_pre_topc(X4)))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))))&((((((m1_subset_1(esk5_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk4_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk6_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk4_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk5_1(X4),u1_pre_topc(X4))|m1_subset_1(esk4_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk6_1(X4),u1_pre_topc(X4))|m1_subset_1(esk4_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk5_1(X4),esk6_1(X4)),u1_pre_topc(X4))|m1_subset_1(esk4_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&((((((m1_subset_1(esk5_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk4_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk6_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk4_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk5_1(X4),u1_pre_topc(X4))|r1_tarski(esk4_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk6_1(X4),u1_pre_topc(X4))|r1_tarski(esk4_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk5_1(X4),esk6_1(X4)),u1_pre_topc(X4))|r1_tarski(esk4_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&(((((m1_subset_1(esk5_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk4_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk6_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk4_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk5_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk4_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk6_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk4_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk5_1(X4),esk6_1(X4)),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk4_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])])])).
cnf(c_0_44,negated_conjecture,(~v3_pre_topc(k2_struct_0(esk1_0),esk1_0)|~r2_hidden(esk2_0,k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])])).
cnf(c_0_45,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46,negated_conjecture,(r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_36])])).
cnf(c_0_47,plain,(m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_48, plain, (![X3]:![X4]:((((~v3_pre_topc(X4,X3)|r2_hidden(X4,u1_pre_topc(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3))&(((~r2_hidden(X4,u1_pre_topc(X3))|v3_pre_topc(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])])])).
cnf(c_0_49,plain,(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50,negated_conjecture,(~l1_struct_0(esk1_0)|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_47])])).
cnf(c_0_51,plain,(v3_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,u1_pre_topc(X1))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_52,negated_conjecture,(r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_37]), c_0_36])])).
cnf(c_0_53,negated_conjecture,(~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_47]), c_0_36])])).
cnf(c_0_54,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_40]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 30
# Proof object formula steps           : 25
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 21
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 40
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 72
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 17
# ...of the previous two non-trivial   : 15
# Contextual simplify-reflections      : 0
# Paramodulations                      : 15
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 94
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 549
# Rec. Clause-clause subsumption calls : 88
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3686

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.004 s
# Total time               : 0.008 s
# Maximum resident set size: 45484 pages
