# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d13_xtuple_0, axiom, (![X1]:![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1)))), file('out/t28_relat_1__6_premises__rerun__rerun.E_input', d13_xtuple_0)).
fof(d8_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(X3=k3_relat_1(X1,X2)<=>![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2)))))), file('out/t28_relat_1__6_premises__rerun__rerun.E_input', d8_relat_1)).
fof(dt_k3_relat_1, axiom, (![X1]:![X2]:v1_relat_1(k3_relat_1(X1,X2))), file('out/t28_relat_1__6_premises__rerun__rerun.E_input', dt_k3_relat_1)).
fof(t28_relat_1, conjecture, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(k9_xtuple_0(X1),k10_xtuple_0(X2))=>k10_xtuple_0(k3_relat_1(X2,X1))=k10_xtuple_0(X1))))), file('out/t28_relat_1__6_premises__rerun__rerun.E_input', t28_relat_1)).
fof(d12_xtuple_0, axiom, (![X1]:![X2]:(X2=k9_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t28_relat_1__6_premises__rerun__rerun.E_input', d12_xtuple_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t28_relat_1__6_premises__rerun__rerun.E_input', d3_tarski)).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(esk1_3(X5,X6,X7),X7),X5))|X6!=k10_xtuple_0(X5))&((~r2_hidden(k4_tarski(X9,X7),X5)|r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5)))&(((~r2_hidden(esk2_2(X5,X6),X6)|~r2_hidden(k4_tarski(X11,esk2_2(X5,X6)),X5))|X6=k10_xtuple_0(X5))&((r2_hidden(esk2_2(X5,X6),X6)|r2_hidden(k4_tarski(esk3_2(X5,X6),esk2_2(X5,X6)),X5))|X6=k10_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_xtuple_0])])])])])])])).
fof(c_0_7, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X10]:![X11]:![X13]:![X16]:((((((r2_hidden(k4_tarski(X10,esk4_5(X7,X8,X9,X10,X11)),X7)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9))&(((r2_hidden(k4_tarski(esk4_5(X7,X8,X9,X10,X11),X11),X8)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9)))&((((~r2_hidden(k4_tarski(X10,X13),X7)|~r2_hidden(k4_tarski(X13,X11),X8))|r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9)))&((((~r2_hidden(k4_tarski(esk5_3(X7,X8,X9),esk6_3(X7,X8,X9)),X9)|(~r2_hidden(k4_tarski(esk5_3(X7,X8,X9),X16),X7)|~r2_hidden(k4_tarski(X16,esk6_3(X7,X8,X9)),X8)))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9))&((((r2_hidden(k4_tarski(esk5_3(X7,X8,X9),esk7_3(X7,X8,X9)),X7)|r2_hidden(k4_tarski(esk5_3(X7,X8,X9),esk6_3(X7,X8,X9)),X9))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9))&(((r2_hidden(k4_tarski(esk7_3(X7,X8,X9),esk6_3(X7,X8,X9)),X8)|r2_hidden(k4_tarski(esk5_3(X7,X8,X9),esk6_3(X7,X8,X9)),X9))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])])).
fof(c_0_8, plain, (![X3]:![X4]:v1_relat_1(k3_relat_1(X3,X4))), inference(variable_rename,[status(thm)],[dt_k3_relat_1])).
cnf(c_0_9,plain,(r2_hidden(X3,X1)|X1!=k10_xtuple_0(X2)|~r2_hidden(k4_tarski(X4,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(r2_hidden(k4_tarski(esk4_5(X2,X3,X1,X4,X5),X5),X3)|~v1_relat_1(X1)|X1!=k3_relat_1(X2,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(v1_relat_1(k3_relat_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(r2_hidden(X1,k10_xtuple_0(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(r2_hidden(k4_tarski(esk4_5(X1,X2,k3_relat_1(X1,X2),X3,X4),X4),X2)|~r2_hidden(k4_tarski(X3,X4),k3_relat_1(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])])).
cnf(c_0_14,plain,(r2_hidden(k4_tarski(esk1_3(X2,X1,X3),X3),X2)|X1!=k10_xtuple_0(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,plain,(r2_hidden(k4_tarski(X4,X5),X1)|~v1_relat_1(X1)|X1!=k3_relat_1(X2,X3)|~r2_hidden(k4_tarski(X6,X5),X3)|~r2_hidden(k4_tarski(X4,X6),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_16, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(k9_xtuple_0(X1),k10_xtuple_0(X2))=>k10_xtuple_0(k3_relat_1(X2,X1))=k10_xtuple_0(X1)))))), inference(assume_negation,[status(cth)],[t28_relat_1])).
fof(c_0_17, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk10_3(X5,X6,X7)),X5))|X6!=k9_xtuple_0(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k9_xtuple_0(X5)))&(((~r2_hidden(esk11_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk11_2(X5,X6),X11),X5))|X6=k9_xtuple_0(X5))&((r2_hidden(esk11_2(X5,X6),X6)|r2_hidden(k4_tarski(esk11_2(X5,X6),esk12_2(X5,X6)),X5))|X6=k9_xtuple_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_xtuple_0])])])])])])])).
cnf(c_0_18,plain,(r2_hidden(X1,k10_xtuple_0(X2))|~r2_hidden(k4_tarski(X3,X1),k3_relat_1(X4,X2))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19,plain,(r2_hidden(k4_tarski(esk1_3(X1,k10_xtuple_0(X1),X2),X2),X1)|~r2_hidden(X2,k10_xtuple_0(X1))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(X1=k10_xtuple_0(X2)|r2_hidden(k4_tarski(esk3_2(X2,X1),esk2_2(X2,X1)),X2)|r2_hidden(esk2_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21,plain,(r2_hidden(k4_tarski(X1,X2),k3_relat_1(X3,X4))|~r2_hidden(k4_tarski(X5,X2),X4)|~r2_hidden(k4_tarski(X1,X5),X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_11])])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk13_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk13_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_23, negated_conjecture, ((v1_relat_1(esk8_0)&(v1_relat_1(esk9_0)&(r1_tarski(k9_xtuple_0(esk8_0),k10_xtuple_0(esk9_0))&k10_xtuple_0(k3_relat_1(esk9_0,esk8_0))!=k10_xtuple_0(esk8_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).
cnf(c_0_24,plain,(r2_hidden(X3,X1)|X1!=k9_xtuple_0(X2)|~r2_hidden(k4_tarski(X3,X4),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(r2_hidden(X1,k10_xtuple_0(X2))|~r2_hidden(X1,k10_xtuple_0(k3_relat_1(X3,X2)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26,plain,(X1=k10_xtuple_0(X2)|r2_hidden(esk2_2(X2,X1),k10_xtuple_0(X2))|r2_hidden(esk2_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_12, c_0_20])).
cnf(c_0_27,plain,(r2_hidden(k4_tarski(X1,X2),k3_relat_1(X3,X4))|~r2_hidden(k4_tarski(X1,esk1_3(X4,k10_xtuple_0(X4),X2)),X3)|~r2_hidden(X2,k10_xtuple_0(X4))), inference(spm,[status(thm)],[c_0_21, c_0_19])).
cnf(c_0_28,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,negated_conjecture,(r1_tarski(k9_xtuple_0(esk8_0),k10_xtuple_0(esk9_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(X1=k10_xtuple_0(X2)|~r2_hidden(k4_tarski(X3,esk2_2(X2,X1)),X2)|~r2_hidden(esk2_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_32,plain,(k10_xtuple_0(k3_relat_1(X1,X2))=k10_xtuple_0(X3)|r2_hidden(esk2_2(X3,k10_xtuple_0(k3_relat_1(X1,X2))),k10_xtuple_0(X3))|r2_hidden(esk2_2(X3,k10_xtuple_0(k3_relat_1(X1,X2))),k10_xtuple_0(X2))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33,plain,(r2_hidden(k4_tarski(esk1_3(X1,k10_xtuple_0(X1),esk1_3(X2,k10_xtuple_0(X2),X3)),X3),k3_relat_1(X1,X2))|~r2_hidden(esk1_3(X2,k10_xtuple_0(X2),X3),k10_xtuple_0(X1))|~r2_hidden(X3,k10_xtuple_0(X2))), inference(spm,[status(thm)],[c_0_27, c_0_19])).
cnf(c_0_34,negated_conjecture,(r2_hidden(X1,k10_xtuple_0(esk9_0))|~r2_hidden(X1,k9_xtuple_0(esk8_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35,plain,(r2_hidden(esk1_3(X1,k10_xtuple_0(X1),X2),k9_xtuple_0(X1))|~r2_hidden(X2,k10_xtuple_0(X1))), inference(spm,[status(thm)],[c_0_30, c_0_19])).
cnf(c_0_36,plain,(X1=k10_xtuple_0(X2)|~r2_hidden(esk2_2(X2,X1),k10_xtuple_0(X2))|~r2_hidden(esk2_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_31, c_0_19])).
cnf(c_0_37,plain,(k10_xtuple_0(k3_relat_1(X1,X2))=k10_xtuple_0(X2)|r2_hidden(esk2_2(X2,k10_xtuple_0(k3_relat_1(X1,X2))),k10_xtuple_0(X2))), inference(ef,[status(thm)],[c_0_32])).
cnf(c_0_38,plain,(r2_hidden(X1,k10_xtuple_0(k3_relat_1(X2,X3)))|~r2_hidden(esk1_3(X3,k10_xtuple_0(X3),X1),k10_xtuple_0(X2))|~r2_hidden(X1,k10_xtuple_0(X3))), inference(spm,[status(thm)],[c_0_12, c_0_33])).
cnf(c_0_39,negated_conjecture,(r2_hidden(esk1_3(esk8_0,k10_xtuple_0(esk8_0),X1),k10_xtuple_0(esk9_0))|~r2_hidden(X1,k10_xtuple_0(esk8_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40,plain,(k10_xtuple_0(k3_relat_1(X1,X2))=k10_xtuple_0(X2)|~r2_hidden(esk2_2(X2,k10_xtuple_0(k3_relat_1(X1,X2))),k10_xtuple_0(k3_relat_1(X1,X2)))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41,negated_conjecture,(r2_hidden(X1,k10_xtuple_0(k3_relat_1(esk9_0,esk8_0)))|~r2_hidden(X1,k10_xtuple_0(esk8_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42,negated_conjecture,(k10_xtuple_0(k3_relat_1(esk9_0,esk8_0))!=k10_xtuple_0(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_43,negated_conjecture,(~r2_hidden(esk2_2(esk8_0,k10_xtuple_0(k3_relat_1(esk9_0,esk8_0))),k10_xtuple_0(esk8_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_44,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_37]), c_0_42]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 32
# Proof object formula steps           : 13
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 2925
# ...of these trivial                  : 0
# ...subsumed                          : 998
# ...remaining for further processing  : 1927
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 3
# Generated clauses                    : 65606
# ...of the previous two non-trivial   : 65587
# Contextual simplify-reflections      : 559
# Paramodulations                      : 65535
# Factorizations                       : 64
# Equation resolutions                 : 7
# Current number of processed clauses  : 1889
#    Positive orientable unit clauses  : 102
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 1785
# Current number of unprocessed clauses: 62678
# ...number of literals in the above   : 185805
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 708281
# Rec. Clause-clause subsumption calls : 549582
# Non-unit clause-clause subsumptions  : 1563
# Unit Clause-clause subsumption calls : 12820
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 458
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2413808

# -------------------------------------------------
# User time                : 1.616 s
# System time              : 0.052 s
# Total time               : 1.668 s
# Maximum resident set size: 74548 pages
