# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l32_xboole_1, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', l32_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t100_xboole_1)).
fof(t83_xboole_1, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1)), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t83_xboole_1)).
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t4_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X1)=X1), file('out/t28_zfmisc_1__20_premises__rerun.E_input', idempotence_k3_xboole_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', d10_xboole_0)).
fof(l97_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', l97_xboole_1)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t91_xboole_1)).
fof(t28_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:![X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4))))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t28_zfmisc_1)).
fof(t65_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t65_enumset1)).
fof(t94_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t94_xboole_1)).
fof(t90_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t90_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t71_enumset1)).
fof(t45_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t45_enumset1)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t28_xboole_1)).
fof(t104_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X3,X2,X4)), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t104_enumset1)).
fof(t64_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', t64_enumset1)).
fof(l142_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k3_enumset1(X5,X6,X7,X8,X9))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', l142_enumset1)).
fof(d2_tarski, axiom, (![X1]:![X2]:![X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2)))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', d2_tarski)).
fof(d7_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:(X10=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)<=>![X11]:(r2_hidden(X11,X10)<=>~((~(X11=X1)&(~(X11=X2)&(~(X11=X3)&(~(X11=X4)&(~(X11=X5)&(~(X11=X6)&(~(X11=X7)&(~(X11=X8)&~(X11=X9))))))))))))), file('out/t28_zfmisc_1__20_premises__rerun.E_input', d7_enumset1)).
fof(c_0_21, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|k4_xboole_0(X3,X4)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])])])).
fof(c_0_22, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_23, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k4_xboole_0(X3,X4)=X3)&(k4_xboole_0(X3,X4)!=X3|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])])])).
fof(c_0_24, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk1_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
fof(c_0_25, plain, (![X3]:k3_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_26, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_27,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(r2_hidden(esk1_2(X1,X2),k3_xboole_0(X1,X2))|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, (![X3]:![X4]:r1_xboole_0(k3_xboole_0(X3,X4),k5_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[l97_xboole_1])).
cnf(c_0_33,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,plain,(k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35,plain,(k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_36,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37,plain,(r2_hidden(esk1_2(X1,X1),X1)|r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_33])).
fof(c_0_40, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_41,plain,(X1=k1_xboole_0|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,plain,(r1_xboole_0(X1,k5_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_38, c_0_31])).
cnf(c_0_44,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_31]), c_0_39])])).
fof(c_0_45, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4)))))), inference(assume_negation,[status(cth)],[t28_zfmisc_1])).
fof(c_0_46, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:k6_enumset1(X9,X10,X11,X12,X13,X14,X15,X16)=k2_xboole_0(k2_enumset1(X9,X10,X11,X12),k2_enumset1(X13,X14,X15,X16))), inference(variable_rename,[status(thm)],[t65_enumset1])).
fof(c_0_47, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t94_xboole_1])).
fof(c_0_48, plain, (![X5]:![X6]:![X7]:![X8]:k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8)=k2_enumset1(X5,X6,X7,X8)), inference(variable_rename,[status(thm)],[t90_enumset1])).
cnf(c_0_49,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_39])])).
cnf(c_0_51,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(rw,[status(thm)],[c_0_43, c_0_44])).
fof(c_0_52, negated_conjecture, ((r1_tarski(k2_tarski(esk2_0,esk3_0),k2_tarski(esk4_0,esk5_0))&(esk2_0!=esk4_0&esk2_0!=esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_45])])])])).
fof(c_0_53, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_54, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_55, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X5,X6),k2_tarski(X7,X8))), inference(variable_rename,[status(thm)],[t45_enumset1])).
cnf(c_0_56,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=k5_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_49, c_0_44])).
cnf(c_0_60,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_61, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_62,negated_conjecture,(r1_tarski(k2_tarski(esk2_0,esk3_0),k2_tarski(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8)),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_58]), c_0_58]), c_0_58])).
fof(c_0_67, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X5,X7,X6,X8)), inference(variable_rename,[status(thm)],[t104_enumset1])).
cnf(c_0_68,plain,(k5_xboole_0(X1,k1_xboole_0)=k5_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_59, c_0_44])).
cnf(c_0_69,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_60]), c_0_51])])).
cnf(c_0_70,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71,negated_conjecture,(r1_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0),k6_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk5_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_63]), c_0_64]), c_0_64]), c_0_58]), c_0_58])).
cnf(c_0_72,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k5_xboole_0(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2),k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4)),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2),k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_63]), c_0_63]), c_0_64]), c_0_64]), c_0_57]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58])).
cnf(c_0_73,plain,(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k5_xboole_0(k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8))))=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[c_0_66, c_0_49])).
cnf(c_0_74,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76,negated_conjecture,(k3_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0),k6_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk5_0))=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_77,plain,(k6_enumset1(X1,X1,X1,X2,X3,X3,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_49]), c_0_73])).
cnf(c_0_78,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X3,X2,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_58]), c_0_58])).
fof(c_0_79, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:k6_enumset1(X9,X10,X11,X12,X13,X14,X15,X16)=k2_xboole_0(k1_enumset1(X9,X10,X11),k3_enumset1(X12,X13,X14,X15,X16))), inference(variable_rename,[status(thm)],[t64_enumset1])).
fof(c_0_80, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:k7_enumset1(X10,X11,X12,X13,X14,X15,X16,X17,X18)=k2_xboole_0(k2_enumset1(X10,X11,X12,X13),k3_enumset1(X14,X15,X16,X17,X18))), inference(variable_rename,[status(thm)],[l142_enumset1])).
fof(c_0_81, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(X8=X5|X8=X6))|X7!=k2_tarski(X5,X6))&(((X8!=X5|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))&((X8!=X6|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))))&((((esk6_3(X5,X6,X7)!=X5|~r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6))&((esk6_3(X5,X6,X7)!=X6|~r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6)))&((r2_hidden(esk6_3(X5,X6,X7),X7)|(esk6_3(X5,X6,X7)=X5|esk6_3(X5,X6,X7)=X6))|X7=k2_tarski(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])])).
cnf(c_0_82,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[c_0_59, c_0_75])).
cnf(c_0_83,negated_conjecture,(k5_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0),k5_xboole_0(k6_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk5_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)))=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk4_0,esk3_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_76]), c_0_77]), c_0_78])).
cnf(c_0_84,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_85,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k3_enumset1(X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_86,plain,(X4=X3|X4=X2|X1!=k2_tarski(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_87,plain,(k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,X2)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_44])).
cnf(c_0_88,negated_conjecture,(k5_xboole_0(k6_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk5_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))=k5_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk4_0,esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_82, c_0_83])).
fof(c_0_89, plain, (![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:![X22]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:![X21]:((((~r2_hidden(X22,X21)|(X22=X12|(X22=X13|(X22=X14|(X22=X15|(X22=X16|(X22=X17|(X22=X18|(X22=X19|X22=X20)))))))))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X12|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X13|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X14|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X15|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X16|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X17|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X18|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((X22!=X19|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&((X22!=X20|r2_hidden(X22,X21))|X21!=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20)))))))))))&((((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X12|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X13|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X14|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X15|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X16|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X17|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X18|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&(((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X19|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))&((esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)!=X20|~r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))))))))))&((r2_hidden(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21),X21)|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X12|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X13|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X14|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X15|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X16|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X17|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X18|(esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X19|esk7_10(X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)=X20)))))))))|X21=k7_enumset1(X12,X13,X14,X15,X16,X17,X18,X19,X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_enumset1])])])])])])])])).
cnf(c_0_90,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8)),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_64]), c_0_57]), c_0_58]), c_0_58])).
cnf(c_0_91,plain,(k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k3_enumset1(X5,X6,X7,X8,X9)),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k3_enumset1(X5,X6,X7,X8,X9)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_57]), c_0_58]), c_0_58])).
cnf(c_0_92,plain,(X3=X4|X2=X4|X1!=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_63]), c_0_64]), c_0_58])).
cnf(c_0_93,plain,(k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_87]), c_0_69])).
cnf(c_0_94,negated_conjecture,(k5_xboole_0(k6_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk5_0),k5_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk4_0,esk3_0,esk5_0)))=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_82, c_0_88])).
cnf(c_0_95,plain,(r2_hidden(X11,X1)|X1!=k7_enumset1(X2,X3,X4,X5,X6,X7,X8,X9,X10)|X11!=X2), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_96,plain,(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k5_xboole_0(k3_enumset1(X4,X5,X6,X7,X8),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8))))=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[c_0_90, c_0_49])).
cnf(c_0_97,plain,(k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k5_xboole_0(k3_enumset1(X5,X6,X7,X8,X9),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k3_enumset1(X5,X6,X7,X8,X9))))=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(rw,[status(thm)],[c_0_91, c_0_49])).
cnf(c_0_98,plain,(X1=X2|X3=X2|~r2_hidden(X2,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X1))), inference(er,[status(thm)],[c_0_92])).
cnf(c_0_99,negated_conjecture,(k6_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk5_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk4_0,esk3_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_49]), c_0_93])).
cnf(c_0_100,plain,(r2_hidden(X1,k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_95])])).
cnf(c_0_101,plain,(k7_enumset1(X1,X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102,negated_conjecture,(esk5_0=X1|esk4_0=X1|~r2_hidden(X1,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk4_0,esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_103,plain,(r2_hidden(X1,k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_104,negated_conjecture,(esk2_0!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_105,negated_conjecture,(esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_106,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104]), c_0_105]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 107
# Proof object clause steps            : 64
# Proof object formula steps           : 43
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 61
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 52
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 47
# Processed clauses                    : 578
# ...of these trivial                  : 31
# ...subsumed                          : 315
# ...remaining for further processing  : 232
# Other redundant clauses eliminated   : 62
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 27
# Generated clauses                    : 2049
# ...of the previous two non-trivial   : 1436
# Contextual simplify-reflections      : 8
# Paramodulations                      : 1994
# Factorizations                       : 4
# Equation resolutions                 : 62
# Current number of processed clauses  : 136
#    Positive orientable unit clauses  : 51
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 70
# Current number of unprocessed clauses: 774
# ...number of literals in the above   : 2196
# Current number of archived formulas  : 0
# Current number of archived clauses   : 84
# Clause-clause subsumption calls (NU) : 1429
# Rec. Clause-clause subsumption calls : 1306
# Non-unit clause-clause subsumptions  : 181
# Unit Clause-clause subsumption calls : 173
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 286
# BW rewrite match successes           : 138
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 31575

# -------------------------------------------------
# User time                : 0.040 s
# System time              : 0.008 s
# Total time               : 0.048 s
# Maximum resident set size: 47632 pages
