# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_yellow19, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>?[X4]:((~(v1_xboole_0(X4))&(v2_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))&(v3_waybel_7(X4,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))&(r2_hidden(X2,X4)&r2_waybel_7(X1,X4,X3)))))))), file('out/t28_yellow19__4_premises__rerun__rerun.E_input', t28_yellow19)).
fof(t26_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>?[X4]:((~(v1_xboole_0(X4))&(v2_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))&(v3_waybel_7(X4,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))&(r2_hidden(X2,X4)&r2_waybel_7(X1,X4,X3)))))))), file('out/t28_yellow19__4_premises__rerun__rerun.E_input', t26_yellow19)).
fof(c_0_2, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>?[X4]:((~(v1_xboole_0(X4))&(v2_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))&(v3_waybel_7(X4,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))&(r2_hidden(X2,X4)&r2_waybel_7(X1,X4,X3))))))))), inference(assume_negation,[status(cth)],[t28_yellow19])).
fof(c_0_3, plain, (![X5]:![X6]:![X7]:![X9]:(((((((~v1_xboole_0(esk5_3(X5,X6,X7))|~r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((v2_waybel_0(esk5_3(X5,X6,X7),k3_yellow_1(k2_struct_0(X5)))|~r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((v13_waybel_0(esk5_3(X5,X6,X7),k3_yellow_1(k2_struct_0(X5)))|~r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((v3_waybel_7(esk5_3(X5,X6,X7),k3_yellow_1(k2_struct_0(X5)))|~r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((m1_subset_1(esk5_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X5)))))|~r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))))))&(((((r2_hidden(X6,esk5_3(X5,X6,X7))|~r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((r2_waybel_7(X5,esk5_3(X5,X6,X7),X7)|~r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))))&((((((v1_xboole_0(X9)|(~v2_waybel_0(X9,k3_yellow_1(k2_struct_0(X5)))|(~v13_waybel_0(X9,k3_yellow_1(k2_struct_0(X5)))|(~v3_waybel_7(X9,k3_yellow_1(k2_struct_0(X5)))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X5)))))))))|(~r2_hidden(X6,X9)|~r2_waybel_7(X5,X9,X7)))|r2_hidden(X7,k2_pre_topc(X5,X6)))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_yellow19])])])])])])])])).
fof(c_0_4, negated_conjecture, (![X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|((v1_xboole_0(X8)|(~v2_waybel_0(X8,k3_yellow_1(k2_struct_0(esk1_0)))|(~v13_waybel_0(X8,k3_yellow_1(k2_struct_0(esk1_0)))|(~v3_waybel_7(X8,k3_yellow_1(k2_struct_0(esk1_0)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))))))|(~r2_hidden(esk2_0,X8)|~r2_waybel_7(esk1_0,X8,esk3_0))))&(((~v1_xboole_0(esk4_0)|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((v2_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk1_0)))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((v13_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk1_0)))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((v3_waybel_7(esk4_0,k3_yellow_1(k2_struct_0(esk1_0)))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))))))&((r2_hidden(esk2_0,esk4_0)|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&(r2_waybel_7(esk1_0,esk4_0,esk3_0)|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_2])])])])])])])])).
cnf(c_0_5,plain,(v2_struct_0(X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|v1_xboole_0(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_waybel_7(X1,X4,X3)|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~v3_waybel_7(X4,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X4,k3_yellow_1(k2_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|v3_waybel_7(esk4_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|v2_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|v13_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(esk1_0,X2))|~r2_waybel_7(esk1_0,esk4_0,X1)|~r2_hidden(X2,esk4_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7]), c_0_8])]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13])).
cnf(c_0_15,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(esk1_0,esk2_0))|~r2_waybel_7(esk1_0,esk4_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_18,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_waybel_7(esk1_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_19,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20,negated_conjecture,(v1_xboole_0(X1)|~r2_waybel_7(esk1_0,X1,esk3_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_21,negated_conjecture,(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_22,negated_conjecture,(v1_xboole_0(X1)|~r2_waybel_7(esk1_0,X1,esk3_0)|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21])])).
cnf(c_0_23,plain,(v2_struct_0(X1)|v3_waybel_7(esk5_3(X1,X2,X3),k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k2_pre_topc(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_24,negated_conjecture,(v1_xboole_0(esk5_3(esk1_0,X1,X2))|~r2_waybel_7(esk1_0,esk5_3(esk1_0,X1,X2),esk3_0)|~v13_waybel_0(esk5_3(esk1_0,X1,X2),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk5_3(esk1_0,X1,X2),k3_yellow_1(k2_struct_0(esk1_0)))|~r2_hidden(esk2_0,esk5_3(esk1_0,X1,X2))|~r2_hidden(X2,k2_pre_topc(esk1_0,X1))|~m1_subset_1(esk5_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_25,plain,(v2_struct_0(X1)|v13_waybel_0(esk5_3(X1,X2,X3),k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k2_pre_topc(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_26,negated_conjecture,(v1_xboole_0(esk5_3(esk1_0,X1,X2))|~r2_waybel_7(esk1_0,esk5_3(esk1_0,X1,X2),esk3_0)|~v2_waybel_0(esk5_3(esk1_0,X1,X2),k3_yellow_1(k2_struct_0(esk1_0)))|~r2_hidden(esk2_0,esk5_3(esk1_0,X1,X2))|~r2_hidden(X2,k2_pre_topc(esk1_0,X1))|~m1_subset_1(esk5_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_27,plain,(v2_struct_0(X1)|v2_waybel_0(esk5_3(X1,X2,X3),k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k2_pre_topc(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_28,negated_conjecture,(v1_xboole_0(esk5_3(esk1_0,X1,X2))|~r2_waybel_7(esk1_0,esk5_3(esk1_0,X1,X2),esk3_0)|~r2_hidden(esk2_0,esk5_3(esk1_0,X1,X2))|~r2_hidden(X2,k2_pre_topc(esk1_0,X1))|~m1_subset_1(esk5_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_29,plain,(v2_struct_0(X1)|m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k2_pre_topc(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_30,negated_conjecture,(v1_xboole_0(esk5_3(esk1_0,X1,X2))|~r2_waybel_7(esk1_0,esk5_3(esk1_0,X1,X2),esk3_0)|~r2_hidden(esk2_0,esk5_3(esk1_0,X1,X2))|~r2_hidden(X2,k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_31,plain,(v2_struct_0(X1)|r2_waybel_7(X1,esk5_3(X1,X2,X3),X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k2_pre_topc(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_32,plain,(v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k2_pre_topc(X1,X2))|~v1_xboole_0(esk5_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_33,negated_conjecture,(v1_xboole_0(esk5_3(esk1_0,X1,esk3_0))|~r2_hidden(esk2_0,esk5_3(esk1_0,X1,esk3_0))|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_19]), c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_34,negated_conjecture,(~r2_hidden(esk2_0,esk5_3(esk1_0,X1,esk3_0))|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_19]), c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_35,plain,(v2_struct_0(X1)|r2_hidden(X2,esk5_3(X1,X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k2_pre_topc(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_36,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_21]), c_0_15]), c_0_19]), c_0_7]), c_0_8])]), c_0_9]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 32
# Proof object formula steps           : 5
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 53
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 9
# Generated clauses                    : 14
# ...of the previous two non-trivial   : 13
# Contextual simplify-reflections      : 9
# Paramodulations                      : 14
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 18
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 9
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 399
# Rec. Clause-clause subsumption calls : 44
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3447

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 43228 pages
