# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t29_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>((m1_pre_topc(X2,X3)&m1_pre_topc(X4,X3))=>m1_pre_topc(k1_tsep_1(X1,X2,X4),X3))))))), file('out/t29_tmap_1__11_premises__rerun.E_input', t29_tmap_1)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t29_tmap_1__11_premises__rerun.E_input', dt_m1_pre_topc)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t29_tmap_1__11_premises__rerun.E_input', d10_xboole_0)).
fof(dt_k1_tsep_1, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k1_tsep_1(X1,X2,X3)))&(v1_pre_topc(k1_tsep_1(X1,X2,X3))&m1_pre_topc(k1_tsep_1(X1,X2,X3),X1))))), file('out/t29_tmap_1__11_premises__rerun.E_input', dt_k1_tsep_1)).
fof(t2_tsep_1, axiom, (![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1))), file('out/t29_tmap_1__11_premises__rerun.E_input', t2_tsep_1)).
fof(t4_tsep_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3)))))), file('out/t29_tmap_1__11_premises__rerun.E_input', t4_tsep_1)).
fof(cc1_pre_topc, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2)))), file('out/t29_tmap_1__11_premises__rerun.E_input', cc1_pre_topc)).
fof(t25_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>k1_tsep_1(X1,X2,X2)=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))), file('out/t29_tmap_1__11_premises__rerun.E_input', t25_tmap_1)).
fof(commutativity_k1_tsep_1, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>k1_tsep_1(X1,X2,X3)=k1_tsep_1(X1,X3,X2))), file('out/t29_tmap_1__11_premises__rerun.E_input', commutativity_k1_tsep_1)).
fof(t22_tsep_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>m1_pre_topc(X2,k1_tsep_1(X1,X2,X3)))))), file('out/t29_tmap_1__11_premises__rerun.E_input', t22_tsep_1)).
fof(t23_tsep_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(m1_pre_topc(X2,X3)<=>k1_tsep_1(X1,X2,X3)=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))))), file('out/t29_tmap_1__11_premises__rerun.E_input', t23_tsep_1)).
fof(t27_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((m1_pre_topc(X2,X3)&m1_pre_topc(X4,X5))=>m1_pre_topc(k1_tsep_1(X1,X2,X4),k1_tsep_1(X1,X3,X5))))))))), file('out/t29_tmap_1__11_premises__rerun.E_input', t27_tmap_1)).
fof(c_0_12, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>((m1_pre_topc(X2,X3)&m1_pre_topc(X4,X3))=>m1_pre_topc(k1_tsep_1(X1,X2,X4),X3)))))))), inference(assume_negation,[status(cth)],[t29_tmap_1])).
fof(c_0_13, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
fof(c_0_14, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&((m1_pre_topc(esk2_0,esk3_0)&m1_pre_topc(esk4_0,esk3_0))&~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk4_0),esk3_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])).
fof(c_0_15, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:((~v2_struct_0(k1_tsep_1(X4,X5,X6))|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))&((v1_pre_topc(k1_tsep_1(X4,X5,X6))|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))&(m1_pre_topc(k1_tsep_1(X4,X5,X6),X4)|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tsep_1])])])])).
fof(c_0_17, plain, (![X2]:(~l1_pre_topc(X2)|m1_pre_topc(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_18,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:(((((~r1_tarski(u1_struct_0(X5),u1_struct_0(X6))|m1_pre_topc(X5,X6))|~m1_pre_topc(X6,X4))|~m1_pre_topc(X5,X4))|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))&((((~m1_pre_topc(X5,X6)|r1_tarski(u1_struct_0(X5),u1_struct_0(X6)))|~m1_pre_topc(X6,X4))|~m1_pre_topc(X5,X4))|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])])])).
cnf(c_0_22,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|m1_pre_topc(k1_tsep_1(X2,X3,X1),X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26,plain,(m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,negated_conjecture,(l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
fof(c_0_28, plain, (![X3]:![X4]:((~v2_pre_topc(X3)|~l1_pre_topc(X3))|(~m1_pre_topc(X4,X3)|v2_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])])])).
fof(c_0_29, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v2_pre_topc(X3)|~l1_pre_topc(X3)))|((v2_struct_0(X4)|~m1_pre_topc(X4,X3))|k1_tsep_1(X3,X4,X4)=g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_tmap_1])])])])])])).
fof(c_0_30, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4))))|k1_tsep_1(X4,X5,X6)=k1_tsep_1(X4,X6,X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k1_tsep_1])])])).
cnf(c_0_31,plain,(m1_pre_topc(X2,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~r1_tarski(u1_struct_0(X2),u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_33,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,X1),esk1_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_19]), c_0_20])]), c_0_24]), c_0_25])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|((v2_struct_0(X6)|~m1_pre_topc(X6,X4))|m1_pre_topc(X5,k1_tsep_1(X4,X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t22_tsep_1])])])])])])).
cnf(c_0_35,negated_conjecture,(m1_pre_topc(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36,plain,(v2_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_38,plain,(k1_tsep_1(X1,X2,X2)=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_39, plain, (![X4]:![X5]:![X6]:(((((~m1_pre_topc(X5,X6)|k1_tsep_1(X4,X5,X6)=g1_pre_topc(u1_struct_0(X6),u1_pre_topc(X6)))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((((k1_tsep_1(X4,X5,X6)!=g1_pre_topc(u1_struct_0(X6),u1_pre_topc(X6))|m1_pre_topc(X5,X6))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t23_tsep_1])])])])])])])).
cnf(c_0_40,plain,(k1_tsep_1(X1,X2,X3)=k1_tsep_1(X1,X3,X2)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41,negated_conjecture,(m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_42,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_43,plain,(m1_pre_topc(X1,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_44,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_19]), c_0_24])).
cnf(c_0_45,plain,(m1_pre_topc(X1,k1_tsep_1(X2,X1,X3))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk3_0,esk3_0,X1),esk3_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_35]), c_0_27])]), c_0_24])).
cnf(c_0_47,negated_conjecture,(v2_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_19]), c_0_20]), c_0_37])])).
cnf(c_0_48,negated_conjecture,(g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))=k1_tsep_1(esk1_0,esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_19]), c_0_20]), c_0_37])]), c_0_24]), c_0_25])).
fof(c_0_49, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|((v2_struct_0(X8)|~m1_pre_topc(X8,X6))|((v2_struct_0(X9)|~m1_pre_topc(X9,X6))|((v2_struct_0(X10)|~m1_pre_topc(X10,X6))|((~m1_pre_topc(X7,X8)|~m1_pre_topc(X9,X10))|m1_pre_topc(k1_tsep_1(X6,X7,X9),k1_tsep_1(X6,X8,X10))))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_tmap_1])])])])])])).
cnf(c_0_50,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|k1_tsep_1(X1,X2,X3)=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51,negated_conjecture,(k1_tsep_1(esk1_0,X1,esk4_0)=k1_tsep_1(esk1_0,esk4_0,X1)|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_20])]), c_0_42]), c_0_25])).
cnf(c_0_52,plain,(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_53,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_20]), c_0_37])])).
cnf(c_0_54,negated_conjecture,(l1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_44]), c_0_20])])).
cnf(c_0_55,negated_conjecture,(v2_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_44]), c_0_20]), c_0_37])])).
cnf(c_0_56,negated_conjecture,(m1_pre_topc(X1,k1_tsep_1(esk1_0,X1,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_19]), c_0_20]), c_0_37])]), c_0_24]), c_0_25])).
cnf(c_0_57,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk3_0,esk3_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_35]), c_0_24])).
cnf(c_0_58,negated_conjecture,(k1_tsep_1(esk3_0,esk3_0,esk3_0)=k1_tsep_1(esk1_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_35]), c_0_27]), c_0_47])]), c_0_24]), c_0_48])).
cnf(c_0_59,plain,(m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X1,X4,X5))|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60,negated_conjecture,(k1_tsep_1(esk1_0,X1,esk3_0)=k1_tsep_1(esk1_0,esk3_0,esk3_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk3_0)|~m1_pre_topc(X1,esk1_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_19]), c_0_20]), c_0_37])]), c_0_24]), c_0_25]), c_0_48])).
cnf(c_0_61,negated_conjecture,(m1_pre_topc(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_62,negated_conjecture,(k1_tsep_1(esk1_0,esk4_0,esk3_0)=k1_tsep_1(esk1_0,esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_19]), c_0_24])).
cnf(c_0_63,negated_conjecture,(r1_tarski(u1_struct_0(X1),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)))|~m1_pre_topc(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_64,negated_conjecture,(m1_pre_topc(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_19]), c_0_24])).
cnf(c_0_65,negated_conjecture,(r1_tarski(u1_struct_0(X1),u1_struct_0(esk3_0))|~m1_pre_topc(X1,esk3_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_19]), c_0_20]), c_0_37])])).
cnf(c_0_66,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)), inference(rw,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_67,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,X1,X2),k1_tsep_1(esk1_0,X3,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X2,esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,X3)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_41]), c_0_20]), c_0_37])]), c_0_42]), c_0_25])).
cnf(c_0_68,negated_conjecture,(k1_tsep_1(esk1_0,esk3_0,esk4_0)=k1_tsep_1(esk1_0,esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_41])]), c_0_42])).
cnf(c_0_69,negated_conjecture,(l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_41]), c_0_20])])).
cnf(c_0_70,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_71,negated_conjecture,(r1_tarski(u1_struct_0(esk3_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_72,negated_conjecture,(r1_tarski(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_44])])).
cnf(c_0_73,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,X1,X2),k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk3_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_19]), c_0_68]), c_0_24])).
cnf(c_0_74,negated_conjecture,(m1_pre_topc(esk4_0,esk4_0)), inference(spm,[status(thm)],[c_0_26, c_0_69])).
cnf(c_0_75,negated_conjecture,(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))=u1_struct_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])])).
cnf(c_0_76,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,X1,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_41])]), c_0_42])).
cnf(c_0_77,negated_conjecture,(m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_78,negated_conjecture,(m1_pre_topc(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_79,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_80,negated_conjecture,(r1_tarski(u1_struct_0(X1),u1_struct_0(esk3_0))|~m1_pre_topc(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0))), inference(rw,[status(thm)],[c_0_63, c_0_75])).
cnf(c_0_81,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])]), c_0_79])).
cnf(c_0_82,negated_conjecture,(r1_tarski(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk4_0)),u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_83,negated_conjecture,(~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk4_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_84,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,X1),esk1_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_77]), c_0_20])]), c_0_79]), c_0_25])).
cnf(c_0_85,negated_conjecture,(~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_82]), c_0_83])).
cnf(c_0_86,negated_conjecture,(m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_41]), c_0_42])).
cnf(c_0_87,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_19]), c_0_20]), c_0_37])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 88
# Proof object clause steps            : 63
# Proof object formula steps           : 25
# Proof object conjectures             : 51
# Proof object clause conjectures      : 48
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 91
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 3605
# ...of these trivial                  : 256
# ...subsumed                          : 478
# ...remaining for further processing  : 2871
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 27
# Backward-rewritten                   : 715
# Generated clauses                    : 28423
# ...of the previous two non-trivial   : 26271
# Contextual simplify-reflections      : 149
# Paramodulations                      : 28421
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 2099
#    Positive orientable unit clauses  : 1279
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 29
#    Non-unit-clauses                  : 791
# Current number of unprocessed clauses: 16955
# ...number of literals in the above   : 47473
# Current number of archived formulas  : 0
# Current number of archived clauses   : 770
# Clause-clause subsumption calls (NU) : 79132
# Rec. Clause-clause subsumption calls : 39326
# Non-unit clause-clause subsumptions  : 512
# Unit Clause-clause subsumption calls : 15017
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19999
# BW rewrite match successes           : 142
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1138071

# -------------------------------------------------
# User time                : 0.560 s
# System time              : 0.016 s
# Total time               : 0.576 s
# Maximum resident set size: 47248 pages
