# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_wellord2, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3]:![X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4))))))), file('out/t29_wellord2__5_premises__rerun.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, (![X1]:v1_relat_1(k1_wellord2(X1))), file('out/t29_wellord2__5_premises__rerun.E_input', dt_k1_wellord2)).
fof(l1_wellord1, axiom, (![X1]:(v1_relat_1(X1)=>(v1_relat_2(X1)<=>![X2]:(r2_hidden(X2,k3_relat_1(X1))=>r2_hidden(k4_tarski(X2,X2),X1))))), file('out/t29_wellord2__5_premises__rerun.E_input', l1_wellord1)).
fof(t2_wellord2, axiom, (![X1]:v1_relat_2(k1_wellord2(X1))), file('out/t29_wellord2__5_premises__rerun.E_input', t2_wellord2)).
fof(d1_relat_2, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(r1_relat_2(X1,X2)<=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(k4_tarski(X3,X3),X1))))), file('out/t29_wellord2__5_premises__rerun.E_input', d1_relat_2)).
fof(t29_wellord2, conjecture, (![X1]:r1_relat_2(k1_wellord2(X1),X1)), file('out/t29_wellord2__5_premises__rerun.E_input', t29_wellord2)).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:((((k3_relat_1(X6)=X5|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&(((((~r2_hidden(k4_tarski(X7,X8),X6)|r1_tarski(X7,X8))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&((((~r1_tarski(X7,X8)|r2_hidden(k4_tarski(X7,X8),X6))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))))&(((((r2_hidden(esk1_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&(((r2_hidden(esk2_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))&(((((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|~r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])])])).
fof(c_0_7, plain, (![X2]:v1_relat_1(k1_wellord2(X2))), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_8, plain, (![X3]:![X4]:(((~v1_relat_2(X3)|(~r2_hidden(X4,k3_relat_1(X3))|r2_hidden(k4_tarski(X4,X4),X3)))|~v1_relat_1(X3))&(((r2_hidden(esk3_1(X3),k3_relat_1(X3))|v1_relat_2(X3))|~v1_relat_1(X3))&((~r2_hidden(k4_tarski(esk3_1(X3),esk3_1(X3)),X3)|v1_relat_2(X3))|~v1_relat_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_wellord1])])])])])])])).
fof(c_0_9, plain, (![X2]:v1_relat_2(k1_wellord2(X2))), inference(variable_rename,[status(thm)],[t2_wellord2])).
cnf(c_0_10,plain,(k3_relat_1(X1)=X2|~v1_relat_1(X1)|X1!=k1_wellord2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:![X5]:(((~r1_relat_2(X4,X5)|(~r2_hidden(X6,X5)|r2_hidden(k4_tarski(X6,X6),X4)))|~v1_relat_1(X4))&(((r2_hidden(esk4_2(X4,X5),X5)|r1_relat_2(X4,X5))|~v1_relat_1(X4))&((~r2_hidden(k4_tarski(esk4_2(X4,X5),esk4_2(X4,X5)),X4)|r1_relat_2(X4,X5))|~v1_relat_1(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_relat_2])])])])])])])).
cnf(c_0_13,plain,(r2_hidden(k4_tarski(X2,X2),X1)|~v1_relat_1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~v1_relat_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(v1_relat_2(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])])).
fof(c_0_16, negated_conjecture, (~(![X1]:r1_relat_2(k1_wellord2(X1),X1))), inference(assume_negation,[status(cth)],[t29_wellord2])).
cnf(c_0_17,plain,(r1_relat_2(X1,X2)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(esk4_2(X1,X2),esk4_2(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(r2_hidden(k4_tarski(X1,X1),k1_wellord2(X2))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_11])])).
fof(c_0_19, negated_conjecture, (~r1_relat_2(k1_wellord2(esk5_0),esk5_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
cnf(c_0_20,plain,(r1_relat_2(k1_wellord2(X1),X2)|~r2_hidden(esk4_2(k1_wellord2(X1),X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_11])])).
cnf(c_0_21,plain,(r1_relat_2(X1,X2)|r2_hidden(esk4_2(X1,X2),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,negated_conjecture,(~r1_relat_2(k1_wellord2(esk5_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,plain,(r1_relat_2(k1_wellord2(X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_11])])).
cnf(c_0_24,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 25
# Proof object clause steps            : 12
# Proof object formula steps           : 13
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 37
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 37
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 13
# ...of the previous two non-trivial   : 10
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6
# Factorizations                       : 0
# Equation resolutions                 : 7
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 9
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 19
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 100
# Rec. Clause-clause subsumption calls : 50
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1547

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45268 pages
