# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_filter_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_filter_0(X1,X2)=k2_filter_0(X1,X3)=>X2=X3))))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', t2_filter_1)).
fof(cc1_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1)))))))))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', cc1_lattices)).
fof(dt_k2_filter_0, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v1_xboole_0(k2_filter_0(X1,X2)))&(v19_lattices(k2_filter_0(X1,X2),X1)&(v20_lattices(k2_filter_0(X1,X2),X1)&m1_subset_1(k2_filter_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))))))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', dt_k2_filter_0)).
fof(t15_filter_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X2,k2_filter_0(X1,X3))<=>r3_lattices(X1,X3,X2)))))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', t15_filter_0)).
fof(reflexivity_r3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_lattices(X1,X2,X2))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', reflexivity_r3_lattices)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', t4_subset)).
fof(redefinition_r3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3)))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', redefinition_r3_lattices)).
fof(dt_l3_lattices, axiom, (![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1)))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', dt_l3_lattices)).
fof(t8_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattices(X1,X2,X3)&r1_lattices(X1,X3,X2))=>X2=X3))))), file('out/t2_filter_1__98_premises__rerun__rerun.E_input', t8_lattices)).
fof(c_0_9, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_filter_0(X1,X2)=k2_filter_0(X1,X3)=>X2=X3)))))), inference(assume_negation,[status(cth)],[t2_filter_1])).
fof(c_0_10, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v4_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v5_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v6_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v7_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v8_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&((v9_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_11, negated_conjecture, (((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&l3_lattices(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(k2_filter_0(esk1_0,esk2_0)=k2_filter_0(esk1_0,esk3_0)&esk2_0!=esk3_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
fof(c_0_12, plain, (![X3]:![X4]:((~v1_xboole_0(k2_filter_0(X3,X4))|((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|~m1_subset_1(X4,u1_struct_0(X3))))&((v19_lattices(k2_filter_0(X3,X4),X3)|((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|~m1_subset_1(X4,u1_struct_0(X3))))&((v20_lattices(k2_filter_0(X3,X4),X3)|((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|~m1_subset_1(X4,u1_struct_0(X3))))&(m1_subset_1(k2_filter_0(X3,X4),k1_zfmisc_1(u1_struct_0(X3)))|((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|~m1_subset_1(X4,u1_struct_0(X3)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_filter_0])])])])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:(((((~r2_hidden(X5,k2_filter_0(X4,X6))|r3_lattices(X4,X6,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4))))&((((~r3_lattices(X4,X6,X5)|r2_hidden(X5,k2_filter_0(X4,X6)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_filter_0])])])])])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|r3_lattices(X4,X5,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_lattices])])])])])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,negated_conjecture,(v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,(v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20,plain,(v2_struct_0(X1)|v6_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_22,plain,(v2_struct_0(X2)|m1_subset_1(k2_filter_0(X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v10_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(k2_filter_0(esk1_0,esk2_0)=k2_filter_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,plain,(v2_struct_0(X1)|r2_hidden(X2,k2_filter_0(X1,X3))|~l3_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r3_lattices(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_26,plain,(r3_lattices(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~v6_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_27,negated_conjecture,(v9_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_28,negated_conjecture,(v8_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_29,negated_conjecture,(v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_30,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31,negated_conjecture,(m1_subset_1(k2_filter_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_17]), c_0_16])]), c_0_18])).
cnf(c_0_32,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_33, plain, (![X4]:![X5]:![X6]:(((~r3_lattices(X4,X5,X6)|r1_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_lattices(X4,X5,X6)|r3_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_34,negated_conjecture,(r2_hidden(X1,k2_filter_0(esk1_0,esk2_0))|~r3_lattices(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_24]), c_0_23]), c_0_17]), c_0_16])]), c_0_18])).
cnf(c_0_35,negated_conjecture,(r3_lattices(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_24]), c_0_17])]), c_0_18]), c_0_27]), c_0_28]), c_0_29])])).
cnf(c_0_36,plain,(v2_struct_0(X1)|r3_lattices(X1,X3,X2)|~l3_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X2,k2_filter_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_37,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k2_filter_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38,negated_conjecture,(r2_hidden(X1,k2_filter_0(esk1_0,esk2_0))|~r3_lattices(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_32]), c_0_17]), c_0_16])]), c_0_18])).
fof(c_0_39, plain, (![X2]:((l1_lattices(X2)|~l3_lattices(X2))&(l2_lattices(X2)|~l3_lattices(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_40,plain,(v2_struct_0(X2)|r1_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r3_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41,negated_conjecture,(r2_hidden(esk3_0,k2_filter_0(esk1_0,esk2_0))|~r3_lattices(esk1_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_24])).
cnf(c_0_42,negated_conjecture,(r3_lattices(esk1_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_35, c_0_24])).
cnf(c_0_43,negated_conjecture,(r3_lattices(esk1_0,esk3_0,X1)|~r2_hidden(X1,k2_filter_0(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_23]), c_0_24]), c_0_17]), c_0_16])]), c_0_18]), c_0_37])).
cnf(c_0_44,negated_conjecture,(r2_hidden(esk2_0,k2_filter_0(esk1_0,esk2_0))|~r3_lattices(esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_32])).
cnf(c_0_45,negated_conjecture,(r3_lattices(esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_32])).
fof(c_0_46, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|(~v4_lattices(X4)|~l2_lattices(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|((~r1_lattices(X4,X5,X6)|~r1_lattices(X4,X6,X5))|X5=X6))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_lattices])])])])])])).
cnf(c_0_47,plain,(l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48,plain,(v2_struct_0(X1)|v4_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_49,negated_conjecture,(r1_lattices(esk1_0,X1,X2)|~r3_lattices(esk1_0,X1,X2)|~r2_hidden(X1,k2_filter_0(esk1_0,esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_37]), c_0_27]), c_0_28]), c_0_29]), c_0_17])]), c_0_18])).
cnf(c_0_50,negated_conjecture,(r2_hidden(esk3_0,k2_filter_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42])])).
cnf(c_0_51,negated_conjecture,(r3_lattices(esk1_0,esk3_0,esk2_0)|~r3_lattices(esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52,negated_conjecture,(r2_hidden(esk2_0,k2_filter_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45])])).
cnf(c_0_53,negated_conjecture,(r3_lattices(esk1_0,esk2_0,esk3_0)|~r3_lattices(esk1_0,esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_41]), c_0_32]), c_0_24]), c_0_17]), c_0_16])]), c_0_18])).
cnf(c_0_54,plain,(X1=X2|v2_struct_0(X3)|~r1_lattices(X3,X2,X1)|~r1_lattices(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l2_lattices(X3)|~v4_lattices(X3)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55,negated_conjecture,(l2_lattices(esk1_0)), inference(spm,[status(thm)],[c_0_47, c_0_17])).
cnf(c_0_56,negated_conjecture,(v4_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_57,negated_conjecture,(r1_lattices(esk1_0,esk3_0,X1)|~r3_lattices(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58,negated_conjecture,(r3_lattices(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_45])])).
cnf(c_0_59,negated_conjecture,(r1_lattices(esk1_0,esk2_0,X1)|~r3_lattices(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_52])).
cnf(c_0_60,negated_conjecture,(r3_lattices(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_42])])).
cnf(c_0_61,negated_conjecture,(X1=esk3_0|~r1_lattices(esk1_0,esk3_0,X1)|~r1_lattices(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_24]), c_0_18]), c_0_55]), c_0_56])])).
cnf(c_0_62,negated_conjecture,(r1_lattices(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_32]), c_0_58])])).
cnf(c_0_63,negated_conjecture,(r1_lattices(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_24]), c_0_60])])).
cnf(c_0_64,negated_conjecture,(esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_65,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_32]), c_0_62]), c_0_63])]), c_0_64]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 47
# Proof object formula steps           : 19
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 71
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 26
# Processed clauses                    : 121
# ...of these trivial                  : 5
# ...subsumed                          : 9
# ...remaining for further processing  : 107
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 9
# Generated clauses                    : 105
# ...of the previous two non-trivial   : 77
# Contextual simplify-reflections      : 3
# Paramodulations                      : 105
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 69
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 1038
# Rec. Clause-clause subsumption calls : 183
# Non-unit clause-clause subsumptions  : 12
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4508

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45364 pages
