# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_yellow_0, conjecture, (![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5)))))))))), file('out/t2_yellow_0__5_premises__rerun.E_input', t2_yellow_0)).
fof(free_g1_orders_2, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3]:![X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4)))), file('out/t2_yellow_0__5_premises__rerun.E_input', free_g1_orders_2)).
fof(dt_u1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))), file('out/t2_yellow_0__5_premises__rerun.E_input', dt_u1_orders_2)).
fof(d8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4))))))), file('out/t2_yellow_0__5_premises__rerun.E_input', d8_lattice3)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t2_yellow_0__5_premises__rerun.E_input', d9_lattice3)).
fof(d5_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))))))), file('out/t2_yellow_0__5_premises__rerun.E_input', d5_orders_2)).
fof(c_0_6, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))))))))))), inference(assume_negation,[status(cth)],[t2_yellow_0])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:(((X5=X7|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5))))&((X6=X8|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])])])).
fof(c_0_8, negated_conjecture, ((l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(esk4_0=esk5_0&(((r1_lattice3(esk1_0,esk3_0,esk4_0)|r2_lattice3(esk1_0,esk3_0,esk4_0))&(~r1_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(esk1_0,esk3_0,esk4_0)))&((r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0))&(~r1_lattice3(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])])).
cnf(c_0_9,plain,(X1=X4|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|g1_orders_2(X2,X1)!=g1_orders_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,negated_conjecture,(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, plain, (![X2]:(~l1_orders_2(X2)|m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_12,negated_conjecture,(u1_orders_2(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X2,X1)|~m1_subset_1(u1_orders_2(esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13,plain,(m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,negated_conjecture,(u1_orders_2(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_16,plain,(X2=X3|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|g1_orders_2(X2,X1)!=g1_orders_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,negated_conjecture,(u1_orders_2(esk2_0)=u1_orders_2(esk1_0)), inference(er,[status(thm)],[c_0_15])).
fof(c_0_18, plain, (![X5]:![X6]:![X7]:![X8]:((((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk6_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk6_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,X7,esk6_3(X5,X6,X7))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])])])).
cnf(c_0_19,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20,negated_conjecture,(esk4_0=esk5_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21,negated_conjecture,(u1_struct_0(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X1,X2)|~m1_subset_1(u1_orders_2(esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(spm,[status(thm)],[c_0_16, c_0_10])).
cnf(c_0_22,negated_conjecture,(m1_subset_1(u1_orders_2(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_17]), c_0_14])])).
cnf(c_0_23,plain,(r1_lattice3(X1,X3,X2)|m1_subset_1(esk6_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25,negated_conjecture,(u1_struct_0(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_17]), c_0_22])])).
cnf(c_0_26,negated_conjecture,(r1_lattice3(esk2_0,X1,esk4_0)|m1_subset_1(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_14])])).
cnf(c_0_27,negated_conjecture,(u1_struct_0(esk2_0)=u1_struct_0(esk1_0)), inference(er,[status(thm)],[c_0_25])).
fof(c_0_28, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk7_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk7_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk7_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
fof(c_0_29, plain, (![X4]:![X5]:![X6]:(((((~r1_orders_2(X4,X5,X6)|r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&((((~r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4))|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_orders_2])])])])])])).
cnf(c_0_30,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31,negated_conjecture,(r1_lattice3(esk2_0,X1,esk4_0)|m1_subset_1(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_33,plain,(r2_lattice3(X1,X3,X2)|m1_subset_1(esk7_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk1_0,X1,esk6_3(esk2_0,X2,esk4_0))|r1_lattice3(esk2_0,X2,esk4_0)|~r2_hidden(esk6_3(esk2_0,X2,esk4_0),X3)|~r1_lattice3(esk1_0,X3,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_37,plain,(r1_lattice3(X1,X3,X2)|r2_hidden(esk6_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_38,negated_conjecture,(r2_lattice3(esk2_0,X1,esk4_0)|m1_subset_1(esk7_3(esk2_0,X1,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_24]), c_0_14])])).
cnf(c_0_39,negated_conjecture,(r1_orders_2(esk2_0,X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_17]), c_0_14])]), c_0_27]), c_0_27])).
cnf(c_0_40,plain,(r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),X2)|~r1_lattice3(esk1_0,X2,esk4_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_43,negated_conjecture,(r2_hidden(esk6_3(esk2_0,X1,X2),X1)|r1_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_27]), c_0_14])])).
cnf(c_0_44,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_45,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_46,negated_conjecture,(r2_lattice3(esk2_0,X1,esk4_0)|m1_subset_1(esk7_3(esk2_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_38, c_0_27])).
cnf(c_0_47,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_48,negated_conjecture,(r1_orders_2(esk2_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_32])])).
cnf(c_0_49,negated_conjecture,(r1_lattice3(esk2_0,X1,X2)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_27]), c_0_14])])).
cnf(c_0_50,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|r2_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51,negated_conjecture,(r2_hidden(esk6_3(esk2_0,X1,esk4_0),X1)|r1_lattice3(esk2_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_43, c_0_36])).
cnf(c_0_52,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_44, c_0_20])).
cnf(c_0_53,negated_conjecture,(r1_orders_2(esk1_0,esk7_3(esk2_0,X1,esk4_0),X2)|r2_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk7_3(esk2_0,X1,esk4_0),X3)|~r2_lattice3(esk1_0,X3,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_32])])).
cnf(c_0_54,negated_conjecture,(r1_lattice3(esk2_0,X1,X2)|~r1_orders_2(esk1_0,X2,esk6_3(esk2_0,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_27]), c_0_14])]), c_0_49])).
cnf(c_0_55,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,esk3_0,esk4_0))|r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_56,plain,(r2_lattice3(X1,X3,X2)|r2_hidden(esk7_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_57,negated_conjecture,(r1_orders_2(esk1_0,esk7_3(esk2_0,X1,esk4_0),esk4_0)|r2_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk7_3(esk2_0,X1,esk4_0),X2)|~r2_lattice3(esk1_0,X2,esk4_0)), inference(spm,[status(thm)],[c_0_53, c_0_36])).
cnf(c_0_58,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_36])]), c_0_52])).
cnf(c_0_59,negated_conjecture,(r2_hidden(esk7_3(esk2_0,X1,X2),X1)|r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_27]), c_0_14])])).
cnf(c_0_60,negated_conjecture,(r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_61,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk7_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_62,negated_conjecture,(r2_lattice3(esk2_0,X1,X2)|m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_27]), c_0_14])])).
cnf(c_0_63,negated_conjecture,(r1_orders_2(esk1_0,esk7_3(esk2_0,X1,esk4_0),esk4_0)|r2_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk7_3(esk2_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64,negated_conjecture,(r2_hidden(esk7_3(esk2_0,X1,esk4_0),X1)|r2_lattice3(esk2_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_59, c_0_36])).
cnf(c_0_65,negated_conjecture,(r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_60, c_0_20])).
cnf(c_0_66,negated_conjecture,(r2_lattice3(esk2_0,X1,X2)|~r1_orders_2(esk1_0,esk7_3(esk2_0,X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_48]), c_0_27]), c_0_14])]), c_0_62])).
cnf(c_0_67,negated_conjecture,(r1_orders_2(esk1_0,esk7_3(esk2_0,esk3_0,esk4_0),esk4_0)|r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68,negated_conjecture,(~r2_lattice3(esk2_0,esk3_0,esk5_0)|~r1_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_69,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)|~r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_41, c_0_65])).
cnf(c_0_70,negated_conjecture,(r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_36])])).
cnf(c_0_71,negated_conjecture,(~r1_lattice3(esk2_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_20]), c_0_20])).
cnf(c_0_72,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_70])])).
cnf(c_0_73,negated_conjecture,(~r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_70])])).
cnf(c_0_74,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_51]), c_0_73])).
cnf(c_0_75,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_74]), c_0_36])]), c_0_73]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 76
# Proof object clause steps            : 63
# Proof object formula steps           : 13
# Proof object conjectures             : 53
# Proof object clause conjectures      : 50
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 58
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 128
# ...of these trivial                  : 1
# ...subsumed                          : 3
# ...remaining for further processing  : 124
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 27
# Generated clauses                    : 179
# ...of the previous two non-trivial   : 184
# Contextual simplify-reflections      : 5
# Paramodulations                      : 172
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 72
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 59
# Current number of unprocessed clauses: 80
# ...number of literals in the above   : 324
# Current number of archived formulas  : 0
# Current number of archived clauses   : 52
# Clause-clause subsumption calls (NU) : 631
# Rec. Clause-clause subsumption calls : 185
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6886

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43188 pages
