# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t30_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>((m1_pre_topc(X2,X4)&m1_pre_topc(X3,X4))=>(r1_tsep_1(X2,X3)|m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)))))))), file('out/t30_tmap_1__6_premises__rerun.E_input', t30_tmap_1)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t30_tmap_1__6_premises__rerun.E_input', dt_m1_pre_topc)).
fof(dt_k2_tsep_1, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k2_tsep_1(X1,X2,X3)))&(v1_pre_topc(k2_tsep_1(X1,X2,X3))&m1_pre_topc(k2_tsep_1(X1,X2,X3),X1))))), file('out/t30_tmap_1__6_premises__rerun.E_input', dt_k2_tsep_1)).
fof(d5_tsep_1, axiom, (![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k2_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))))))), file('out/t30_tmap_1__6_premises__rerun.E_input', d5_tsep_1)).
fof(t1_tsep_1, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t30_tmap_1__6_premises__rerun.E_input', t1_tsep_1)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t30_tmap_1__6_premises__rerun.E_input', t3_subset)).
fof(t4_tsep_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3)))))), file('out/t30_tmap_1__6_premises__rerun.E_input', t4_tsep_1)).
fof(c_0_7, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>((m1_pre_topc(X2,X4)&m1_pre_topc(X3,X4))=>(r1_tsep_1(X2,X3)|m1_pre_topc(k2_tsep_1(X1,X2,X3),X4))))))))), inference(assume_negation,[status(cth)],[t30_tmap_1])).
fof(c_0_8, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
fof(c_0_9, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&((m1_pre_topc(esk2_0,esk4_0)&m1_pre_topc(esk3_0,esk4_0))&(~r1_tsep_1(esk2_0,esk3_0)&~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])])])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:((~v2_struct_0(k2_tsep_1(X4,X5,X6))|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))&((v1_pre_topc(k2_tsep_1(X4,X5,X6))|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))&(m1_pre_topc(k2_tsep_1(X4,X5,X6),X4)|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tsep_1])])])])).
cnf(c_0_11,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,negated_conjecture,(m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X8!=k2_tsep_1(X5,X6,X7)|u1_struct_0(X8)=k3_xboole_0(u1_struct_0(X6),u1_struct_0(X7)))|(v2_struct_0(X8)|(~v1_pre_topc(X8)|~m1_pre_topc(X8,X5))))|r1_tsep_1(X6,X7))|(v2_struct_0(X7)|~m1_pre_topc(X7,X5)))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|~l1_pre_topc(X5)))&((((((u1_struct_0(X8)!=k3_xboole_0(u1_struct_0(X6),u1_struct_0(X7))|X8=k2_tsep_1(X5,X6,X7))|(v2_struct_0(X8)|(~v1_pre_topc(X8)|~m1_pre_topc(X8,X5))))|r1_tsep_1(X6,X7))|(v2_struct_0(X7)|~m1_pre_topc(X7,X5)))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|~l1_pre_topc(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tsep_1])])])])])])])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|m1_pre_topc(k2_tsep_1(X2,X3,X1),X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,negated_conjecture,(m1_pre_topc(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19,negated_conjecture,(l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_20,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_tsep_1(X2,X3)|v2_struct_0(X4)|u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_pre_topc(X4)|X4!=k2_tsep_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v1_pre_topc(k2_tsep_1(X2,X3,X1))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22,plain,(v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_struct_0(k2_tsep_1(X2,X3,X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_23, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(X3)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])])])).
cnf(c_0_24,negated_conjecture,(m1_pre_topc(k2_tsep_1(esk4_0,esk2_0,X1),esk4_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_25,negated_conjecture,(m1_pre_topc(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27,plain,(u1_struct_0(k2_tsep_1(X1,X2,X3))=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_20]), c_0_15]), c_0_21]), c_0_22])).
fof(c_0_28, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_29,plain,(m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,negated_conjecture,(m1_pre_topc(k2_tsep_1(esk4_0,esk2_0,esk3_0),esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_31,negated_conjecture,(m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_33,negated_conjecture,(k3_xboole_0(u1_struct_0(X1),u1_struct_0(esk3_0))=u1_struct_0(k2_tsep_1(esk4_0,X1,esk3_0))|r1_tsep_1(X1,esk3_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_25]), c_0_19])]), c_0_18]), c_0_26])).
cnf(c_0_34,negated_conjecture,(~r1_tsep_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(u1_struct_0(k2_tsep_1(esk4_0,esk2_0,esk3_0)),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_19])])).
cnf(c_0_37,negated_conjecture,(k3_xboole_0(u1_struct_0(X1),u1_struct_0(esk3_0))=u1_struct_0(k2_tsep_1(esk1_0,X1,esk3_0))|r1_tsep_1(X1,esk3_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_31]), c_0_13])]), c_0_32]), c_0_26])).
cnf(c_0_38,negated_conjecture,(m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39,negated_conjecture,(k3_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0))=u1_struct_0(k2_tsep_1(esk4_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_16]), c_0_34]), c_0_17])).
fof(c_0_40, plain, (![X4]:![X5]:![X6]:(((((~r1_tarski(u1_struct_0(X5),u1_struct_0(X6))|m1_pre_topc(X5,X6))|~m1_pre_topc(X6,X4))|~m1_pre_topc(X5,X4))|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))&((((~m1_pre_topc(X5,X6)|r1_tarski(u1_struct_0(X5),u1_struct_0(X6)))|~m1_pre_topc(X6,X4))|~m1_pre_topc(X5,X4))|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])])])).
cnf(c_0_41,negated_conjecture,(r1_tarski(u1_struct_0(k2_tsep_1(esk4_0,esk2_0,esk3_0)),u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42,negated_conjecture,(u1_struct_0(k2_tsep_1(esk4_0,esk2_0,esk3_0))=u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_34]), c_0_17])).
cnf(c_0_43,plain,(m1_pre_topc(X2,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~r1_tarski(u1_struct_0(X2),u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44,negated_conjecture,(r1_tarski(u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk4_0))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45,negated_conjecture,(~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_46,negated_conjecture,(m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,X1),esk1_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_38]), c_0_13])]), c_0_17]), c_0_32])).
cnf(c_0_47,negated_conjecture,(~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
cnf(c_0_48,negated_conjecture,(m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_31]), c_0_26])).
cnf(c_0_49,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_50,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_12]), c_0_13]), c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 36
# Proof object formula steps           : 15
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 261
# ...of these trivial                  : 4
# ...subsumed                          : 0
# ...remaining for further processing  : 257
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 996
# ...of the previous two non-trivial   : 962
# Contextual simplify-reflections      : 3
# Paramodulations                      : 995
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 225
#    Positive orientable unit clauses  : 87
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 132
# Current number of unprocessed clauses: 748
# ...number of literals in the above   : 2732
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 4644
# Rec. Clause-clause subsumption calls : 3396
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 190
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 208
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 29287

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.000 s
# Total time               : 0.032 s
# Maximum resident set size: 45280 pages
