# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d9_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))))), file('out/t30_yellow_0__2_premises__rerun.E_input', d9_yellow_0)).
fof(t30_yellow_0, conjecture, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))))))), file('out/t30_yellow_0__2_premises__rerun.E_input', t30_yellow_0)).
fof(t15_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('out/t30_yellow_0__2_premises__rerun.E_input', t15_yellow_0)).
fof(c_0_3, plain, (![X5]:![X6]:![X7]:![X8]:((((((r2_lattice3(X5,X6,X7)|X7!=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X6,X8)|r1_orders_2(X5,X7,X8)))|X7!=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))&((((((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((((r2_lattice3(X5,X6,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~r1_orders_2(X5,X7,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])])])).
fof(c_0_4, negated_conjecture, (~(![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))))), inference(assume_negation,[status(cth)],[t30_yellow_0])).
cnf(c_0_5,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|X2!=k1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_6, negated_conjecture, (![X9]:((v5_orders_2(esk2_0)&l1_orders_2(esk2_0))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(((((r2_lattice3(esk2_0,esk6_0,esk3_0)|esk3_0=k1_yellow_0(esk2_0,esk4_0))&((~m1_subset_1(X9,u1_struct_0(esk2_0))|(~r2_lattice3(esk2_0,esk6_0,X9)|r1_orders_2(esk2_0,esk3_0,X9)))|esk3_0=k1_yellow_0(esk2_0,esk4_0)))&((esk3_0!=k1_yellow_0(esk2_0,esk6_0)|~r1_yellow_0(esk2_0,esk6_0))|esk3_0=k1_yellow_0(esk2_0,esk4_0)))&(((r2_lattice3(esk2_0,esk6_0,esk3_0)|r1_yellow_0(esk2_0,esk4_0))&((~m1_subset_1(X9,u1_struct_0(esk2_0))|(~r2_lattice3(esk2_0,esk6_0,X9)|r1_orders_2(esk2_0,esk3_0,X9)))|r1_yellow_0(esk2_0,esk4_0)))&((esk3_0!=k1_yellow_0(esk2_0,esk6_0)|~r1_yellow_0(esk2_0,esk6_0))|r1_yellow_0(esk2_0,esk4_0))))&((((r2_lattice3(esk2_0,esk6_0,esk3_0)|(m1_subset_1(esk5_0,u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,esk4_0,esk3_0)))&((~m1_subset_1(X9,u1_struct_0(esk2_0))|(~r2_lattice3(esk2_0,esk6_0,X9)|r1_orders_2(esk2_0,esk3_0,X9)))|(m1_subset_1(esk5_0,u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,esk4_0,esk3_0))))&((esk3_0!=k1_yellow_0(esk2_0,esk6_0)|~r1_yellow_0(esk2_0,esk6_0))|(m1_subset_1(esk5_0,u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,esk4_0,esk3_0))))&((((r2_lattice3(esk2_0,esk6_0,esk3_0)|(r2_lattice3(esk2_0,esk4_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)))&((~m1_subset_1(X9,u1_struct_0(esk2_0))|(~r2_lattice3(esk2_0,esk6_0,X9)|r1_orders_2(esk2_0,esk3_0,X9)))|(r2_lattice3(esk2_0,esk4_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0))))&((esk3_0!=k1_yellow_0(esk2_0,esk6_0)|~r1_yellow_0(esk2_0,esk6_0))|(r2_lattice3(esk2_0,esk4_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0))))&(((r2_lattice3(esk2_0,esk6_0,esk3_0)|(~r1_orders_2(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)))&((~m1_subset_1(X9,u1_struct_0(esk2_0))|(~r2_lattice3(esk2_0,esk6_0,X9)|r1_orders_2(esk2_0,esk3_0,X9)))|(~r1_orders_2(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0))))&((esk3_0!=k1_yellow_0(esk2_0,esk6_0)|~r1_yellow_0(esk2_0,esk6_0))|(~r1_orders_2(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])])).
fof(c_0_7, plain, (![X5]:![X6]:![X8]:![X6]:![X9]:((((m1_subset_1(esk7_2(X5,X6),u1_struct_0(X5))|~r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_lattice3(X5,X6,esk7_2(X5,X6))|~r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X6,X8)|r1_orders_2(X5,esk7_2(X5,X6),X8)))|~r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))&(((((m1_subset_1(esk8_3(X5,X6,X9),u1_struct_0(X5))|~r2_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((((r2_lattice3(X5,X6,esk8_3(X5,X6,X9))|~r2_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~r1_orders_2(X5,X9,esk8_3(X5,X6,X9))|~r2_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])])).
cnf(c_0_8,plain,(r1_orders_2(X1,X2,X4)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|X2!=k1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_9,plain,(r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_5])).
cnf(c_0_10,negated_conjecture,(esk3_0=k1_yellow_0(esk2_0,esk4_0)|r2_lattice3(esk2_0,esk6_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,negated_conjecture,(r1_yellow_0(esk2_0,esk4_0)|r2_lattice3(esk2_0,esk6_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,plain,(r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,esk8_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,negated_conjecture,(esk3_0=k1_yellow_0(esk2_0,esk4_0)|r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,negated_conjecture,(v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17,plain,(r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_18,negated_conjecture,(r2_lattice3(esk2_0,esk6_0,esk3_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r1_orders_2(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk5_0)|r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_20,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk5_0)|r2_lattice3(esk2_0,esk6_0,esk3_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21,negated_conjecture,(r2_lattice3(esk2_0,esk6_0,esk3_0)|r2_lattice3(esk2_0,esk4_0,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])]), c_0_13])).
cnf(c_0_22,negated_conjecture,(k1_yellow_0(esk2_0,esk4_0)=esk3_0|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,esk6_0,esk8_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~m1_subset_1(esk8_3(esk2_0,X1,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_23,plain,(r1_yellow_0(X1,X2)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|r2_lattice3(esk2_0,esk6_0,esk3_0)|~r2_lattice3(esk2_0,esk4_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_10]), c_0_11]), c_0_12])]), c_0_13])).
cnf(c_0_25,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk2_0))|r2_lattice3(esk2_0,esk6_0,esk3_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk5_0)|r2_lattice3(esk2_0,esk6_0,esk3_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_20]), c_0_21])).
cnf(c_0_27,negated_conjecture,(r1_yellow_0(esk2_0,esk4_0)|r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_28,negated_conjecture,(k1_yellow_0(esk2_0,esk4_0)=esk3_0|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,esk6_0,esk8_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_29,plain,(r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30,negated_conjecture,(r2_lattice3(esk2_0,esk6_0,esk3_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_24]), c_0_25]), c_0_26]), c_0_21])).
cnf(c_0_31,negated_conjecture,(r1_yellow_0(esk2_0,esk4_0)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,esk6_0,esk8_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~m1_subset_1(esk8_3(esk2_0,X1,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_27]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_32,negated_conjecture,(k1_yellow_0(esk2_0,esk4_0)=esk3_0|r1_yellow_0(esk2_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_33,negated_conjecture,(r1_yellow_0(esk2_0,esk4_0)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,esk6_0,esk8_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_23]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_34,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r1_orders_2(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|r1_yellow_0(esk2_0,esk6_0)|~r2_lattice3(esk2_0,esk4_0,X1)|~r1_yellow_0(esk2_0,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_32]), c_0_11]), c_0_12])])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk2_0))|r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_37,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk3_0)|r1_yellow_0(esk2_0,esk6_0)|~r1_yellow_0(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_32]), c_0_11]), c_0_12])])).
cnf(c_0_38,negated_conjecture,(r1_yellow_0(esk2_0,esk6_0)|r1_yellow_0(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_29]), c_0_30]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_39,negated_conjecture,(esk3_0=k1_yellow_0(esk2_0,esk4_0)|~r1_yellow_0(esk2_0,esk6_0)|esk3_0!=k1_yellow_0(esk2_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_40,negated_conjecture,(r1_yellow_0(esk2_0,esk4_0)|~r1_yellow_0(esk2_0,esk6_0)|esk3_0!=k1_yellow_0(esk2_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_41,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|r1_yellow_0(esk2_0,esk6_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_27]), c_0_19])).
cnf(c_0_42,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk3_0)|r1_yellow_0(esk2_0,esk6_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43,negated_conjecture,(~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r1_orders_2(esk2_0,esk3_0,esk5_0)|~r1_yellow_0(esk2_0,esk6_0)|esk3_0!=k1_yellow_0(esk2_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_44,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r1_yellow_0(esk2_0,esk6_0)|esk3_0!=k1_yellow_0(esk2_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_45,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk3_0)|k1_yellow_0(esk2_0,esk6_0)!=esk3_0|~r1_yellow_0(esk2_0,esk6_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_39]), c_0_11]), c_0_12])]), c_0_40])).
cnf(c_0_46,negated_conjecture,(r1_yellow_0(esk2_0,esk6_0)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,esk6_0,esk8_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~m1_subset_1(esk8_3(esk2_0,X1,esk3_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_41]), c_0_16]), c_0_11]), c_0_12])]), c_0_42])).
cnf(c_0_47,plain,(X2=k1_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk1_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_48,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|k1_yellow_0(esk2_0,esk6_0)!=esk3_0|~r2_lattice3(esk2_0,esk4_0,X1)|~r1_yellow_0(esk2_0,esk6_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_39]), c_0_11]), c_0_12])]), c_0_40])).
cnf(c_0_49,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r1_yellow_0(esk2_0,esk6_0)|esk3_0!=k1_yellow_0(esk2_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_50,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk5_0)|k1_yellow_0(esk2_0,esk6_0)!=esk3_0|~r1_yellow_0(esk2_0,esk6_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_19]), c_0_44]), c_0_44]), c_0_45])).
cnf(c_0_51,negated_conjecture,(r1_yellow_0(esk2_0,esk6_0)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,esk6_0,esk8_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_23]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_52,negated_conjecture,(k1_yellow_0(esk2_0,esk4_0)=esk3_0|k1_yellow_0(esk2_0,X1)=esk3_0|~r2_lattice3(esk2_0,esk6_0,esk1_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~r1_yellow_0(esk2_0,X1)|~m1_subset_1(esk1_3(esk2_0,X1,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_15]), c_0_11]), c_0_12])])).
cnf(c_0_53,plain,(X2=k1_yellow_0(X1,X3)|m1_subset_1(esk1_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_54,negated_conjecture,(k1_yellow_0(esk2_0,esk6_0)!=esk3_0|~r1_yellow_0(esk2_0,esk6_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_48]), c_0_49]), c_0_50]), c_0_45])).
cnf(c_0_55,negated_conjecture,(r1_yellow_0(esk2_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_29]), c_0_30]), c_0_16]), c_0_11]), c_0_12])])).
cnf(c_0_56,negated_conjecture,(k1_yellow_0(esk2_0,esk4_0)=esk3_0|k1_yellow_0(esk2_0,X1)=esk3_0|~r2_lattice3(esk2_0,esk6_0,esk1_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~r1_yellow_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_11]), c_0_12])])).
cnf(c_0_57,plain,(X2=k1_yellow_0(X1,X3)|r2_lattice3(X1,X3,esk1_3(X1,X3,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_58,negated_conjecture,(k1_yellow_0(esk2_0,esk6_0)!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_59,negated_conjecture,(k1_yellow_0(esk2_0,X1)=esk3_0|r1_yellow_0(esk2_0,esk4_0)|~r2_lattice3(esk2_0,esk6_0,esk1_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~r1_yellow_0(esk2_0,X1)|~m1_subset_1(esk1_3(esk2_0,X1,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_27]), c_0_11]), c_0_12])])).
cnf(c_0_60,negated_conjecture,(k1_yellow_0(esk2_0,esk4_0)=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_30]), c_0_55]), c_0_11]), c_0_12])]), c_0_58])).
cnf(c_0_61,negated_conjecture,(k1_yellow_0(esk2_0,X1)=esk3_0|r1_yellow_0(esk2_0,esk4_0)|~r2_lattice3(esk2_0,esk6_0,esk1_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~r1_yellow_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_53]), c_0_11]), c_0_12])])).
cnf(c_0_62,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk4_0,X1)|~r1_yellow_0(esk2_0,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_60]), c_0_11]), c_0_12])])).
cnf(c_0_63,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk3_0)|~r1_yellow_0(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_60]), c_0_11]), c_0_12])])).
cnf(c_0_64,negated_conjecture,(r1_yellow_0(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_57]), c_0_30]), c_0_55]), c_0_11]), c_0_12])]), c_0_58])).
cnf(c_0_65,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk4_0,esk3_0)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_62]), c_0_36]), c_0_27]), c_0_19])).
cnf(c_0_66,negated_conjecture,(r2_lattice3(esk2_0,esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_67,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,esk6_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_66])])).
cnf(c_0_68,negated_conjecture,(k1_yellow_0(esk2_0,X1)=esk3_0|~r2_lattice3(esk2_0,esk6_0,esk1_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~r1_yellow_0(esk2_0,X1)|~m1_subset_1(esk1_3(esk2_0,X1,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_67]), c_0_11]), c_0_12])])).
cnf(c_0_69,negated_conjecture,(k1_yellow_0(esk2_0,X1)=esk3_0|~r2_lattice3(esk2_0,esk6_0,esk1_3(esk2_0,X1,esk3_0))|~r2_lattice3(esk2_0,X1,esk3_0)|~r1_yellow_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_53]), c_0_11]), c_0_12])])).
cnf(c_0_70,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_57]), c_0_30]), c_0_55]), c_0_11]), c_0_12])]), c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 64
# Proof object formula steps           : 7
# Proof object conjectures             : 57
# Proof object clause conjectures      : 54
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 130
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 124
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 111
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 24
# Backward-rewritten                   : 26
# Generated clauses                    : 90
# ...of the previous two non-trivial   : 85
# Contextual simplify-reflections      : 54
# Paramodulations                      : 88
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 37
# Current number of archived formulas  : 0
# Current number of archived clauses   : 79
# Clause-clause subsumption calls (NU) : 470
# Rec. Clause-clause subsumption calls : 154
# Non-unit clause-clause subsumptions  : 90
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5429

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43120 pages
