# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t30_filter_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>((r6_binop_1(X1,X3,X4)&r6_binop_1(X2,X5,X6))<=>r6_binop_1(k2_zfmisc_1(X1,X2),k6_filter_1(X1,X2,X3,X5),k6_filter_1(X1,X2,X4,X6)))))))))), file('out/t30_filter_1__4_premises__rerun__rerun.E_input', t30_filter_1)).
fof(t29_filter_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>((r6_binop_1(X1,X3,X4)&r6_binop_1(X2,X5,X6))<=>r6_binop_1(k2_zfmisc_1(X1,X2),k6_filter_1(X1,X2,X3,X5),k6_filter_1(X1,X2,X4,X6)))))))))), file('out/t30_filter_1__4_premises__rerun__rerun.E_input', t29_filter_1)).
fof(c_0_2, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>((r6_binop_1(X1,X3,X4)&r6_binop_1(X2,X5,X6))<=>r6_binop_1(k2_zfmisc_1(X1,X2),k6_filter_1(X1,X2,X3,X5),k6_filter_1(X1,X2,X4,X6))))))))))), inference(assume_negation,[status(cth)],[t30_filter_1])).
fof(c_0_3, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((((((((~r6_binop_1(X7,X9,X10)|~r6_binop_1(X8,X11,X12))|r6_binop_1(k2_zfmisc_1(X7,X8),k6_filter_1(X7,X8,X9,X11),k6_filter_1(X7,X8,X10,X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,k2_zfmisc_1(X8,X8),X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X8,X8),X8))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,k2_zfmisc_1(X8,X8),X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X8,X8),X8))))))|(~v1_funct_1(X10)|(~v1_funct_2(X10,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7))))))|(~v1_funct_1(X9)|(~v1_funct_2(X9,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7))))))|v1_xboole_0(X8))|v1_xboole_0(X7))&((((((((r6_binop_1(X7,X9,X10)|~r6_binop_1(k2_zfmisc_1(X7,X8),k6_filter_1(X7,X8,X9,X11),k6_filter_1(X7,X8,X10,X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,k2_zfmisc_1(X8,X8),X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X8,X8),X8))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,k2_zfmisc_1(X8,X8),X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X8,X8),X8))))))|(~v1_funct_1(X10)|(~v1_funct_2(X10,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7))))))|(~v1_funct_1(X9)|(~v1_funct_2(X9,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7))))))|v1_xboole_0(X8))|v1_xboole_0(X7))&(((((((r6_binop_1(X8,X11,X12)|~r6_binop_1(k2_zfmisc_1(X7,X8),k6_filter_1(X7,X8,X9,X11),k6_filter_1(X7,X8,X10,X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,k2_zfmisc_1(X8,X8),X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X8,X8),X8))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,k2_zfmisc_1(X8,X8),X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X8,X8),X8))))))|(~v1_funct_1(X10)|(~v1_funct_2(X10,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7))))))|(~v1_funct_1(X9)|(~v1_funct_2(X9,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7))))))|v1_xboole_0(X8))|v1_xboole_0(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_filter_1])])])])])])])).
fof(c_0_4, negated_conjecture, ((~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))))&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,k2_zfmisc_1(esk2_0,esk2_0),esk2_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0),esk2_0)))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,k2_zfmisc_1(esk2_0,esk2_0),esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0),esk2_0)))))&(((~r6_binop_1(esk1_0,esk3_0,esk4_0)|~r6_binop_1(esk2_0,esk5_0,esk6_0))|~r6_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k6_filter_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_filter_1(esk1_0,esk2_0,esk4_0,esk6_0)))&((r6_binop_1(esk1_0,esk3_0,esk4_0)|r6_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k6_filter_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_filter_1(esk1_0,esk2_0,esk4_0,esk6_0)))&(r6_binop_1(esk2_0,esk5_0,esk6_0)|r6_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k6_filter_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_filter_1(esk1_0,esk2_0,esk4_0,esk6_0)))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_2])])])])])])])).
cnf(c_0_5,plain,(v1_xboole_0(X1)|v1_xboole_0(X2)|r6_binop_1(X2,X5,X6)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X6)|~r6_binop_1(k2_zfmisc_1(X1,X2),k6_filter_1(X1,X2,X3,X5),k6_filter_1(X1,X2,X4,X6))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(r6_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k6_filter_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_filter_1(esk1_0,esk2_0,esk4_0,esk6_0))|r6_binop_1(esk2_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,negated_conjecture,(m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,negated_conjecture,(v1_funct_2(esk6_0,k2_zfmisc_1(esk2_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(v1_funct_2(esk5_0,k2_zfmisc_1(esk2_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13,negated_conjecture,(v1_funct_2(esk4_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14,negated_conjecture,(v1_funct_2(esk3_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,negated_conjecture,(v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_19,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_21,plain,(v1_xboole_0(X1)|v1_xboole_0(X2)|r6_binop_1(X1,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X6)|~r6_binop_1(k2_zfmisc_1(X1,X2),k6_filter_1(X1,X2,X3,X5),k6_filter_1(X1,X2,X4,X6))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_22,negated_conjecture,(r6_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k6_filter_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_filter_1(esk1_0,esk2_0,esk4_0,esk6_0))|r6_binop_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_23,negated_conjecture,(~r6_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k6_filter_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_filter_1(esk1_0,esk2_0,esk4_0,esk6_0))|~r6_binop_1(esk2_0,esk5_0,esk6_0)|~r6_binop_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_24,negated_conjecture,(r6_binop_1(esk2_0,esk5_0,esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18])]), c_0_19]), c_0_20])).
cnf(c_0_25,negated_conjecture,(r6_binop_1(esk1_0,esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18])]), c_0_19]), c_0_20])).
cnf(c_0_26,negated_conjecture,(~r6_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k6_filter_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_filter_1(esk1_0,esk2_0,esk4_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])]), c_0_25])])).
cnf(c_0_27,plain,(v1_xboole_0(X1)|v1_xboole_0(X2)|r6_binop_1(k2_zfmisc_1(X1,X2),k6_filter_1(X1,X2,X3,X5),k6_filter_1(X1,X2,X4,X6))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X6)|~r6_binop_1(X2,X5,X6)|~r6_binop_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_28,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_24]), c_0_25]), c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18])]), c_0_19]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 24
# Proof object formula steps           : 5
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 43
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 43
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 9
# ...of the previous two non-trivial   : 5
# Contextual simplify-reflections      : 0
# Paramodulations                      : 9
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 20
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 4
# Rec. Clause-clause subsumption calls : 4
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2826

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 43128 pages
