# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t30_lattice3, conjecture, (![X1]:![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>(r4_lattice3(X2,X3,X1)<=>r2_lattice3(k3_lattice3(X2),X1,k4_lattice3(X2,X3)))))), file('out/t30_lattice3__16_premises__rerun.E_input', t30_lattice3)).
fof(t7_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)<=>r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3))))))), file('out/t30_lattice3__16_premises__rerun.E_input', t7_lattice3)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t30_lattice3__16_premises__rerun.E_input', redefinition_r3_orders_2)).
fof(dt_k4_lattice3, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k4_lattice3(X1,X2),u1_struct_0(k3_lattice3(X1))))), file('out/t30_lattice3__16_premises__rerun.E_input', dt_k4_lattice3)).
fof(fc4_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>(~(v2_struct_0(k3_lattice3(X1)))&(v1_orders_2(k3_lattice3(X1))&(v3_orders_2(k3_lattice3(X1))&(v4_orders_2(k3_lattice3(X1))&v5_orders_2(k3_lattice3(X1)))))))), file('out/t30_lattice3__16_premises__rerun.E_input', fc4_lattice3)).
fof(dt_k3_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>(v1_orders_2(k3_lattice3(X1))&(v3_orders_2(k3_lattice3(X1))&(v4_orders_2(k3_lattice3(X1))&(v5_orders_2(k3_lattice3(X1))&l1_orders_2(k3_lattice3(X1)))))))), file('out/t30_lattice3__16_premises__rerun.E_input', dt_k3_lattice3)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t30_lattice3__16_premises__rerun.E_input', d9_lattice3)).
fof(d3_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k4_lattice3(X1,X2)=X2))), file('out/t30_lattice3__16_premises__rerun.E_input', d3_lattice3)).
fof(redefinition_r3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3)))), file('out/t30_lattice3__16_premises__rerun.E_input', redefinition_r3_lattices)).
fof(d17_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2))))))), file('out/t30_lattice3__16_premises__rerun.E_input', d17_lattice3)).
fof(dt_k5_lattice3, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))&m1_subset_1(X2,u1_struct_0(k3_lattice3(X1))))=>m1_subset_1(k5_lattice3(X1,X2),u1_struct_0(X1)))), file('out/t30_lattice3__16_premises__rerun.E_input', dt_k5_lattice3)).
fof(d4_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))=>k5_lattice3(X1,X2)=X2))), file('out/t30_lattice3__16_premises__rerun.E_input', d4_lattice3)).
fof(cc1_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1)))))))))), file('out/t30_lattice3__16_premises__rerun.E_input', cc1_lattices)).
fof(free_g1_orders_2, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3]:![X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4)))), file('out/t30_lattice3__16_premises__rerun.E_input', free_g1_orders_2)).
fof(dt_k2_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>(v1_partfun1(k2_lattice3(X1),u1_struct_0(X1))&(v1_relat_2(k2_lattice3(X1))&(v4_relat_2(k2_lattice3(X1))&(v8_relat_2(k2_lattice3(X1))&m1_subset_1(k2_lattice3(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))))))), file('out/t30_lattice3__16_premises__rerun.E_input', dt_k2_lattice3)).
fof(d2_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>k3_lattice3(X1)=g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))), file('out/t30_lattice3__16_premises__rerun.E_input', d2_lattice3)).
fof(abstractness_v1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))), file('out/t30_lattice3__16_premises__rerun.E_input', abstractness_v1_orders_2)).
fof(c_0_17, negated_conjecture, (~(![X1]:![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>(r4_lattice3(X2,X3,X1)<=>r2_lattice3(k3_lattice3(X2),X1,k4_lattice3(X2,X3))))))), inference(assume_negation,[status(cth)],[t30_lattice3])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:(((((~r3_lattices(X4,X5,X6)|r3_orders_2(k3_lattice3(X4),k4_lattice3(X4,X5),k4_lattice3(X4,X6)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4))))&((((~r3_orders_2(k3_lattice3(X4),k4_lattice3(X4,X5),k4_lattice3(X4,X6))|r3_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_lattice3])])])])])])])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_20, plain, (![X3]:![X4]:(((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|~m1_subset_1(X4,u1_struct_0(X3)))|m1_subset_1(k4_lattice3(X3,X4),u1_struct_0(k3_lattice3(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_lattice3])])])).
fof(c_0_21, plain, (![X2]:((~v2_struct_0(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v1_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v3_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v4_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&(v5_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_lattice3])])])])).
fof(c_0_22, plain, (![X2]:((v1_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v3_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v4_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v5_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&(l1_orders_2(k3_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_lattice3])])])])).
fof(c_0_23, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk4_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk4_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
fof(c_0_24, negated_conjecture, (((~v2_struct_0(esk2_0)&(v10_lattices(esk2_0)&l3_lattices(esk2_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&((~r4_lattice3(esk2_0,esk3_0,esk1_0)|~r2_lattice3(k3_lattice3(esk2_0),esk1_0,k4_lattice3(esk2_0,esk3_0)))&(r4_lattice3(esk2_0,esk3_0,esk1_0)|r2_lattice3(k3_lattice3(esk2_0),esk1_0,k4_lattice3(esk2_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
cnf(c_0_25,plain,(v2_struct_0(X1)|r3_lattices(X1,X2,X3)|~l3_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(v2_struct_0(X2)|r3_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r1_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(m1_subset_1(k4_lattice3(X1,X2),u1_struct_0(k3_lattice3(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(v2_struct_0(X1)|v3_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,plain,(v2_struct_0(X1)|l1_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,plain,(v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v2_struct_0(k3_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,negated_conjecture,(r2_lattice3(k3_lattice3(esk2_0),esk1_0,k4_lattice3(esk2_0,esk3_0))|r4_lattice3(esk2_0,esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_34,negated_conjecture,(r1_orders_2(k3_lattice3(esk2_0),X1,k4_lattice3(esk2_0,esk3_0))|r4_lattice3(esk2_0,esk3_0,esk1_0)|~r2_hidden(X1,esk1_0)|~l1_orders_2(k3_lattice3(esk2_0))|~m1_subset_1(k4_lattice3(esk2_0,esk3_0),u1_struct_0(k3_lattice3(esk2_0)))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(esk2_0)))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_35,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36,negated_conjecture,(l3_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37,negated_conjecture,(v10_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_39,negated_conjecture,(r3_lattices(esk2_0,X1,esk3_0)|r4_lattice3(esk2_0,esk3_0,esk1_0)|~r2_hidden(k4_lattice3(esk2_0,X1),esk1_0)|~l1_orders_2(k3_lattice3(esk2_0))|~m1_subset_1(k4_lattice3(esk2_0,esk3_0),u1_struct_0(k3_lattice3(esk2_0)))|~m1_subset_1(k4_lattice3(esk2_0,X1),u1_struct_0(k3_lattice3(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_40,negated_conjecture,(r3_lattices(esk2_0,X1,esk3_0)|r4_lattice3(esk2_0,esk3_0,esk1_0)|~r2_hidden(k4_lattice3(esk2_0,X1),esk1_0)|~l1_orders_2(k3_lattice3(esk2_0))|~m1_subset_1(k4_lattice3(esk2_0,X1),u1_struct_0(k3_lattice3(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_27]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_41, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|(~m1_subset_1(X4,u1_struct_0(X3))|k4_lattice3(X3,X4)=X4))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattice3])])])])])])).
cnf(c_0_42,negated_conjecture,(r3_lattices(esk2_0,X1,esk3_0)|r4_lattice3(esk2_0,esk3_0,esk1_0)|~r2_hidden(k4_lattice3(esk2_0,X1),esk1_0)|~l1_orders_2(k3_lattice3(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_27]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_43,plain,(k4_lattice3(X1,X2)=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_44, plain, (![X4]:![X5]:![X6]:(((~r3_lattices(X4,X5,X6)|r1_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_lattices(X4,X5,X6)|r3_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_45,negated_conjecture,(r3_lattices(esk2_0,X1,esk3_0)|r4_lattice3(esk2_0,esk3_0,esk1_0)|~r2_hidden(X1,esk1_0)|~l1_orders_2(k3_lattice3(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_46, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:((((~r4_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|r1_lattices(X5,X8,X6))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((m1_subset_1(esk5_3(X5,X6,X7),u1_struct_0(X5))|r4_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((r2_hidden(esk5_3(X5,X6,X7),X7)|r4_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((~r1_lattices(X5,esk5_3(X5,X6,X7),X6)|r4_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])])).
cnf(c_0_47,plain,(v2_struct_0(X2)|r1_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r3_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_48,negated_conjecture,(r3_lattices(esk2_0,X1,esk3_0)|r4_lattice3(esk2_0,esk3_0,esk1_0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_29]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_49, plain, (![X3]:![X4]:(((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|~m1_subset_1(X4,u1_struct_0(k3_lattice3(X3))))|m1_subset_1(k5_lattice3(X3,X4),u1_struct_0(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattice3])])])).
fof(c_0_50, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3)))|(~m1_subset_1(X4,u1_struct_0(k3_lattice3(X3)))|k5_lattice3(X3,X4)=X4))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_lattice3])])])])])])).
cnf(c_0_51,plain,(v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,esk5_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52,negated_conjecture,(r1_lattices(esk2_0,X1,esk3_0)|r4_lattice3(esk2_0,esk3_0,esk1_0)|~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~v6_lattices(esk2_0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_35]), c_0_36])]), c_0_38])).
cnf(c_0_53,plain,(m1_subset_1(k5_lattice3(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54,plain,(k5_lattice3(X1,X2)=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_55,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|r4_lattice3(esk2_0,esk3_0,X1)|~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~v6_lattices(esk2_0)|~r2_hidden(esk5_3(esk2_0,esk3_0,X1),esk1_0)|~m1_subset_1(esk5_3(esk2_0,esk3_0,X1),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_35]), c_0_36])]), c_0_38])).
cnf(c_0_56,plain,(v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|r2_hidden(esk5_3(X1,X2,X3),X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57,plain,(v2_struct_0(X1)|r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3))|~l3_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r3_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_58,plain,(m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(k3_lattice3(X2)))|~l3_lattices(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~v6_lattices(esk2_0)|~m1_subset_1(esk5_3(esk2_0,esk3_0,esk1_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_35]), c_0_36])]), c_0_38])).
cnf(c_0_60,plain,(v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_61, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v4_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v5_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v6_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v7_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v8_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&((v9_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_62, plain, (![X5]:![X6]:![X7]:![X8]:(((X5=X7|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5))))&((X6=X8|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])])])).
fof(c_0_63, plain, (![X2]:((v1_partfun1(k2_lattice3(X2),u1_struct_0(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v1_relat_2(k2_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v4_relat_2(k2_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&((v8_relat_2(k2_lattice3(X2))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))&(m1_subset_1(k2_lattice3(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|(v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_lattice3])])])])).
cnf(c_0_64,plain,(r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_57, c_0_43])).
cnf(c_0_65,plain,(m1_subset_1(k4_lattice3(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_58, c_0_27])).
cnf(c_0_66,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~v6_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_35]), c_0_36])]), c_0_38])).
cnf(c_0_67,plain,(v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68,plain,(X2=X3|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|g1_orders_2(X2,X1)!=g1_orders_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69,plain,(v2_struct_0(X1)|m1_subset_1(k2_lattice3(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_70, plain, (![X2]:((v2_struct_0(X2)|(~v10_lattices(X2)|~l3_lattices(X2)))|k3_lattice3(X2)=g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_lattice3])])])).
cnf(c_0_71,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,k4_lattice3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_64]), c_0_65])).
cnf(c_0_72,plain,(v2_struct_0(X2)|r3_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r1_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_73,plain,(v2_struct_0(X1)|v6_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_74,plain,(v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_75,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|~v8_lattices(esk2_0)|~v6_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_76,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_77,plain,(u1_struct_0(X1)=X2|v2_struct_0(X1)|g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))!=g1_orders_2(X2,X3)|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_78,plain,(k3_lattice3(X1)=g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))|v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_79, plain, (![X2]:(~l1_orders_2(X2)|(~v1_orders_2(X2)|X2=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
cnf(c_0_80,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_25, c_0_43])).
cnf(c_0_81,plain,(r3_orders_2(k3_lattice3(X1),X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_64, c_0_43])).
cnf(c_0_82,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,k4_lattice3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_65]), c_0_73]), c_0_74]), c_0_67])).
cnf(c_0_83,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)|~v6_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_74]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_84,plain,(r1_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_57]), c_0_27]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_85,plain,(u1_struct_0(X1)=X2|v2_struct_0(X1)|k3_lattice3(X1)!=g1_orders_2(X2,X3)|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_86,plain,(X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_87,plain,(v2_struct_0(X1)|v1_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_88,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,k4_lattice3(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_65])).
cnf(c_0_89,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_82, c_0_43])).
cnf(c_0_90,plain,(v2_struct_0(X1)|r1_lattices(X1,X3,X2)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r4_lattice3(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_91,negated_conjecture,(r4_lattice3(esk2_0,esk3_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_73]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_92,plain,(r1_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_84, c_0_43])).
cnf(c_0_93,plain,(r2_lattice3(X1,X3,X2)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_94,plain,(u1_struct_0(k3_lattice3(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86])]), c_0_29]), c_0_87])).
cnf(c_0_95,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,k4_lattice3(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_65])).
cnf(c_0_96,negated_conjecture,(r1_lattices(esk2_0,X1,esk3_0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_35]), c_0_36])]), c_0_38])).
cnf(c_0_97,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk4_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_98,plain,(r1_orders_2(k3_lattice3(X1),X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_92, c_0_43])).
cnf(c_0_99,plain,(r2_lattice3(k3_lattice3(X1),X2,X3)|m1_subset_1(esk4_3(k3_lattice3(X1),X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_29])).
cnf(c_0_100,negated_conjecture,(r3_lattices(esk2_0,X1,esk3_0)|~r2_hidden(k4_lattice3(esk2_0,X1),esk1_0)|~m1_subset_1(k4_lattice3(esk2_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_101,negated_conjecture,(~r2_lattice3(k3_lattice3(esk2_0),esk1_0,k4_lattice3(esk2_0,esk3_0))|~r4_lattice3(esk2_0,esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_102,plain,(r2_lattice3(k3_lattice3(X1),X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,esk4_3(k3_lattice3(X1),X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99]), c_0_58]), c_0_29])).
cnf(c_0_103,negated_conjecture,(r3_lattices(esk2_0,X1,esk3_0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_43]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_104,negated_conjecture,(~r2_lattice3(k3_lattice3(esk2_0),esk1_0,esk3_0)|~r4_lattice3(esk2_0,esk3_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_43]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_105,negated_conjecture,(r2_lattice3(k3_lattice3(esk2_0),X1,esk3_0)|~r2_hidden(esk4_3(k3_lattice3(esk2_0),X1,esk3_0),esk1_0)|~m1_subset_1(esk4_3(k3_lattice3(esk2_0),X1,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(esk3_0,u1_struct_0(k3_lattice3(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_106,plain,(r2_lattice3(X1,X3,X2)|r2_hidden(esk4_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_107,negated_conjecture,(~r2_lattice3(k3_lattice3(esk2_0),esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_91])])).
cnf(c_0_108,negated_conjecture,(~l1_orders_2(k3_lattice3(esk2_0))|~m1_subset_1(esk4_3(k3_lattice3(esk2_0),esk1_0,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(esk3_0,u1_struct_0(k3_lattice3(esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])).
cnf(c_0_109,negated_conjecture,(~l1_orders_2(k3_lattice3(esk2_0))|~m1_subset_1(esk3_0,u1_struct_0(k3_lattice3(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_99]), c_0_35]), c_0_36]), c_0_37])]), c_0_107]), c_0_38])).
cnf(c_0_110,negated_conjecture,(~m1_subset_1(esk3_0,u1_struct_0(k3_lattice3(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_29]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_111,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_94]), c_0_35]), c_0_36]), c_0_37])]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 112
# Proof object clause steps            : 77
# Proof object formula steps           : 35
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 114
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 49
# Processed clauses                    : 1451
# ...of these trivial                  : 0
# ...subsumed                          : 974
# ...remaining for further processing  : 477
# Other redundant clauses eliminated   : 16
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 118
# Backward-rewritten                   : 12
# Generated clauses                    : 2339
# ...of the previous two non-trivial   : 2251
# Contextual simplify-reflections      : 1417
# Paramodulations                      : 2297
# Factorizations                       : 0
# Equation resolutions                 : 42
# Current number of processed clauses  : 302
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 294
# Current number of unprocessed clauses: 731
# ...number of literals in the above   : 11263
# Current number of archived formulas  : 0
# Current number of archived clauses   : 175
# Clause-clause subsumption calls (NU) : 207918
# Rec. Clause-clause subsumption calls : 13525
# Non-unit clause-clause subsumptions  : 2508
# Unit Clause-clause subsumption calls : 183
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 144541

# -------------------------------------------------
# User time                : 0.200 s
# System time              : 0.000 s
# Total time               : 0.200 s
# Maximum resident set size: 45352 pages
