# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_wellord2, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3]:![X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4))))))), file('out/t31_wellord2__4_premises__rerun.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, (![X1]:v1_relat_1(k1_wellord2(X1))), file('out/t31_wellord2__4_premises__rerun.E_input', dt_k1_wellord2)).
fof(t31_wellord2, conjecture, (![X1]:r4_relat_2(k1_wellord2(X1),X1)), file('out/t31_wellord2__4_premises__rerun.E_input', t31_wellord2)).
fof(d12_relat_2, axiom, (![X1]:(v1_relat_1(X1)=>(v4_relat_2(X1)<=>r4_relat_2(X1,k3_relat_1(X1))))), file('out/t31_wellord2__4_premises__rerun.E_input', d12_relat_2)).
fof(t5_wellord2, axiom, (![X1]:v4_relat_2(k1_wellord2(X1))), file('out/t31_wellord2__4_premises__rerun.E_input', t5_wellord2)).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:((((k3_relat_1(X6)=X5|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&(((((~r2_hidden(k4_tarski(X7,X8),X6)|r1_tarski(X7,X8))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&((((~r1_tarski(X7,X8)|r2_hidden(k4_tarski(X7,X8),X6))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))))&(((((r2_hidden(esk1_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&(((r2_hidden(esk2_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))&(((((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|~r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])])])).
fof(c_0_6, plain, (![X2]:v1_relat_1(k1_wellord2(X2))), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_7, negated_conjecture, (~(![X1]:r4_relat_2(k1_wellord2(X1),X1))), inference(assume_negation,[status(cth)],[t31_wellord2])).
fof(c_0_8, plain, (![X2]:(((~v4_relat_2(X2)|r4_relat_2(X2,k3_relat_1(X2)))|~v1_relat_1(X2))&((~r4_relat_2(X2,k3_relat_1(X2))|v4_relat_2(X2))|~v1_relat_1(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_relat_2])])])).
cnf(c_0_9,plain,(k3_relat_1(X1)=X2|~v1_relat_1(X1)|X1!=k1_wellord2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, plain, (![X2]:v4_relat_2(k1_wellord2(X2))), inference(variable_rename,[status(thm)],[t5_wellord2])).
fof(c_0_12, negated_conjecture, (~r4_relat_2(k1_wellord2(esk3_0),esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_13,plain,(r4_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)|~v4_relat_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_9]), c_0_10])])).
cnf(c_0_15,plain,(v4_relat_2(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(~r4_relat_2(k1_wellord2(esk3_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(r4_relat_2(k1_wellord2(X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_10])])).
cnf(c_0_18,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 19
# Proof object clause steps            : 8
# Proof object formula steps           : 11
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 1
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 26
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 26
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 10
# ...of the previous two non-trivial   : 8
# Contextual simplify-reflections      : 0
# Paramodulations                      : 3
# Factorizations                       : 0
# Equation resolutions                 : 7
# Current number of processed clauses  : 6
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 2
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 22
# Current number of archived formulas  : 0
# Current number of archived clauses   : 13
# Clause-clause subsumption calls (NU) : 19
# Rec. Clause-clause subsumption calls : 15
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1000

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 45280 pages
