# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t32_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&(v3_pre_topc(X3,X1)&(v2_tex_2(X2,X1)&v2_tex_2(X3,X1))))=>v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)))))), file('out/t32_tex_2__208_premises__rerun__rerun.E_input', t32_tex_2)).
fof(t24_tex_2, axiom, (![X1]:(l1_pre_topc(X1)=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&v3_pre_topc(X3,X1))=>(v3_pre_topc(k9_subset_1(u1_struct_0(X1),X2,X3),X1)&v3_pre_topc(k4_subset_1(u1_struct_0(X1),X2,X3),X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&(v3_pre_topc(X3,X1)&(v2_tex_2(X2,X1)&v2_tex_2(X3,X1))))=>v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1))))))), file('out/t32_tex_2__208_premises__rerun__rerun.E_input', t24_tex_2)).
fof(fc7_tops_1, axiom, (![X1]:![X2]:![X3]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&((v3_pre_topc(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(v3_pre_topc(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))=>v3_pre_topc(k2_xboole_0(X2,X3),X1))), file('out/t32_tex_2__208_premises__rerun__rerun.E_input', fc7_tops_1)).
fof(redefinition_k4_subset_1, axiom, (![X1]:![X2]:![X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3))), file('out/t32_tex_2__208_premises__rerun__rerun.E_input', redefinition_k4_subset_1)).
fof(fc6_tops_1, axiom, (![X1]:![X2]:![X3]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&((v3_pre_topc(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(v3_pre_topc(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))=>v3_pre_topc(k3_xboole_0(X2,X3),X1))), file('out/t32_tex_2__208_premises__rerun__rerun.E_input', fc6_tops_1)).
fof(redefinition_k9_subset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3))), file('out/t32_tex_2__208_premises__rerun__rerun.E_input', redefinition_k9_subset_1)).
fof(c_0_6, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&(v3_pre_topc(X3,X1)&(v2_tex_2(X2,X1)&v2_tex_2(X3,X1))))=>v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1))))))), inference(assume_negation,[status(cth)],[t32_tex_2])).
fof(c_0_7, plain, (![X4]:![X7]:![X8]:(((m1_subset_1(esk4_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4))&(((m1_subset_1(esk5_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4))&((((v3_pre_topc(esk4_1(X4),X4)|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4))&((v3_pre_topc(esk5_1(X4),X4)|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4)))&(((~v3_pre_topc(k9_subset_1(u1_struct_0(X4),esk4_1(X4),esk5_1(X4)),X4)|~v3_pre_topc(k4_subset_1(u1_struct_0(X4),esk4_1(X4),esk5_1(X4)),X4))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_tex_2])])])])])])])).
fof(c_0_8, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((v3_pre_topc(esk2_0,esk1_0)&(v3_pre_topc(esk3_0,esk1_0)&(v2_tex_2(esk2_0,esk1_0)&v2_tex_2(esk3_0,esk1_0))))&~v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0),esk1_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
cnf(c_0_9,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|m1_subset_1(esk5_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11,negated_conjecture,(v2_tex_2(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,negated_conjecture,(v3_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|v3_pre_topc(esk5_1(X1),X1)|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(((~v2_pre_topc(X4)|~l1_pre_topc(X4))|((~v3_pre_topc(X5,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(~v3_pre_topc(X6,X4)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4))))))|v3_pre_topc(k2_xboole_0(X5,X6),X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_tops_1])])).
cnf(c_0_16,negated_conjecture,(v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),X1,esk3_0),esk1_0)|m1_subset_1(esk5_1(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_tex_2(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18,negated_conjecture,(v2_tex_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19,negated_conjecture,(v3_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20,negated_conjecture,(~v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21,negated_conjecture,(v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),X1,esk3_0),esk1_0)|v3_pre_topc(esk5_1(esk1_0),esk1_0)|~v2_tex_2(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_22,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|v3_pre_topc(esk4_1(X1),X1)|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_24, plain, (![X4]:![X5]:![X6]:((~m1_subset_1(X5,k1_zfmisc_1(X4))|~m1_subset_1(X6,k1_zfmisc_1(X4)))|k4_subset_1(X4,X5,X6)=k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
cnf(c_0_25,plain,(v3_pre_topc(k2_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X1,X3)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26,negated_conjecture,(m1_subset_1(esk5_1(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_27,negated_conjecture,(v3_pre_topc(esk5_1(esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_28,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_29,negated_conjecture,(v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),X1,esk3_0),esk1_0)|m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_tex_2(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_30,negated_conjecture,(v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),X1,esk3_0),esk1_0)|v3_pre_topc(esk4_1(esk1_0),esk1_0)|~v2_tex_2(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_31,plain,(k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, (![X4]:![X5]:![X6]:(((~v2_pre_topc(X4)|~l1_pre_topc(X4))|((~v3_pre_topc(X5,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(~v3_pre_topc(X6,X4)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4))))))|v3_pre_topc(k3_xboole_0(X5,X6),X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_tops_1])])).
cnf(c_0_33,negated_conjecture,(v3_pre_topc(k2_xboole_0(X1,esk5_1(esk1_0)),esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_13]), c_0_28])])).
cnf(c_0_34,negated_conjecture,(m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_35,negated_conjecture,(v3_pre_topc(esk4_1(esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_36,negated_conjecture,(k2_xboole_0(X1,esk5_1(esk1_0))=k4_subset_1(u1_struct_0(esk1_0),X1,esk5_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_31, c_0_26])).
fof(c_0_37, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(X4))|k9_subset_1(X4,X5,X6)=k3_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
cnf(c_0_38,plain,(v3_pre_topc(k3_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X1,X3)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39,negated_conjecture,(v3_pre_topc(k2_xboole_0(esk4_1(esk1_0),esk5_1(esk1_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_40,negated_conjecture,(k2_xboole_0(esk4_1(esk1_0),esk5_1(esk1_0))=k4_subset_1(u1_struct_0(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0))), inference(spm,[status(thm)],[c_0_36, c_0_34])).
cnf(c_0_41,plain,(k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_42,negated_conjecture,(v3_pre_topc(k3_xboole_0(X1,esk5_1(esk1_0)),esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_26]), c_0_27]), c_0_13]), c_0_28])])).
cnf(c_0_43,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(k4_subset_1(u1_struct_0(X1),esk4_1(X1),esk5_1(X1)),X1)|~v3_pre_topc(k9_subset_1(u1_struct_0(X1),esk4_1(X1),esk5_1(X1)),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_44,negated_conjecture,(v3_pre_topc(k4_subset_1(u1_struct_0(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),esk1_0)), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45,negated_conjecture,(k9_subset_1(u1_struct_0(esk1_0),X1,esk5_1(esk1_0))=k3_xboole_0(X1,esk5_1(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_26])).
cnf(c_0_46,negated_conjecture,(v3_pre_topc(k3_xboole_0(esk4_1(esk1_0),esk5_1(esk1_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_34]), c_0_35])])).
cnf(c_0_47,negated_conjecture,(v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),X1,X2),esk1_0)|~v2_tex_2(X2,esk1_0)|~v2_tex_2(X1,esk1_0)|~v3_pre_topc(X2,esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46]), c_0_13])])).
cnf(c_0_48,negated_conjecture,(v2_tex_2(k4_subset_1(u1_struct_0(esk1_0),X1,esk3_0),esk1_0)|~v2_tex_2(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_10]), c_0_11]), c_0_12])])).
cnf(c_0_49,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_17]), c_0_18]), c_0_19])]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 37
# Proof object formula steps           : 13
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 56
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 128
# ...of these trivial                  : 5
# ...subsumed                          : 0
# ...remaining for further processing  : 123
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 14
# Generated clauses                    : 98
# ...of the previous two non-trivial   : 100
# Contextual simplify-reflections      : 0
# Paramodulations                      : 98
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 90
#    Positive orientable unit clauses  : 64
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 14
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 544
# Rec. Clause-clause subsumption calls : 84
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 55
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 44
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4838

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 44884 pages
