# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_tex_2, axiom, (![X1]:(l1_pre_topc(X1)=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&v3_pre_topc(X3,X1))=>(v3_pre_topc(k9_subset_1(u1_struct_0(X1),X2,X3),X1)&v3_pre_topc(k4_subset_1(u1_struct_0(X1),X2,X3),X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&(v3_pre_topc(X3,X1)&(v2_tex_2(X2,X1)&v2_tex_2(X3,X1))))=>v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1))))))), file('out/t32_tex_2__3_premises__rerun.E_input', t24_tex_2)).
fof(t11_tops_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&v3_pre_topc(X3,X1))=>v3_pre_topc(k9_subset_1(u1_struct_0(X1),X2,X3),X1)))))), file('out/t32_tex_2__3_premises__rerun.E_input', t11_tops_1)).
fof(t32_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&(v3_pre_topc(X3,X1)&(v2_tex_2(X2,X1)&v2_tex_2(X3,X1))))=>v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)))))), file('out/t32_tex_2__3_premises__rerun.E_input', t32_tex_2)).
fof(t10_tops_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&v3_pre_topc(X3,X1))=>v3_pre_topc(k4_subset_1(u1_struct_0(X1),X2,X3),X1)))))), file('out/t32_tex_2__3_premises__rerun.E_input', t10_tops_1)).
fof(c_0_4, plain, (![X4]:![X7]:![X8]:(((m1_subset_1(esk1_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4))&(((m1_subset_1(esk2_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4))&((((v3_pre_topc(esk1_1(X4),X4)|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4))&((v3_pre_topc(esk2_1(X4),X4)|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4)))&(((~v3_pre_topc(k9_subset_1(u1_struct_0(X4),esk1_1(X4),esk2_1(X4)),X4)|~v3_pre_topc(k4_subset_1(u1_struct_0(X4),esk1_1(X4),esk2_1(X4)),X4))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X7,X4)|(~v3_pre_topc(X8,X4)|(~v2_tex_2(X7,X4)|~v2_tex_2(X8,X4))))|v2_tex_2(k4_subset_1(u1_struct_0(X4),X7,X8),X4)))))|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_tex_2])])])])])])])).
fof(c_0_5, plain, (![X4]:![X5]:![X6]:((~v2_pre_topc(X4)|~l1_pre_topc(X4))|(~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X5,X4)|~v3_pre_topc(X6,X4))|v3_pre_topc(k9_subset_1(u1_struct_0(X4),X5,X6),X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_tops_1])])])])])).
fof(c_0_6, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&(v3_pre_topc(X3,X1)&(v2_tex_2(X2,X1)&v2_tex_2(X3,X1))))=>v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1))))))), inference(assume_negation,[status(cth)],[t32_tex_2])).
cnf(c_0_7,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(k4_subset_1(u1_struct_0(X1),esk1_1(X1),esk2_1(X1)),X1)|~v3_pre_topc(k9_subset_1(u1_struct_0(X1),esk1_1(X1),esk2_1(X1)),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(v3_pre_topc(k9_subset_1(u1_struct_0(X1),X2,X3),X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|m1_subset_1(esk1_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|m1_subset_1(esk2_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|v3_pre_topc(esk1_1(X1),X1)|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|v3_pre_topc(esk2_1(X1),X1)|~l1_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:((~v2_pre_topc(X4)|~l1_pre_topc(X4))|(~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|((~v3_pre_topc(X5,X4)|~v3_pre_topc(X6,X4))|v3_pre_topc(k4_subset_1(u1_struct_0(X4),X5,X6),X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_tops_1])])])])])).
fof(c_0_14, negated_conjecture, (((~v2_struct_0(esk3_0)&(v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))&(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk3_0)))&((v3_pre_topc(esk4_0,esk3_0)&(v3_pre_topc(esk5_0,esk3_0)&(v2_tex_2(esk4_0,esk3_0)&v2_tex_2(esk5_0,esk3_0))))&~v2_tex_2(k4_subset_1(u1_struct_0(esk3_0),esk4_0,esk5_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
cnf(c_0_15,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~v2_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(k4_subset_1(u1_struct_0(X1),esk1_1(X1),esk2_1(X1)),X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12])).
cnf(c_0_16,plain,(v3_pre_topc(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(~v2_tex_2(k4_subset_1(u1_struct_0(esk3_0),esk4_0,esk5_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18,plain,(v2_tex_2(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~v2_pre_topc(X1)|~v2_tex_2(X3,X1)|~v2_tex_2(X2,X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_9]), c_0_10]), c_0_11]), c_0_12])).
cnf(c_0_19,negated_conjecture,(v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,negated_conjecture,(v2_tex_2(esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,negated_conjecture,(v2_tex_2(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,negated_conjecture,(v3_pre_topc(esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,negated_conjecture,(v3_pre_topc(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26,negated_conjecture,(l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_27,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 19
# Proof object formula steps           : 9
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 45
# ...of these trivial                  : 3
# ...subsumed                          : 0
# ...remaining for further processing  : 42
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 8
# Paramodulations                      : 18
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 22
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 8
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 105
# Rec. Clause-clause subsumption calls : 10
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2723

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 44868 pages
