# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C45_F1_PI_AE_Q7_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_yellow19, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3)))))), file('out/t33_yellow19__2_premises__rerun.E_input', t33_yellow19)).
fof(l35_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3)))))), file('out/t33_yellow19__2_premises__rerun.E_input', l35_yellow19)).
fof(l36_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>~((r2_hidden(X2,k6_yellow_6(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~(r3_waybel_9(X1,X2,X3))))))=>v1_compts_1(X1)))), file('out/t33_yellow19__2_premises__rerun.E_input', l36_yellow19)).
fof(c_0_3, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3))))))), inference(assume_negation,[status(cth)],[t33_yellow19])).
fof(c_0_4, negated_conjecture, (![X6]:![X7]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((((~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0))&((v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))&((v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0))&(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)))))&((~m1_subset_1(X6,u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X6))|~v1_compts_1(esk1_0)))&(((m1_subset_1(esk3_1(X7),u1_struct_0(esk1_0))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,esk1_0)))))|v1_compts_1(esk1_0))&((r3_waybel_9(esk1_0,X7,esk3_1(X7))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,esk1_0)))))|v1_compts_1(esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])])])).
fof(c_0_5, plain, (![X4]:![X5]:((((m1_subset_1(esk4_2(X4,X5),u1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|~v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((r3_waybel_9(X4,X5,esk4_2(X4,X5))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|~v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l35_yellow19])])])])])])])])).
cnf(c_0_6,negated_conjecture,(~v1_compts_1(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_11, plain, (![X4]:![X6]:((((~v2_struct_0(esk5_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((v4_orders_2(esk5_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((v7_waybel_0(esk5_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((l1_waybel_0(esk5_1(X4),X4)|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))))&(((r2_hidden(esk5_1(X4),k6_yellow_6(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((~m1_subset_1(X6,u1_struct_0(X4))|~r3_waybel_9(X4,esk5_1(X4),X6))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l36_yellow19])])])])])])])])).
cnf(c_0_12,negated_conjecture,(v2_struct_0(X1)|~r3_waybel_9(esk1_0,esk2_0,esk4_2(esk1_0,X1))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_13,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,esk4_2(X1,X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,negated_conjecture,(v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15,negated_conjecture,(v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(~v1_compts_1(esk1_0)|~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18,plain,(v2_struct_0(X1)|v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_waybel_9(X1,esk5_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(X1)|r3_waybel_9(esk1_0,X1,esk3_1(X1))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20,negated_conjecture,(~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_8]), c_0_9])]), c_0_10]), c_0_14]), c_0_15]), c_0_16]), c_0_17])).
cnf(c_0_21,negated_conjecture,(v2_struct_0(esk5_1(esk1_0))|~m1_subset_1(esk3_1(esk5_1(esk1_0)),u1_struct_0(esk1_0))|~l1_waybel_0(esk5_1(esk1_0),esk1_0)|~v7_waybel_0(esk5_1(esk1_0))|~v4_orders_2(esk5_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_8]), c_0_9])]), c_0_10]), c_0_20])).
cnf(c_0_22,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(esk1_0))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_23,negated_conjecture,(v2_struct_0(esk5_1(esk1_0))|~l1_waybel_0(esk5_1(esk1_0),esk1_0)|~v7_waybel_0(esk5_1(esk1_0))|~v4_orders_2(esk5_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_20])).
cnf(c_0_24,plain,(v2_struct_0(X1)|v1_compts_1(X1)|l1_waybel_0(esk5_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,negated_conjecture,(v2_struct_0(esk5_1(esk1_0))|~v7_waybel_0(esk5_1(esk1_0))|~v4_orders_2(esk5_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_8]), c_0_9])]), c_0_20]), c_0_10])).
cnf(c_0_26,plain,(v2_struct_0(X1)|v1_compts_1(X1)|v7_waybel_0(esk5_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_27,negated_conjecture,(v2_struct_0(esk5_1(esk1_0))|~v4_orders_2(esk5_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_8]), c_0_9])]), c_0_20]), c_0_10])).
cnf(c_0_28,plain,(v2_struct_0(X1)|v1_compts_1(X1)|v4_orders_2(esk5_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_29,plain,(v2_struct_0(X1)|v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(esk5_1(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_30,negated_conjecture,(v2_struct_0(esk5_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_8]), c_0_9])]), c_0_20]), c_0_10])).
cnf(c_0_31,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_8]), c_0_9])]), c_0_20]), c_0_10]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 32
# Proof object clause steps            : 25
# Proof object formula steps           : 7
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 43
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 43
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 1
# Generated clauses                    : 10
# ...of the previous two non-trivial   : 7
# Contextual simplify-reflections      : 4
# Paramodulations                      : 10
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 63
# Rec. Clause-clause subsumption calls : 18
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2068

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 43268 pages
