# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_ordinal1, conjecture, (![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2))))), file('out/t33_ordinal1__11_premises__rerun.E_input', t33_ordinal1)).
fof(d1_ordinal1, axiom, (![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), file('out/t33_ordinal1__11_premises__rerun.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t33_ordinal1__11_premises__rerun.E_input', t69_enumset1)).
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t33_ordinal1__11_premises__rerun.E_input', l51_zfmisc_1)).
fof(t11_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3))), file('out/t33_ordinal1__11_premises__rerun.E_input', t11_xboole_1)).
fof(redefinition_r1_ordinal1, axiom, (![X1]:![X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2)))), file('out/t33_ordinal1__11_premises__rerun.E_input', redefinition_r1_ordinal1)).
fof(t29_ordinal1, axiom, (![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1)))), file('out/t33_ordinal1__11_premises__rerun.E_input', t29_ordinal1)).
fof(t26_ordinal1, axiom, (![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1))))), file('out/t33_ordinal1__11_premises__rerun.E_input', t26_ordinal1)).
fof(t13_ordinal1, axiom, (![X1]:![X2]:(r2_hidden(X1,k1_ordinal1(X2))<=>(r2_hidden(X1,X2)|X1=X2))), file('out/t33_ordinal1__11_premises__rerun.E_input', t13_ordinal1)).
fof(t7_ordinal1, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1)))), file('out/t33_ordinal1__11_premises__rerun.E_input', t7_ordinal1)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t33_ordinal1__11_premises__rerun.E_input', d10_xboole_0)).
fof(antisymmetry_r2_hidden, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1)))), file('out/t33_ordinal1__11_premises__rerun.E_input', antisymmetry_r2_hidden)).
fof(c_0_12, negated_conjecture, (~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2)))))), inference(assume_negation,[status(cth)],[t33_ordinal1])).
fof(c_0_13, plain, (![X2]:k1_ordinal1(X2)=k2_xboole_0(X2,k1_tarski(X2))), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_14, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_15, negated_conjecture, ((v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&((~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0))&(r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
cnf(c_0_16,plain,(k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:(~r1_tarski(k2_xboole_0(X4,X5),X6)|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
fof(c_0_20, plain, (![X3]:![X4]:(((~r1_ordinal1(X3,X4)|r1_tarski(X3,X4))|(~v3_ordinal1(X3)|~v3_ordinal1(X4)))&((~r1_tarski(X3,X4)|r1_ordinal1(X3,X4))|(~v3_ordinal1(X3)|~v3_ordinal1(X4))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_21,negated_conjecture,(r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)|r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(r1_tarski(X1,X2)|~r1_tarski(k2_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,plain,(r1_tarski(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_27,negated_conjecture,(v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_28, plain, (![X2]:(~v3_ordinal1(X2)|v3_ordinal1(k1_ordinal1(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
cnf(c_0_29,negated_conjecture,(~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)|~r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_30, plain, (![X3]:![X4]:(~v3_ordinal1(X3)|(~v3_ordinal1(X4)|(r1_ordinal1(X3,X4)|r2_hidden(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])])])).
cnf(c_0_31,plain,(r1_tarski(X1,X2)|~r1_tarski(k3_tarski(k2_tarski(X1,X3)),X2)), inference(rw,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_32,negated_conjecture,(r1_tarski(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))),esk2_0)|r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_33,plain,(v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_34, plain, (![X3]:![X4]:![X3]:![X4]:((~r2_hidden(X3,k1_ordinal1(X4))|(r2_hidden(X3,X4)|X3=X4))&((~r2_hidden(X3,X4)|r2_hidden(X3,k1_ordinal1(X4)))&(X3!=X4|r2_hidden(X3,k1_ordinal1(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_ordinal1])])])])])).
cnf(c_0_35,negated_conjecture,(~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_22]), c_0_23])).
cnf(c_0_36,plain,(r2_hidden(X1,X2)|r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_37, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~r1_tarski(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_38, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_39,negated_conjecture,(r1_tarski(esk1_0,esk2_0)|r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40,plain,(v3_ordinal1(k3_tarski(k2_tarski(X1,k2_tarski(X1,X1))))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_22]), c_0_23])).
cnf(c_0_41,negated_conjecture,(v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_42,plain,(X1=X2|r2_hidden(X1,X2)|~r2_hidden(X1,k1_ordinal1(X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43,negated_conjecture,(r2_hidden(esk2_0,k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))|~r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_27])])).
fof(c_0_44, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~r2_hidden(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_45,plain,(~r1_tarski(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47,negated_conjecture,(r1_tarski(esk1_0,esk2_0)|r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
cnf(c_0_48,plain,(X1=X2|r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k2_tarski(X2,k2_tarski(X2,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_22]), c_0_23])).
cnf(c_0_49,negated_conjecture,(r2_hidden(esk2_0,k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))|~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_40]), c_0_41])])).
cnf(c_0_50,plain,(~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51,plain,(r2_hidden(X1,k1_ordinal1(X2))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_52,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk2_0,k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))|~v3_ordinal1(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))), inference(spm,[status(thm)],[c_0_45, c_0_32])).
cnf(c_0_54,negated_conjecture,(esk2_0=esk1_0|~r1_tarski(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_45])).
cnf(c_0_55,plain,(r1_tarski(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_25, c_0_36])).
cnf(c_0_56,negated_conjecture,(esk2_0=esk1_0|~r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_57,plain,(r2_hidden(X1,k3_tarski(k2_tarski(X2,k2_tarski(X2,X2))))|X1!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_22]), c_0_23])).
cnf(c_0_58,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_59,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk2_0,k3_tarski(k2_tarski(esk1_0,k2_tarski(esk1_0,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_40]), c_0_41])])).
cnf(c_0_60,negated_conjecture,(esk2_0=esk1_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_27]), c_0_41])]), c_0_56])).
cnf(c_0_61,plain,(r2_hidden(X1,k3_tarski(k2_tarski(X1,k2_tarski(X1,X1))))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_62,plain,(~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_45, c_0_58])).
cnf(c_0_63,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60]), c_0_60]), c_0_61])]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 39
# Proof object formula steps           : 25
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 17
# Processed clauses                    : 62
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 59
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 13
# Generated clauses                    : 48
# ...of the previous two non-trivial   : 42
# Contextual simplify-reflections      : 5
# Paramodulations                      : 45
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 11
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 93
# Rec. Clause-clause subsumption calls : 60
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1829

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45508 pages
