# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t34_ordinal1, conjecture, (![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2))))), file('out/t34_ordinal1__16_premises__rerun.E_input', t34_ordinal1)).
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t34_ordinal1__16_premises__rerun.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t34_ordinal1__16_premises__rerun.E_input', t70_enumset1)).
fof(connectedness_r1_ordinal1, axiom, (![X1]:![X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1)))), file('out/t34_ordinal1__16_premises__rerun.E_input', connectedness_r1_ordinal1)).
fof(d3_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2))))), file('out/t34_ordinal1__16_premises__rerun.E_input', d3_xboole_0)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t34_ordinal1__16_premises__rerun.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), file('out/t34_ordinal1__16_premises__rerun.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), file('out/t34_ordinal1__16_premises__rerun.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), file('out/t34_ordinal1__16_premises__rerun.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), file('out/t34_ordinal1__16_premises__rerun.E_input', t75_enumset1)).
fof(d1_ordinal1, axiom, (![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), file('out/t34_ordinal1__16_premises__rerun.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, (![X1]:k2_tarski(X1,X1)=k1_tarski(X1)), file('out/t34_ordinal1__16_premises__rerun.E_input', t69_enumset1)).
fof(t7_ordinal1, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1)))), file('out/t34_ordinal1__16_premises__rerun.E_input', t7_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, (![X1]:![X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2)))), file('out/t34_ordinal1__16_premises__rerun.E_input', redefinition_r1_ordinal1)).
fof(l1_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t34_ordinal1__16_premises__rerun.E_input', l1_zfmisc_1)).
fof(t24_ordinal1, axiom, (![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1)))))))), file('out/t34_ordinal1__16_premises__rerun.E_input', t24_ordinal1)).
fof(d6_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8)))))))))))), file('out/t34_ordinal1__16_premises__rerun.E_input', d6_enumset1)).
fof(c_0_17, negated_conjecture, (~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2)))))), inference(assume_negation,[status(cth)],[t34_ordinal1])).
fof(c_0_18, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_19, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_20, plain, (![X3]:![X4]:((~v3_ordinal1(X3)|~v3_ordinal1(X4))|(r1_ordinal1(X3,X4)|r1_ordinal1(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_ordinal1])])).
fof(c_0_21, negated_conjecture, ((v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&((~r2_hidden(esk1_0,k1_ordinal1(esk2_0))|~r1_ordinal1(esk1_0,esk2_0))&(r2_hidden(esk1_0,k1_ordinal1(esk2_0))|r1_ordinal1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).
fof(c_0_22, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(r2_hidden(X8,X5)|r2_hidden(X8,X6)))|X7!=k2_xboole_0(X5,X6))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X7))|X7!=k2_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|r2_hidden(X8,X7))|X7!=k2_xboole_0(X5,X6))))&((((~r2_hidden(esk3_3(X5,X6,X7),X5)|~r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k2_xboole_0(X5,X6))&((~r2_hidden(esk3_3(X5,X6,X7),X6)|~r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k2_xboole_0(X5,X6)))&((r2_hidden(esk3_3(X5,X6,X7),X7)|(r2_hidden(esk3_3(X5,X6,X7),X5)|r2_hidden(esk3_3(X5,X6,X7),X6)))|X7=k2_xboole_0(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])])).
cnf(c_0_23,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_26, plain, (![X5]:![X6]:![X7]:![X8]:k3_enumset1(X5,X5,X6,X7,X8)=k2_enumset1(X5,X6,X7,X8)), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_27, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k4_enumset1(X6,X6,X7,X8,X9,X10)=k3_enumset1(X6,X7,X8,X9,X10)), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_28, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:k5_enumset1(X7,X7,X8,X9,X10,X11,X12)=k4_enumset1(X7,X8,X9,X10,X11,X12)), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_29, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:k6_enumset1(X8,X8,X9,X10,X11,X12,X13,X14)=k5_enumset1(X8,X9,X10,X11,X12,X13,X14)), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_30, plain, (![X2]:k1_ordinal1(X2)=k2_xboole_0(X2,k1_tarski(X2))), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_31, plain, (![X2]:k2_tarski(X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_32, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~r1_tarski(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_33, plain, (![X3]:![X4]:(((~r1_ordinal1(X3,X4)|r1_tarski(X3,X4))|(~v3_ordinal1(X3)|~v3_ordinal1(X4)))&((~r1_tarski(X3,X4)|r1_ordinal1(X3,X4))|(~v3_ordinal1(X3)|~v3_ordinal1(X4))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_34,plain,(r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_35,negated_conjecture,(v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_36, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])])])).
cnf(c_0_37,plain,(r2_hidden(X4,X3)|r2_hidden(X4,X2)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_38,plain,(k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_39,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_40,plain,(k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41,plain,(k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_42,plain,(k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43,plain,(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44,plain,(k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_45,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_46,plain,(~r1_tarski(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47,plain,(r1_tarski(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48,negated_conjecture,(r1_ordinal1(X1,esk2_0)|r1_ordinal1(esk2_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_49,negated_conjecture,(v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_50,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_51,plain,(r2_hidden(X4,X3)|r2_hidden(X4,X2)|X1!=k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3))|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_52,negated_conjecture,(r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk1_0,k1_ordinal1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_53,plain,(k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54,plain,(~r1_ordinal1(X1,X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55,negated_conjecture,(r1_ordinal1(esk2_0,esk1_0)|r1_ordinal1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_56,plain,(r2_hidden(X4,X1)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_57, plain, (![X3]:![X4]:(~v3_ordinal1(X3)|(~v3_ordinal1(X4)|(r2_hidden(X3,X4)|(X3=X4|r2_hidden(X4,X3)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])])])).
cnf(c_0_58,plain,(r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_45]), c_0_24]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_59,plain,(r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2)))), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_60,negated_conjecture,(r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk1_0,k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53]), c_0_24]), c_0_38]), c_0_39]), c_0_39]), c_0_40]), c_0_40]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_43]), c_0_43])).
cnf(c_0_61,negated_conjecture,(r1_ordinal1(esk2_0,esk1_0)|~r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_49]), c_0_35])])).
cnf(c_0_62,negated_conjecture,(~r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,k1_ordinal1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_63,plain,(r2_hidden(X4,X1)|X1!=k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3))|~r2_hidden(X4,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_64,plain,(r2_hidden(X1,X2)|X2=X1|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65,plain,(~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_58])).
cnf(c_0_66,negated_conjecture,(r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_67,negated_conjecture,(~r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_61]), c_0_35]), c_0_49])])).
cnf(c_0_68,negated_conjecture,(~r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_53]), c_0_24]), c_0_38]), c_0_39]), c_0_39]), c_0_40]), c_0_40]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_43]), c_0_43])).
cnf(c_0_69,plain,(r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_70,negated_conjecture,(X1=esk2_0|r2_hidden(X1,esk2_0)|r2_hidden(esk2_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_64, c_0_35])).
cnf(c_0_71,negated_conjecture,(r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_72,negated_conjecture,(~r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_73,negated_conjecture,(esk1_0=esk2_0|r2_hidden(esk2_0,esk1_0)|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_49])).
cnf(c_0_74,negated_conjecture,(~r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_71]), c_0_49]), c_0_35])])).
cnf(c_0_75,negated_conjecture,(r1_ordinal1(esk2_0,esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_72, c_0_55])).
cnf(c_0_76,negated_conjecture,(esk1_0=esk2_0|r2_hidden(esk1_0,esk2_0)), inference(sr,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_77,negated_conjecture,(~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_75]), c_0_35]), c_0_49])])).
cnf(c_0_78,plain,(r2_hidden(X4,X1)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_79, plain, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:![X20]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:((((~r2_hidden(X20,X19)|(X20=X11|(X20=X12|(X20=X13|(X20=X14|(X20=X15|(X20=X16|(X20=X17|X20=X18))))))))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((X20!=X11|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((X20!=X12|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((X20!=X13|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((X20!=X14|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((X20!=X15|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((X20!=X16|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((X20!=X17|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&((X20!=X18|r2_hidden(X20,X19))|X19!=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))))))))))&((((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X11|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X12|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X13|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X14|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X15|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X16|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&(((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X17|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))&((esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)!=X18|~r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18)))))))))&((r2_hidden(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19),X19)|(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X11|(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X12|(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X13|(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X14|(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X15|(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X16|(esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X17|esk4_9(X11,X12,X13,X14,X15,X16,X17,X18,X19)=X18))))))))|X19=k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])])).
cnf(c_0_80,negated_conjecture,(esk1_0=esk2_0), inference(sr,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_81,negated_conjecture,(r1_ordinal1(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_35])).
cnf(c_0_82,plain,(r2_hidden(X4,X1)|X1!=k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3))|~r2_hidden(X4,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_83,plain,(r2_hidden(X10,X1)|X1!=k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9)|X10!=X2), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_84,negated_conjecture,(~r2_hidden(esk2_0,k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_80]), c_0_81]), c_0_80])])).
cnf(c_0_85,plain,(r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_82])).
cnf(c_0_86,plain,(r2_hidden(X1,k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_83])])).
cnf(c_0_87,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 88
# Proof object clause steps            : 53
# Proof object formula steps           : 35
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 79
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 35
# Processed clauses                    : 112
# ...of these trivial                  : 0
# ...subsumed                          : 17
# ...remaining for further processing  : 95
# Other redundant clauses eliminated   : 20
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 15
# Generated clauses                    : 62
# ...of the previous two non-trivial   : 62
# Contextual simplify-reflections      : 1
# Paramodulations                      : 47
# Factorizations                       : 0
# Equation resolutions                 : 20
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 68
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 226
# Rec. Clause-clause subsumption calls : 168
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 45
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3460

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 47168 pages
