# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t35_yellow_0, conjecture, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:((r1_tarski(X2,X3)&(r2_yellow_0(X1,X2)&r2_yellow_0(X1,X3)))=>r1_orders_2(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))))), file('out/t35_yellow_0__3_premises__rerun.E_input', t35_yellow_0)).
fof(d10_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))))), file('out/t35_yellow_0__3_premises__rerun.E_input', d10_yellow_0)).
fof(d8_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3)))))))), file('out/t35_yellow_0__3_premises__rerun.E_input', d8_yellow_0)).
fof(t9_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4))))))), file('out/t35_yellow_0__3_premises__rerun.E_input', t9_yellow_0)).
fof(c_0_4, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>![X2]:![X3]:((r1_tarski(X2,X3)&(r2_yellow_0(X1,X2)&r2_yellow_0(X1,X3)))=>r1_orders_2(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2)))))), inference(assume_negation,[status(cth)],[t35_yellow_0])).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:((((((r1_lattice3(X5,X6,X7)|X7!=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X6,X8)|r1_orders_2(X5,X8,X7)))|X7!=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))&((((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X6,X7))|X7=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((((r1_lattice3(X5,X6,esk4_3(X5,X6,X7))|~r1_lattice3(X5,X6,X7))|X7=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~r1_orders_2(X5,esk4_3(X5,X6,X7),X7)|~r1_lattice3(X5,X6,X7))|X7=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])])])).
fof(c_0_6, plain, (![X6]:![X7]:![X9]:![X10]:![X7]:![X12]:![X15]:((((m1_subset_1(esk5_2(X6,X7),u1_struct_0(X6))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((r1_lattice3(X6,X7,esk5_2(X6,X7))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((~m1_subset_1(X9,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X9)|r1_orders_2(X6,X9,esk5_2(X6,X7))))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((m1_subset_1(esk6_3(X6,X7,X10),u1_struct_0(X6))|~r1_lattice3(X6,X7,X10))|X10=esk5_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((r1_lattice3(X6,X7,esk6_3(X6,X7,X10))|~r1_lattice3(X6,X7,X10))|X10=esk5_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((~r1_orders_2(X6,esk6_3(X6,X7,X10),X10)|~r1_lattice3(X6,X7,X10))|X10=esk5_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6)))))))&(((((((m1_subset_1(esk8_3(X6,X7,X12),u1_struct_0(X6))|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r1_lattice3(X6,X7,esk8_3(X6,X7,X12))|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X15)|r1_orders_2(X6,X15,esk8_3(X6,X7,X12))))|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk8_3(X6,X7,X12)!=X12|m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))))&(((((((m1_subset_1(esk8_3(X6,X7,X12),u1_struct_0(X6))|r1_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r1_lattice3(X6,X7,esk8_3(X6,X7,X12))|r1_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X15)|r1_orders_2(X6,X15,esk8_3(X6,X7,X12))))|r1_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk8_3(X6,X7,X12)!=X12|r1_lattice3(X6,X7,esk7_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))))&((((((m1_subset_1(esk8_3(X6,X7,X12),u1_struct_0(X6))|~r1_orders_2(X6,esk7_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r1_lattice3(X6,X7,esk8_3(X6,X7,X12))|~r1_orders_2(X6,esk7_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X15)|r1_orders_2(X6,X15,esk8_3(X6,X7,X12))))|~r1_orders_2(X6,esk7_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk8_3(X6,X7,X12)!=X12|~r1_orders_2(X6,esk7_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])])).
fof(c_0_7, negated_conjecture, ((l1_orders_2(esk1_0)&((r1_tarski(esk2_0,esk3_0)&(r2_yellow_0(esk1_0,esk2_0)&r2_yellow_0(esk1_0,esk3_0)))&~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_8,plain,(r1_orders_2(X1,X4,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|X2!=k2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:(((((~r1_lattice3(X5,X7,X8)|r1_lattice3(X5,X6,X8))|~m1_subset_1(X8,u1_struct_0(X5)))|~r1_tarski(X6,X7))|~l1_orders_2(X5))&((((~r2_lattice3(X5,X7,X8)|r2_lattice3(X5,X6,X8))|~m1_subset_1(X8,u1_struct_0(X5)))|~r1_tarski(X6,X7))|~l1_orders_2(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])])])).
cnf(c_0_10,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|X2!=k2_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,plain,(X2=k2_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,esk4_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,plain,(r1_orders_2(X1,X3,esk5_2(X1,X2))|~l1_orders_2(X1)|~r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,plain,(X2=k2_yellow_0(X1,X3)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,plain,(m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16,plain,(r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_17,negated_conjecture,(r2_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_18,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19,plain,(r1_lattice3(X1,X2,X4)|~l1_orders_2(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_lattice3(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20,plain,(r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_21,plain,(esk5_2(X1,X2)=k2_yellow_0(X1,X3)|~r1_lattice3(X1,X2,esk4_3(X1,X3,esk5_2(X1,X2)))|~r1_lattice3(X1,X3,esk5_2(X1,X2))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])).
cnf(c_0_22,plain,(X2=k2_yellow_0(X1,X3)|r1_lattice3(X1,X3,esk4_3(X1,X3,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_23,plain,(r1_lattice3(X1,X2,esk5_2(X1,X2))|~l1_orders_2(X1)|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24,negated_conjecture,(~r1_lattice3(esk1_0,esk2_0,k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_25,plain,(r1_lattice3(X1,X2,k2_yellow_0(X1,X3))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_tarski(X2,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26,negated_conjecture,(r2_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27,negated_conjecture,(r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_28,plain,(esk5_2(X1,X2)=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14]), c_0_23])).
cnf(c_0_29,negated_conjecture,(~m1_subset_1(k2_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_18])])).
cnf(c_0_30,plain,(m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_14, c_0_28])).
cnf(c_0_31,negated_conjecture,(~m1_subset_1(k2_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_26]), c_0_18])])).
cnf(c_0_32,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_30]), c_0_17]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 24
# Proof object formula steps           : 9
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 98
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 84
# ...of the previous two non-trivial   : 81
# Contextual simplify-reflections      : 23
# Paramodulations                      : 82
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 64
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 58
# Current number of unprocessed clauses: 41
# ...number of literals in the above   : 335
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 2009
# Rec. Clause-clause subsumption calls : 225
# Non-unit clause-clause subsumptions  : 25
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6084

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45304 pages
