# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t92_zfmisc_1, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2)))), file('out/t35_subset_1__20_premises__rerun.E_input', t92_zfmisc_1)).
fof(t99_zfmisc_1, axiom, (![X1]:k3_tarski(k1_zfmisc_1(X1))=X1), file('out/t35_subset_1__20_premises__rerun.E_input', t99_zfmisc_1)).
fof(t4_xboole_1, axiom, (![X1]:![X2]:![X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t35_subset_1__20_premises__rerun.E_input', t4_xboole_1)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t35_subset_1__20_premises__rerun.E_input', t12_xboole_1)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t35_subset_1__20_premises__rerun.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, (![X1]:~(v1_xboole_0(k1_zfmisc_1(X1)))), file('out/t35_subset_1__20_premises__rerun.E_input', fc1_subset_1)).
fof(t35_subset_1, conjecture, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X3))=>r1_tarski(X3,k3_subset_1(X1,X2)))))), file('out/t35_subset_1__20_premises__rerun.E_input', t35_subset_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t35_subset_1__20_premises__rerun.E_input', t7_xboole_1)).
fof(t104_xboole_1, axiom, (![X1]:![X2]:(~((~(r2_xboole_0(X1,X2))&(~(X1=X2)&~(r2_xboole_0(X2,X1)))))<=>r3_xboole_0(X1,X2))), file('out/t35_subset_1__20_premises__rerun.E_input', t104_xboole_1)).
fof(d9_xboole_0, axiom, (![X1]:![X2]:(r3_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)|r1_tarski(X2,X1)))), file('out/t35_subset_1__20_premises__rerun.E_input', d9_xboole_0)).
fof(t14_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&(r1_tarski(X3,X2)&![X4]:((r1_tarski(X1,X4)&r1_tarski(X3,X4))=>r1_tarski(X2,X4))))=>X2=k2_xboole_0(X1,X3))), file('out/t35_subset_1__20_premises__rerun.E_input', t14_xboole_1)).
fof(t107_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k5_xboole_0(X2,X3))<=>(r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,k3_xboole_0(X2,X3))))), file('out/t35_subset_1__20_premises__rerun.E_input', t107_xboole_1)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t35_subset_1__20_premises__rerun.E_input', t95_xboole_1)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t35_subset_1__20_premises__rerun.E_input', t28_xboole_1)).
fof(d5_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2))), file('out/t35_subset_1__20_premises__rerun.E_input', d5_subset_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t35_subset_1__20_premises__rerun.E_input', t100_xboole_1)).
fof(t63_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3))), file('out/t35_subset_1__20_premises__rerun.E_input', t63_xboole_1)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t35_subset_1__20_premises__rerun.E_input', t91_xboole_1)).
fof(commutativity_k5_xboole_0, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), file('out/t35_subset_1__20_premises__rerun.E_input', commutativity_k5_xboole_0)).
fof(t76_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_xboole_0(X1,X2)=>r1_xboole_0(k3_xboole_0(X3,X1),k3_xboole_0(X3,X2)))), file('out/t35_subset_1__20_premises__rerun.E_input', t76_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1))), file('out/t35_subset_1__20_premises__rerun.E_input', symmetry_r1_xboole_0)).
fof(c_0_21, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|r1_tarski(X3,k3_tarski(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t92_zfmisc_1])])).
fof(c_0_22, plain, (![X2]:k3_tarski(k1_zfmisc_1(X2))=X2), inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:k2_xboole_0(k2_xboole_0(X4,X5),X6)=k2_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t4_xboole_1])).
fof(c_0_24, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_25,plain,(r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26,plain,(k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
fof(c_0_28, plain, (![X2]:~v1_xboole_0(k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_29, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X3))=>r1_tarski(X3,k3_subset_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t35_subset_1])).
cnf(c_0_30,plain,(k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33,plain,(v1_xboole_0(X1)|r2_hidden(X2,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34,plain,(~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_35, negated_conjecture, ((m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&(r1_tarski(esk2_0,k3_subset_1(esk1_0,esk3_0))&~r1_tarski(esk3_0,k3_subset_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])])).
fof(c_0_36, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_37,plain,(k2_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(X2,X3)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_39,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41,plain,(k2_xboole_0(X1,k2_xboole_0(X2,X3))=X3|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_42,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_37, c_0_31])).
cnf(c_0_43,negated_conjecture,(r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_44, plain, (![X3]:![X4]:![X3]:![X4]:(((~r2_xboole_0(X3,X4)|r3_xboole_0(X3,X4))&((X3!=X4|r3_xboole_0(X3,X4))&(~r2_xboole_0(X4,X3)|r3_xboole_0(X3,X4))))&(~r3_xboole_0(X3,X4)|(r2_xboole_0(X3,X4)|(X3=X4|r2_xboole_0(X4,X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_xboole_1])])])])])])).
cnf(c_0_45,plain,(r1_tarski(X1,X2)|~r1_tarski(k2_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46,negated_conjecture,(k2_xboole_0(esk3_0,X1)=X1|~r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_47, plain, (![X3]:![X4]:![X3]:![X4]:((~r3_xboole_0(X3,X4)|(r1_tarski(X3,X4)|r1_tarski(X4,X3)))&((~r1_tarski(X3,X4)|r3_xboole_0(X3,X4))&(~r1_tarski(X4,X3)|r3_xboole_0(X3,X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_xboole_0])])])])])).
cnf(c_0_48,plain,(r3_xboole_0(X1,X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_50,negated_conjecture,(r1_tarski(esk3_0,X1)|~r1_tarski(esk1_0,X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51,plain,(r1_tarski(X1,X2)|r1_tarski(X2,X1)|~r3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_52,plain,(r3_xboole_0(X1,X1)), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_53,negated_conjecture,(r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_38, c_0_49])).
cnf(c_0_54,negated_conjecture,(r1_tarski(esk3_0,X1)|~r1_tarski(k2_xboole_0(esk1_0,X2),X1)), inference(spm,[status(thm)],[c_0_50, c_0_40])).
cnf(c_0_55,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56,negated_conjecture,(k2_xboole_0(esk2_0,X1)=X1|~r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_42, c_0_53])).
cnf(c_0_57,negated_conjecture,(r1_tarski(esk3_0,k2_xboole_0(esk1_0,X1))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_58, plain, (![X5]:![X6]:![X7]:(((((r1_tarski(X5,esk4_3(X5,X6,X7))|~r1_tarski(X7,X6))|~r1_tarski(X5,X6))|X6=k2_xboole_0(X5,X7))&(((r1_tarski(X7,esk4_3(X5,X6,X7))|~r1_tarski(X7,X6))|~r1_tarski(X5,X6))|X6=k2_xboole_0(X5,X7)))&(((~r1_tarski(X6,esk4_3(X5,X6,X7))|~r1_tarski(X7,X6))|~r1_tarski(X5,X6))|X6=k2_xboole_0(X5,X7)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_xboole_1])])])])])])).
cnf(c_0_59,negated_conjecture,(r1_tarski(esk2_0,X1)|~r1_tarski(esk1_0,X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_45, c_0_56])).
cnf(c_0_60,negated_conjecture,(k2_xboole_0(esk3_0,X1)=X1|~r1_tarski(k2_xboole_0(esk1_0,X2),X1)), inference(spm,[status(thm)],[c_0_42, c_0_57])).
cnf(c_0_61,plain,(X1=k2_xboole_0(X2,X3)|~r1_tarski(X2,X1)|~r1_tarski(X3,X1)|~r1_tarski(X1,esk4_3(X2,X1,X3))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_62,plain,(X1=k2_xboole_0(X2,X3)|r1_tarski(X2,esk4_3(X2,X1,X3))|~r1_tarski(X2,X1)|~r1_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_63,negated_conjecture,(r1_tarski(esk2_0,X1)|~r1_tarski(k2_xboole_0(esk1_0,X2),X1)), inference(spm,[status(thm)],[c_0_59, c_0_40])).
cnf(c_0_64,negated_conjecture,(k2_xboole_0(esk3_0,k2_xboole_0(esk1_0,X1))=k2_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_60, c_0_55])).
cnf(c_0_65,plain,(k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_55])])).
fof(c_0_66, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(((r1_tarski(X4,k2_xboole_0(X5,X6))|~r1_tarski(X4,k5_xboole_0(X5,X6)))&(r1_xboole_0(X4,k3_xboole_0(X5,X6))|~r1_tarski(X4,k5_xboole_0(X5,X6))))&((~r1_tarski(X4,k2_xboole_0(X5,X6))|~r1_xboole_0(X4,k3_xboole_0(X5,X6)))|r1_tarski(X4,k5_xboole_0(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t107_xboole_1])])])])])).
fof(c_0_67, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
fof(c_0_68, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_69, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(X3))|k3_subset_1(X3,X4)=k4_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_70, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_71,negated_conjecture,(r1_tarski(esk2_0,k2_xboole_0(esk1_0,X1))), inference(spm,[status(thm)],[c_0_63, c_0_55])).
cnf(c_0_72,plain,(r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X1,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_40, c_0_30])).
cnf(c_0_73,negated_conjecture,(k2_xboole_0(esk3_0,esk1_0)=esk1_0|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74,negated_conjecture,(r1_tarski(esk3_0,X1)|~r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_40, c_0_46])).
fof(c_0_75, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_xboole_0(X5,X6))|r1_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_76,plain,(r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_78, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_79,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_80, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k5_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_81,plain,(k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_82,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_83,negated_conjecture,(k2_xboole_0(esk2_0,X1)=X1|~r1_tarski(k2_xboole_0(esk1_0,X2),X1)), inference(spm,[status(thm)],[c_0_42, c_0_71])).
fof(c_0_84, plain, (![X4]:![X5]:![X6]:(~r1_xboole_0(X4,X5)|r1_xboole_0(k3_xboole_0(X6,X4),k3_xboole_0(X6,X5)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t76_xboole_1])])])])).
cnf(c_0_85,plain,(r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3))|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_72, c_0_31])).
cnf(c_0_86,negated_conjecture,(k2_xboole_0(esk3_0,esk1_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_55])])).
cnf(c_0_87,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_88,negated_conjecture,(r1_tarski(esk2_0,k3_subset_1(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_89,plain,(r1_xboole_0(X1,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3)))|~r1_tarski(X1,k5_xboole_0(X2,X3))), inference(rw,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_90,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_91,plain,(k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_79, c_0_77])).
cnf(c_0_92,plain,(k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_93,plain,(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)))=k3_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82]), c_0_77])).
cnf(c_0_94,negated_conjecture,(k2_xboole_0(esk2_0,k2_xboole_0(esk1_0,X1))=k2_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_83, c_0_55])).
cnf(c_0_95,plain,(r1_tarski(X1,k5_xboole_0(X2,X3))|~r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_96,plain,(r1_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_97,plain,(r1_tarski(k2_xboole_0(X1,X2),X1)|~r1_tarski(k2_xboole_0(X2,X3),X1)), inference(spm,[status(thm)],[c_0_72, c_0_65])).
cnf(c_0_98,negated_conjecture,(r1_tarski(k2_xboole_0(esk3_0,X1),esk1_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_99,negated_conjecture,(r1_xboole_0(esk2_0,X1)|~r1_xboole_0(k3_subset_1(esk1_0,esk3_0),X1)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_100,plain,(r1_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k2_xboole_0(X2,X3))))|~r1_tarski(X1,k5_xboole_0(X2,X3))), inference(rw,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_101,plain,(k5_xboole_0(X1,k5_xboole_0(X2,k2_xboole_0(X1,X2)))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_91, c_0_90])).
cnf(c_0_102,plain,(k5_xboole_0(X1,k5_xboole_0(X2,X3))=k5_xboole_0(X3,k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_92, c_0_90])).
cnf(c_0_103,plain,(k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k2_xboole_0(X1,X2))))=k3_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_93, c_0_90])).
cnf(c_0_104,negated_conjecture,(k2_xboole_0(esk2_0,esk1_0)=esk1_0|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_94, c_0_65])).
cnf(c_0_105,plain,(r1_tarski(X1,k5_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X1,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_95, c_0_77])).
cnf(c_0_106,plain,(r1_xboole_0(k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)),k5_xboole_0(k5_xboole_0(X1,X3),k2_xboole_0(X1,X3)))|~r1_xboole_0(X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_77]), c_0_77])).
cnf(c_0_107,negated_conjecture,(r1_tarski(k2_xboole_0(esk1_0,esk3_0),esk1_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_108,negated_conjecture,(r1_xboole_0(esk2_0,k5_xboole_0(X1,k5_xboole_0(X2,k2_xboole_0(X1,X2))))|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_109,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk3_0))=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_86]), c_0_43])]), c_0_102]), c_0_92])).
cnf(c_0_110,plain,(k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,X1)))=k3_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_65]), c_0_38])).
cnf(c_0_111,negated_conjecture,(k2_xboole_0(esk2_0,esk1_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_74]), c_0_55])])).
cnf(c_0_112,plain,(r1_tarski(X1,k5_xboole_0(X2,X3))|~r1_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k2_xboole_0(X2,X3))))|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(rw,[status(thm)],[c_0_105, c_0_90])).
cnf(c_0_113,plain,(r1_xboole_0(k5_xboole_0(X1,k5_xboole_0(X2,k2_xboole_0(X1,X2))),k5_xboole_0(X1,k5_xboole_0(X3,k2_xboole_0(X1,X3))))|~r1_xboole_0(X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_90]), c_0_90])).
cnf(c_0_114,plain,(X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_65])).
cnf(c_0_115,negated_conjecture,(r1_tarski(k2_xboole_0(esk1_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_74]), c_0_55])])).
fof(c_0_116, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|r1_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_117,negated_conjecture,(r1_xboole_0(esk2_0,esk3_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k5_xboole_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_86]), c_0_102]), c_0_92]), c_0_109]), c_0_92])).
cnf(c_0_118,negated_conjecture,(k3_subset_1(esk1_0,esk3_0)=k5_xboole_0(esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_39]), c_0_92]), c_0_109])).
cnf(c_0_119,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_111]), c_0_53])]), c_0_102]), c_0_92])).
cnf(c_0_120,plain,(r1_tarski(k5_xboole_0(X1,k5_xboole_0(X2,k2_xboole_0(X1,X2))),k5_xboole_0(X1,X3))|~r1_xboole_0(X2,X3)|~r1_tarski(k5_xboole_0(X1,k5_xboole_0(X2,k2_xboole_0(X1,X2))),k2_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_121,negated_conjecture,(k2_xboole_0(esk1_0,esk3_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_40])])).
cnf(c_0_122,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_123,negated_conjecture,(r1_xboole_0(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_117, c_0_118]), c_0_55])])).
cnf(c_0_124,negated_conjecture,(~r1_tarski(esk3_0,k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_125,negated_conjecture,(k3_subset_1(esk1_0,esk2_0)=k5_xboole_0(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_49]), c_0_92]), c_0_119])).
cnf(c_0_126,negated_conjecture,(r1_tarski(esk3_0,k5_xboole_0(esk1_0,X1))|~r1_xboole_0(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_92]), c_0_109]), c_0_92]), c_0_109]), c_0_57])])).
cnf(c_0_127,negated_conjecture,(r1_xboole_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_128,negated_conjecture,(~r1_tarski(esk3_0,k5_xboole_0(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_129,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_128]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 130
# Proof object clause steps            : 87
# Proof object formula steps           : 43
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 48
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 34
# Processed clauses                    : 2135
# ...of these trivial                  : 14
# ...subsumed                          : 1449
# ...remaining for further processing  : 672
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 88
# Backward-rewritten                   : 57
# Generated clauses                    : 13394
# ...of the previous two non-trivial   : 12548
# Contextual simplify-reflections      : 618
# Paramodulations                      : 13393
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 492
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 446
# Current number of unprocessed clauses: 9486
# ...number of literals in the above   : 31372
# Current number of archived formulas  : 0
# Current number of archived clauses   : 181
# Clause-clause subsumption calls (NU) : 108576
# Rec. Clause-clause subsumption calls : 89857
# Non-unit clause-clause subsumptions  : 1707
# Unit Clause-clause subsumption calls : 2102
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 51
# BW rewrite match successes           : 45
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 187936

# -------------------------------------------------
# User time                : 0.164 s
# System time              : 0.008 s
# Total time               : 0.172 s
# Maximum resident set size: 45460 pages
