# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t35_yellow19, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m2_yellow_6(X3,X1,X2)&v3_yellow_6(X3,X1)))))), file('out/t35_yellow19__15_premises__rerun.E_input', t35_yellow19)).
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t35_yellow19__15_premises__rerun.E_input', t7_boole)).
fof(t32_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r3_waybel_9(X1,X2,X3)&![X4]:(m2_yellow_6(X4,X1,X2)=>~(r2_hidden(X3,k10_yellow_6(X1,X4)))))))))), file('out/t35_yellow19__15_premises__rerun.E_input', t32_waybel_9)).
fof(d16_yellow_6, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(v3_yellow_6(X2,X1)<=>~(k10_yellow_6(X1,X2)=k1_xboole_0))))), file('out/t35_yellow19__15_premises__rerun.E_input', d16_yellow_6)).
fof(dt_o_0_0_xboole_0, axiom, (v1_xboole_0(o_0_0_xboole_0)), file('out/t35_yellow19__15_premises__rerun.E_input', dt_o_0_0_xboole_0)).
fof(d2_xboole_0, axiom, (k1_xboole_0=o_0_0_xboole_0), file('out/t35_yellow19__15_premises__rerun.E_input', d2_xboole_0)).
fof(dt_m2_yellow_6, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>(~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))))), file('out/t35_yellow19__15_premises__rerun.E_input', dt_m2_yellow_6)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t35_yellow19__15_premises__rerun.E_input', dt_l1_pre_topc)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t35_yellow19__15_premises__rerun.E_input', t4_subset)).
fof(dt_k10_yellow_6, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t35_yellow19__15_premises__rerun.E_input', dt_k10_yellow_6)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t35_yellow19__15_premises__rerun.E_input', t2_subset)).
fof(dt_o_3_4_yellow19, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))&m2_yellow_6(X3,X1,X2)))=>m1_subset_1(o_3_4_yellow19(X1,X2,X3),k10_yellow_6(X1,X3)))), file('out/t35_yellow19__15_premises__rerun.E_input', dt_o_3_4_yellow19)).
fof(t33_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3)))))), file('out/t35_yellow19__15_premises__rerun.E_input', t33_yellow19)).
fof(t31_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_waybel_9(X1,X3,X4)=>r3_waybel_9(X1,X2,X4))))))), file('out/t35_yellow19__15_premises__rerun.E_input', t31_waybel_9)).
fof(t29_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k10_yellow_6(X1,X2))=>r3_waybel_9(X1,X2,X3)))))), file('out/t35_yellow19__15_premises__rerun.E_input', t29_waybel_9)).
fof(t6_boole, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t35_yellow19__15_premises__rerun.E_input', t6_boole)).
fof(c_0_16, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m2_yellow_6(X3,X1,X2)&v3_yellow_6(X3,X1))))))), inference(assume_negation,[status(cth)],[t35_yellow19])).
fof(c_0_17, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_18, plain, (![X5]:![X6]:![X7]:(((((m2_yellow_6(esk4_3(X5,X6,X7),X5,X6)|~r3_waybel_9(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((r2_hidden(X7,k10_yellow_6(X5,esk4_3(X5,X6,X7)))|~r3_waybel_9(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t32_waybel_9])])])])])])])])).
fof(c_0_19, negated_conjecture, (![X6]:![X7]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((((~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0))&((v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))&((v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0))&(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)))))&((~m2_yellow_6(X6,esk1_0,esk2_0)|~v3_yellow_6(X6,esk1_0))|~v1_compts_1(esk1_0)))&(((m2_yellow_6(esk3_1(X7),esk1_0,X7)|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,esk1_0)))))|v1_compts_1(esk1_0))&((v3_yellow_6(esk3_1(X7),esk1_0)|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,esk1_0)))))|v1_compts_1(esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])])])])).
cnf(c_0_20,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,esk4_3(X1,X2,X3)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~r3_waybel_9(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_22, plain, (![X3]:![X4]:((((~v3_yellow_6(X4,X3)|k10_yellow_6(X3,X4)!=k1_xboole_0)|(v2_struct_0(X4)|(~v4_orders_2(X4)|(~v7_waybel_0(X4)|~l1_waybel_0(X4,X3)))))|(v2_struct_0(X3)|(~v2_pre_topc(X3)|~l1_pre_topc(X3))))&(((k10_yellow_6(X3,X4)=k1_xboole_0|v3_yellow_6(X4,X3))|(v2_struct_0(X4)|(~v4_orders_2(X4)|(~v7_waybel_0(X4)|~l1_waybel_0(X4,X3)))))|(v2_struct_0(X3)|(~v2_pre_topc(X3)|~l1_pre_topc(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_yellow_6])])])])])])])).
cnf(c_0_23,plain,(v1_xboole_0(o_0_0_xboole_0)), inference(split_conjunct,[status(thm)],[dt_o_0_0_xboole_0])).
cnf(c_0_24,plain,(k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
cnf(c_0_25,negated_conjecture,(~v1_compts_1(esk1_0)|~v3_yellow_6(X1,esk1_0)|~m2_yellow_6(X1,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m2_yellow_6(esk4_3(X1,X2,X3),X1,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~r3_waybel_9(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,negated_conjecture,(v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31,negated_conjecture,(v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33,negated_conjecture,(~v1_compts_1(esk1_0)|~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~v1_xboole_0(k10_yellow_6(X2,esk4_3(X2,X1,X3)))|~l1_waybel_0(X1,X2)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_35,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v3_yellow_6(X2,X1)|k10_yellow_6(X1,X2)=k1_xboole_0|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_36,plain,(v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_37, plain, (![X4]:![X5]:![X6]:(((~v2_struct_0(X6)|~m2_yellow_6(X6,X4,X5))|((v2_struct_0(X4)|~l1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4))))))&(((v4_orders_2(X6)|~m2_yellow_6(X6,X4,X5))|((v2_struct_0(X4)|~l1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4))))))&(((v7_waybel_0(X6)|~m2_yellow_6(X6,X4,X5))|((v2_struct_0(X4)|~l1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4))))))&((l1_waybel_0(X6,X4)|~m2_yellow_6(X6,X4,X5))|((v2_struct_0(X4)|~l1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_yellow_6])])])])])])])).
fof(c_0_38, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_39,negated_conjecture,(~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v3_yellow_6(esk4_3(esk1_0,esk2_0,X1),esk1_0)|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_40,plain,(v3_yellow_6(esk4_3(X1,X2,X3),X1)|v2_struct_0(esk4_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(esk4_3(X1,X2,X3),X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(esk4_3(X1,X2,X3))|~v7_waybel_0(X2)|~v4_orders_2(esk4_3(X1,X2,X3))|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_41,plain,(v2_struct_0(X1)|v2_struct_0(X2)|l1_waybel_0(X3,X2)|~l1_waybel_0(X1,X2)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_struct_0(X2)|~m2_yellow_6(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_42,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X1,X2)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_struct_0(X2)|~m2_yellow_6(X3,X2,X1)|~v2_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44,negated_conjecture,(v2_struct_0(esk4_3(esk1_0,esk2_0,X1))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_waybel_0(esk4_3(esk1_0,esk2_0,X1),esk1_0)|~v7_waybel_0(esk4_3(esk1_0,esk2_0,X1))|~v4_orders_2(esk4_3(esk1_0,esk2_0,X1))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_45,plain,(l1_waybel_0(esk4_3(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_26]), c_0_42])).
cnf(c_0_46,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(esk4_3(X1,X2,X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_26]), c_0_42])).
cnf(c_0_47,negated_conjecture,(v2_struct_0(esk4_3(esk1_0,esk2_0,X1))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v7_waybel_0(esk4_3(esk1_0,esk2_0,X1))|~v4_orders_2(esk4_3(esk1_0,esk2_0,X1))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_48,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v7_waybel_0(X3)|~l1_waybel_0(X1,X2)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_struct_0(X2)|~m2_yellow_6(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_49, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_50, plain, (![X3]:![X4]:(((v2_struct_0(X3)|(~v2_pre_topc(X3)|~l1_pre_topc(X3)))|(v2_struct_0(X4)|(~v4_orders_2(X4)|(~v7_waybel_0(X4)|~l1_waybel_0(X4,X3)))))|m1_subset_1(k10_yellow_6(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_yellow_6])])])).
fof(c_0_51, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_52, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))|((v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4))))|~m2_yellow_6(X6,X4,X5)))|m1_subset_1(o_3_4_yellow19(X4,X5,X6),k10_yellow_6(X4,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_o_3_4_yellow19])])])).
cnf(c_0_53,negated_conjecture,(~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v7_waybel_0(esk4_3(esk1_0,esk2_0,X1))|~v4_orders_2(esk4_3(esk1_0,esk2_0,X1))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_54,plain,(v7_waybel_0(esk4_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_26]), c_0_42])).
cnf(c_0_55,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v4_orders_2(X3)|~l1_waybel_0(X1,X2)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_struct_0(X2)|~m2_yellow_6(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_56,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57,plain,(m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59,plain,(m1_subset_1(o_3_4_yellow19(X1,X2,X3),k10_yellow_6(X1,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m2_yellow_6(X3,X1,X2)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60,negated_conjecture,(~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v4_orders_2(esk4_3(esk1_0,esk2_0,X1))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_61,plain,(v4_orders_2(esk4_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_26]), c_0_42])).
fof(c_0_62, plain, (![X4]:![X5]:![X8]:(((((m1_subset_1(esk5_2(X4,X5),u1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|~v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((r3_waybel_9(X4,X5,esk5_2(X4,X5))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|~v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))&((((~v2_struct_0(esk6_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((v4_orders_2(esk6_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((v7_waybel_0(esk6_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((l1_waybel_0(esk6_1(X4),X4)|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))))&(((~m1_subset_1(X8,u1_struct_0(X4))|~r3_waybel_9(X4,esk6_1(X4),X8))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t33_yellow19])])])])])])])])).
fof(c_0_63, plain, (![X5]:![X6]:![X7]:![X8]:((v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))|((v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5))))|(~m2_yellow_6(X7,X5,X6)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r3_waybel_9(X5,X7,X8)|r3_waybel_9(X5,X6,X8))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_waybel_9])])])])])])).
fof(c_0_64, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))|((v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4))))|(~m1_subset_1(X6,u1_struct_0(X4))|(~r2_hidden(X6,k10_yellow_6(X4,X5))|r3_waybel_9(X4,X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_waybel_9])])])])])])).
cnf(c_0_65,plain,(m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,k10_yellow_6(X2,X3))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_66,plain,(v1_xboole_0(k10_yellow_6(X1,X2))|r2_hidden(o_3_4_yellow19(X1,X3,X2),k10_yellow_6(X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|~m2_yellow_6(X2,X1,X3)|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_67,negated_conjecture,(~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_68,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,esk5_2(X1,X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69,plain,(r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_yellow_6(X4,X1,X2)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_70,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(X1)|m2_yellow_6(esk3_1(X1),esk1_0,X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_71,plain,(r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(X3,k10_yellow_6(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72,plain,(m1_subset_1(o_3_4_yellow19(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(k10_yellow_6(X1,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~m2_yellow_6(X3,X1,X2)|~l1_waybel_0(X3,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X3)|~v7_waybel_0(X2)|~v4_orders_2(X3)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73,negated_conjecture,(~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_74,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_75,negated_conjecture,(r3_waybel_9(esk1_0,X1,X2)|v1_compts_1(esk1_0)|v2_struct_0(X1)|~r3_waybel_9(esk1_0,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_27]), c_0_28])]), c_0_29])).
cnf(c_0_76,plain,(r3_waybel_9(X1,X2,o_3_4_yellow19(X1,X3,X2))|v1_xboole_0(k10_yellow_6(X1,X2))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m2_yellow_6(X2,X1,X3)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X3,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X3)|~v4_orders_2(X2)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_66]), c_0_72])).
cnf(c_0_77,negated_conjecture,(~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_78,plain,(v2_struct_0(X1)|v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_waybel_9(X1,esk6_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_79,negated_conjecture,(r3_waybel_9(esk1_0,X1,o_3_4_yellow19(esk1_0,X2,esk3_1(X1)))|v1_xboole_0(k10_yellow_6(esk1_0,esk3_1(X1)))|v2_struct_0(esk3_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(o_3_4_yellow19(esk1_0,X2,esk3_1(X1)),u1_struct_0(esk1_0))|~m2_yellow_6(esk3_1(X1),esk1_0,X2)|~l1_waybel_0(esk3_1(X1),esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(esk3_1(X1))|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(esk3_1(X1))|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_27]), c_0_28])]), c_0_77]), c_0_29])).
cnf(c_0_80,negated_conjecture,(v1_xboole_0(k10_yellow_6(esk1_0,esk3_1(esk6_1(esk1_0))))|v2_struct_0(esk3_1(esk6_1(esk1_0)))|v2_struct_0(esk6_1(esk1_0))|v2_struct_0(X1)|~m1_subset_1(o_3_4_yellow19(esk1_0,X1,esk3_1(esk6_1(esk1_0))),u1_struct_0(esk1_0))|~m2_yellow_6(esk3_1(esk6_1(esk1_0)),esk1_0,X1)|~l1_waybel_0(esk3_1(esk6_1(esk1_0)),esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v7_waybel_0(X1)|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_27]), c_0_28])]), c_0_77]), c_0_29])).
fof(c_0_81, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_82,negated_conjecture,(v1_xboole_0(k10_yellow_6(esk1_0,esk3_1(esk6_1(esk1_0))))|v2_struct_0(esk3_1(esk6_1(esk1_0)))|v2_struct_0(esk6_1(esk1_0))|v2_struct_0(X1)|~m2_yellow_6(esk3_1(esk6_1(esk1_0)),esk1_0,X1)|~l1_waybel_0(esk3_1(esk6_1(esk1_0)),esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v7_waybel_0(X1)|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_72]), c_0_27]), c_0_28])]), c_0_29])).
cnf(c_0_83,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_84,negated_conjecture,(v1_xboole_0(k10_yellow_6(esk1_0,esk3_1(esk6_1(esk1_0))))|v2_struct_0(esk3_1(esk6_1(esk1_0)))|v2_struct_0(esk6_1(esk1_0))|~l1_waybel_0(esk3_1(esk6_1(esk1_0)),esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_70]), c_0_77])).
cnf(c_0_85,plain,(v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|k10_yellow_6(X1,X2)!=k1_xboole_0|~v3_yellow_6(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_86,negated_conjecture,(k10_yellow_6(esk1_0,esk3_1(esk6_1(esk1_0)))=k1_xboole_0|v2_struct_0(esk3_1(esk6_1(esk1_0)))|v2_struct_0(esk6_1(esk1_0))|~l1_waybel_0(esk3_1(esk6_1(esk1_0)),esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_87,negated_conjecture,(v2_struct_0(esk3_1(esk6_1(esk1_0)))|v2_struct_0(esk6_1(esk1_0))|~v3_yellow_6(esk3_1(esk6_1(esk1_0)),esk1_0)|~l1_waybel_0(esk3_1(esk6_1(esk1_0)),esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_27]), c_0_28])]), c_0_29])).
cnf(c_0_88,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(X1)|v3_yellow_6(esk3_1(X1),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_89,negated_conjecture,(v2_struct_0(esk3_1(esk6_1(esk1_0)))|v2_struct_0(esk6_1(esk1_0))|~l1_waybel_0(esk3_1(esk6_1(esk1_0)),esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_77])).
cnf(c_0_90,negated_conjecture,(l1_waybel_0(esk3_1(X1),esk1_0)|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_70]), c_0_29])).
cnf(c_0_91,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v2_struct_0(esk3_1(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_70]), c_0_29])).
cnf(c_0_92,negated_conjecture,(v2_struct_0(esk3_1(esk6_1(esk1_0)))|v2_struct_0(esk6_1(esk1_0))|~l1_struct_0(esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_77])).
cnf(c_0_93,negated_conjecture,(v2_struct_0(esk6_1(esk1_0))|~l1_struct_0(esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk3_1(esk6_1(esk1_0)))|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_77])).
cnf(c_0_94,negated_conjecture,(v7_waybel_0(esk3_1(X1))|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_70]), c_0_29])).
cnf(c_0_95,negated_conjecture,(v2_struct_0(esk6_1(esk1_0))|~l1_struct_0(esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk3_1(esk6_1(esk1_0)))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_77])).
cnf(c_0_96,negated_conjecture,(v4_orders_2(esk3_1(X1))|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_70]), c_0_29])).
cnf(c_0_97,negated_conjecture,(v2_struct_0(esk6_1(esk1_0))|~l1_struct_0(esk1_0)|~l1_waybel_0(esk6_1(esk1_0),esk1_0)|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_77])).
cnf(c_0_98,plain,(v2_struct_0(X1)|v1_compts_1(X1)|l1_waybel_0(esk6_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_99,plain,(v2_struct_0(X1)|v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_100,negated_conjecture,(v2_struct_0(esk6_1(esk1_0))|~l1_struct_0(esk1_0)|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_27]), c_0_28])]), c_0_77]), c_0_29])).
cnf(c_0_101,negated_conjecture,(~l1_struct_0(esk1_0)|~v7_waybel_0(esk6_1(esk1_0))|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_27]), c_0_28])]), c_0_77]), c_0_29])).
cnf(c_0_102,plain,(v2_struct_0(X1)|v1_compts_1(X1)|v7_waybel_0(esk6_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_103,negated_conjecture,(~l1_struct_0(esk1_0)|~v4_orders_2(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_27]), c_0_28])]), c_0_77]), c_0_29])).
cnf(c_0_104,plain,(v2_struct_0(X1)|v1_compts_1(X1)|v4_orders_2(esk6_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_105,negated_conjecture,(~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_27]), c_0_28])]), c_0_77]), c_0_29])).
cnf(c_0_106,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_42]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 107
# Proof object clause steps            : 76
# Proof object formula steps           : 31
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 126
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 99
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 92
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 0
# Generated clauses                    : 94
# ...of the previous two non-trivial   : 76
# Contextual simplify-reflections      : 93
# Paramodulations                      : 94
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 76
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 69
# Current number of unprocessed clauses: 8
# ...number of literals in the above   : 140
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 6050
# Rec. Clause-clause subsumption calls : 216
# Non-unit clause-clause subsumptions  : 113
# Unit Clause-clause subsumption calls : 25
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8418

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 45296 pages
