# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t34_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2)))))))), file('out/t35_lattice3__9_premises__rerun.E_input', t34_lattice3)).
fof(dt_k16_lattice3, axiom, (![X1]:![X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1)))), file('out/t35_lattice3__9_premises__rerun.E_input', dt_k16_lattice3)).
fof(d16_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4))))))), file('out/t35_lattice3__9_premises__rerun.E_input', d16_lattice3)).
fof(t1_filter_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)=>r3_lattices(X1,k3_lattices(X1,X4,X2),k3_lattices(X1,X4,X3)))))))), file('out/t35_lattice3__9_premises__rerun.E_input', t1_filter_0)).
fof(fraenkel_a_3_4_lattice3, axiom, (![X1]:![X2]:![X3]:![X4]:(((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))&m1_subset_1(X3,u1_struct_0(X2)))=>(r2_hidden(X1,a_3_4_lattice3(X2,X3,X4))<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(X1=k3_lattices(X2,X3,X5)&r2_hidden(X5,X4)))))), file('out/t35_lattice3__9_premises__rerun.E_input', fraenkel_a_3_4_lattice3)).
fof(cc1_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1)))))))))), file('out/t35_lattice3__9_premises__rerun.E_input', cc1_lattices)).
fof(redefinition_r3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3)))), file('out/t35_lattice3__9_premises__rerun.E_input', redefinition_r3_lattices)).
fof(t35_lattice3, conjecture, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:r3_lattices(X1,k3_lattices(X1,X2,k16_lattice3(X1,X3)),k16_lattice3(X1,a_3_4_lattice3(X1,X2,X3)))))), file('out/t35_lattice3__9_premises__rerun.E_input', t35_lattice3)).
fof(dt_k3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1)))), file('out/t35_lattice3__9_premises__rerun.E_input', dt_k3_lattices)).
fof(dt_l3_lattices, axiom, (![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1)))), file('out/t35_lattice3__9_premises__rerun.E_input', dt_l3_lattices)).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r3_lattice3(X5,X6,X7)|X6!=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r3_lattice3(X5,X8,X7)|r3_lattices(X5,X8,X6)))|X6!=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))))&(((((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&(((((r3_lattice3(X5,esk1_3(X5,X6,X7),X7)|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&((((~r3_lattices(X5,esk1_3(X5,X6,X7),X6)|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])])).
fof(c_0_11, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l3_lattices(X3))|m1_subset_1(k16_lattice3(X3,X4),u1_struct_0(X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k16_lattice3])])])])])).
fof(c_0_12, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:((((~r3_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|r1_lattices(X5,X6,X8))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((m1_subset_1(esk2_3(X5,X6,X7),u1_struct_0(X5))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((r2_hidden(esk2_3(X5,X6,X7),X7)|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((~r1_lattices(X5,X6,esk2_3(X5,X6,X7))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])])).
cnf(c_0_13,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X2!=k16_lattice3(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:![X8]:((v2_struct_0(X5)|(~v4_lattices(X5)|(~v5_lattices(X5)|(~v6_lattices(X5)|(~v8_lattices(X5)|(~v9_lattices(X5)|~l3_lattices(X5)))))))|(~m1_subset_1(X6,u1_struct_0(X5))|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|(~r3_lattices(X5,X6,X7)|r3_lattices(X5,k3_lattices(X5,X8,X6),k3_lattices(X5,X8,X7)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_filter_0])])])])])])).
fof(c_0_16, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:((((m1_subset_1(esk3_4(X6,X7,X8,X9),u1_struct_0(X7))|~r2_hidden(X6,a_3_4_lattice3(X7,X8,X9)))|((v2_struct_0(X7)|(~v10_lattices(X7)|(~v4_lattice3(X7)|~l3_lattices(X7))))|~m1_subset_1(X8,u1_struct_0(X7))))&(((X6=k3_lattices(X7,X8,esk3_4(X6,X7,X8,X9))|~r2_hidden(X6,a_3_4_lattice3(X7,X8,X9)))|((v2_struct_0(X7)|(~v10_lattices(X7)|(~v4_lattice3(X7)|~l3_lattices(X7))))|~m1_subset_1(X8,u1_struct_0(X7))))&((r2_hidden(esk3_4(X6,X7,X8,X9),X9)|~r2_hidden(X6,a_3_4_lattice3(X7,X8,X9)))|((v2_struct_0(X7)|(~v10_lattices(X7)|(~v4_lattice3(X7)|~l3_lattices(X7))))|~m1_subset_1(X8,u1_struct_0(X7))))))&(((~m1_subset_1(X11,u1_struct_0(X7))|(X6!=k3_lattices(X7,X8,X11)|~r2_hidden(X11,X9)))|r2_hidden(X6,a_3_4_lattice3(X7,X8,X9)))|((v2_struct_0(X7)|(~v10_lattices(X7)|(~v4_lattice3(X7)|~l3_lattices(X7))))|~m1_subset_1(X8,u1_struct_0(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_3_4_lattice3])])])])])])])])).
fof(c_0_17, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v4_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v5_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v6_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v7_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v8_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&((v9_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:(((~r3_lattices(X4,X5,X6)|r1_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_lattices(X4,X5,X6)|r3_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_19,plain,(v2_struct_0(X1)|r1_lattices(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r3_lattice3(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(r3_lattice3(X1,k16_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_14])).
cnf(c_0_21,plain,(r3_lattices(X1,k3_lattices(X1,X2,X3),k3_lattices(X1,X2,X4))|v2_struct_0(X1)|~r3_lattices(X1,X3,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~v6_lattices(X1)|~v5_lattices(X1)|~v4_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(v2_struct_0(X2)|X3=k3_lattices(X2,X1,esk3_4(X3,X2,X1,X4))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)|~r2_hidden(X3,a_3_4_lattice3(X2,X1,X4))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(v2_struct_0(X2)|m1_subset_1(esk3_4(X3,X2,X1,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)|~r2_hidden(X3,a_3_4_lattice3(X2,X1,X4))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,plain,(v2_struct_0(X1)|v4_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(v2_struct_0(X1)|v5_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26,plain,(v2_struct_0(X1)|v6_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,plain,(v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28,plain,(v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29,plain,(v2_struct_0(X2)|r3_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r1_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30,plain,(r1_lattices(X1,k16_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_14])).
cnf(c_0_31,plain,(r3_lattices(X1,k3_lattices(X1,X2,X3),X4)|v2_struct_0(X1)|~r2_hidden(X4,a_3_4_lattice3(X1,X2,X5))|~r3_lattices(X1,X3,esk3_4(X4,X1,X2,X5))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_32,plain,(r3_lattices(X1,k16_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_14]), c_0_26]), c_0_27]), c_0_28])).
fof(c_0_33, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:r3_lattices(X1,k3_lattices(X1,X2,k16_lattice3(X1,X3)),k16_lattice3(X1,a_3_4_lattice3(X1,X2,X3))))))), inference(assume_negation,[status(cth)],[t35_lattice3])).
cnf(c_0_34,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,X2,esk2_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_35,plain,(v2_struct_0(X2)|r1_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r3_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_36,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37,plain,(r3_lattices(X1,k3_lattices(X1,X2,k16_lattice3(X1,X3)),X4)|v2_struct_0(X1)|~r2_hidden(esk3_4(X4,X1,X2,X5),X3)|~r2_hidden(X4,a_3_4_lattice3(X1,X2,X5))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_23]), c_0_14])).
cnf(c_0_38,plain,(v2_struct_0(X2)|r2_hidden(esk3_4(X3,X2,X1,X4),X4)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)|~r2_hidden(X3,a_3_4_lattice3(X2,X1,X4))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_39, negated_conjecture, (((~v2_struct_0(esk4_0)&(v10_lattices(esk4_0)&(v4_lattice3(esk4_0)&l3_lattices(esk4_0))))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&~r3_lattices(esk4_0,k3_lattices(esk4_0,esk5_0,k16_lattice3(esk4_0,esk6_0)),k16_lattice3(esk4_0,a_3_4_lattice3(esk4_0,esk5_0,esk6_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_33])])])])])])).
cnf(c_0_40,plain,(v2_struct_0(X1)|r3_lattices(X1,X4,X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X2!=k16_lattice3(X1,X3)|~r3_lattice3(X1,X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_41,plain,(r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~v6_lattices(X1)|~r3_lattices(X1,X2,esk2_3(X1,X2,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_42,plain,(r3_lattices(X1,k3_lattices(X1,X2,k16_lattice3(X1,X3)),X4)|v2_struct_0(X1)|~r2_hidden(X4,a_3_4_lattice3(X1,X2,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43,negated_conjecture,(~r3_lattices(esk4_0,k3_lattices(esk4_0,esk5_0,k16_lattice3(esk4_0,esk6_0)),k16_lattice3(esk4_0,a_3_4_lattice3(esk4_0,esk5_0,esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44,plain,(r3_lattices(X1,X2,k16_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_40]), c_0_14])).
cnf(c_0_45,negated_conjecture,(l3_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46,negated_conjecture,(v4_lattice3(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47,negated_conjecture,(v10_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49,plain,(r3_lattice3(X1,k3_lattices(X1,X2,k16_lattice3(X1,X3)),X4)|v2_struct_0(X1)|~r2_hidden(esk2_3(X1,k3_lattices(X1,X2,k16_lattice3(X1,X3)),X4),a_3_4_lattice3(X1,X2,X3))|~m1_subset_1(k3_lattices(X1,X2,k16_lattice3(X1,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_50,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|r2_hidden(esk2_3(X1,X2,X3),X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51,negated_conjecture,(~r3_lattice3(esk4_0,k3_lattices(esk4_0,esk5_0,k16_lattice3(esk4_0,esk6_0)),a_3_4_lattice3(esk4_0,esk5_0,esk6_0))|~m1_subset_1(k3_lattices(esk4_0,esk5_0,k16_lattice3(esk4_0,esk6_0)),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_52,plain,(r3_lattice3(X1,k3_lattices(X1,X2,k16_lattice3(X1,X3)),a_3_4_lattice3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(k3_lattices(X1,X2,k16_lattice3(X1,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_53,negated_conjecture,(m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_54, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v4_lattices(X4)|~l2_lattices(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k3_lattices(X4,X5,X6),u1_struct_0(X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_lattices])])])).
cnf(c_0_55,negated_conjecture,(~m1_subset_1(k3_lattices(esk4_0,esk5_0,k16_lattice3(esk4_0,esk6_0)),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_56,plain,(m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)|~v4_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_57, plain, (![X2]:((l1_lattices(X2)|~l3_lattices(X2))&(l2_lattices(X2)|~l3_lattices(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_58,negated_conjecture,(~l2_lattices(esk4_0)|~v4_lattices(esk4_0)|~m1_subset_1(k16_lattice3(esk4_0,esk6_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_53])]), c_0_48])).
cnf(c_0_59,plain,(l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_60,negated_conjecture,(~v4_lattices(esk4_0)|~m1_subset_1(k16_lattice3(esk4_0,esk6_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_45])])).
cnf(c_0_61,negated_conjecture,(~m1_subset_1(k16_lattice3(esk4_0,esk6_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_24]), c_0_45]), c_0_47])]), c_0_48])).
cnf(c_0_62,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_14]), c_0_45])]), c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 42
# Proof object formula steps           : 21
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 110
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 106
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 63
# ...of the previous two non-trivial   : 60
# Contextual simplify-reflections      : 50
# Paramodulations                      : 60
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 70
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 188
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 6451
# Rec. Clause-clause subsumption calls : 327
# Non-unit clause-clause subsumptions  : 55
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7059

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.004 s
# Total time               : 0.024 s
# Maximum resident set size: 45352 pages
