# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_yellow19, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>~((r2_hidden(X2,k6_yellow_6(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~(r3_waybel_9(X1,X2,X3))))))))), file('out/t36_yellow19__2_premises__rerun.E_input', t36_yellow19)).
fof(l38_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>~((r2_hidden(X2,k6_yellow_6(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~(r3_waybel_9(X1,X2,X3))))))=>v1_compts_1(X1)))), file('out/t36_yellow19__2_premises__rerun.E_input', l38_yellow19)).
fof(l37_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3)))))), file('out/t36_yellow19__2_premises__rerun.E_input', l37_yellow19)).
fof(c_0_3, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>~((r2_hidden(X2,k6_yellow_6(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~(r3_waybel_9(X1,X2,X3)))))))))), inference(assume_negation,[status(cth)],[t36_yellow19])).
fof(c_0_4, negated_conjecture, (![X6]:![X7]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((((~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0))&((v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))&((v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0))&(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)))))&((r2_hidden(esk2_0,k6_yellow_6(esk1_0))|~v1_compts_1(esk1_0))&((~m1_subset_1(X6,u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X6))|~v1_compts_1(esk1_0))))&((((m1_subset_1(esk3_1(X7),u1_struct_0(esk1_0))|~r2_hidden(X7,k6_yellow_6(esk1_0)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,esk1_0)))))|v1_compts_1(esk1_0))&(((r3_waybel_9(esk1_0,X7,esk3_1(X7))|~r2_hidden(X7,k6_yellow_6(esk1_0)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,esk1_0)))))|v1_compts_1(esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])])])).
fof(c_0_5, plain, (![X4]:![X6]:((((~v2_struct_0(esk4_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((v4_orders_2(esk4_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((v7_waybel_0(esk4_1(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((l1_waybel_0(esk4_1(X4),X4)|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))))&(((r2_hidden(esk4_1(X4),k6_yellow_6(X4))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((~m1_subset_1(X6,u1_struct_0(X4))|~r3_waybel_9(X4,esk4_1(X4),X6))|v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l38_yellow19])])])])])])])])).
cnf(c_0_6,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(X1)|r3_waybel_9(esk1_0,X1,esk3_1(X1))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~r2_hidden(X1,k6_yellow_6(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v1_compts_1(X1)|l1_waybel_0(esk4_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(esk1_0))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~r2_hidden(X1,k6_yellow_6(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(r3_waybel_9(esk1_0,esk4_1(esk1_0),esk3_1(esk4_1(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))|~r2_hidden(esk4_1(esk1_0),k6_yellow_6(esk1_0))|~v7_waybel_0(esk4_1(esk1_0))|~v4_orders_2(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_13,plain,(v2_struct_0(X1)|v1_compts_1(X1)|r2_hidden(esk4_1(X1),k6_yellow_6(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,negated_conjecture,(m1_subset_1(esk3_1(esk4_1(esk1_0)),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))|~r2_hidden(esk4_1(esk1_0),k6_yellow_6(esk1_0))|~v7_waybel_0(esk4_1(esk1_0))|~v4_orders_2(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_7]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_15,negated_conjecture,(r3_waybel_9(esk1_0,esk4_1(esk1_0),esk3_1(esk4_1(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))|~v7_waybel_0(esk4_1(esk1_0))|~v4_orders_2(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_16,plain,(v2_struct_0(X1)|v1_compts_1(X1)|v4_orders_2(esk4_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk3_1(esk4_1(esk1_0)),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))|~v7_waybel_0(esk4_1(esk1_0))|~v4_orders_2(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_13]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_18,negated_conjecture,(r3_waybel_9(esk1_0,esk4_1(esk1_0),esk3_1(esk4_1(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))|~v7_waybel_0(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_19,plain,(v2_struct_0(X1)|v1_compts_1(X1)|v7_waybel_0(esk4_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_20,negated_conjecture,(m1_subset_1(esk3_1(esk4_1(esk1_0)),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))|~v7_waybel_0(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_16]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_21,plain,(v2_struct_0(X1)|v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_waybel_9(X1,esk4_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22,negated_conjecture,(r3_waybel_9(esk1_0,esk4_1(esk1_0),esk3_1(esk4_1(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk3_1(esk4_1(esk1_0)),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_19]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_24,plain,(v2_struct_0(X1)|v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(esk4_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25,negated_conjecture,(v1_compts_1(esk1_0)|v2_struct_0(esk4_1(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_8]), c_0_9])]), c_0_10]), c_0_23])).
fof(c_0_26, plain, (![X4]:![X5]:((((m1_subset_1(esk5_2(X4,X5),u1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|~v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((r3_waybel_9(X4,X5,esk5_2(X4,X5))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|~v1_compts_1(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l37_yellow19])])])])])])])])).
cnf(c_0_27,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_28,negated_conjecture,(v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_29,negated_conjecture,(v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_30,negated_conjecture,(v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_31,negated_conjecture,(~v1_compts_1(esk1_0)|~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_32,negated_conjecture,(~v1_compts_1(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_33,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,esk5_2(X1,X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28])])).
cnf(c_0_35,negated_conjecture,(v7_waybel_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_28])])).
cnf(c_0_36,negated_conjecture,(v4_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_28])])).
cnf(c_0_37,negated_conjecture,(~v2_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_28])])).
cnf(c_0_38,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39,negated_conjecture,(~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_28])])).
cnf(c_0_40,negated_conjecture,(r3_waybel_9(esk1_0,esk2_0,esk5_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_28]), c_0_8]), c_0_9])]), c_0_10]), c_0_37])).
cnf(c_0_41,negated_conjecture,(m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_34]), c_0_35]), c_0_36]), c_0_28]), c_0_8]), c_0_9])]), c_0_10]), c_0_37])).
cnf(c_0_42,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 36
# Proof object formula steps           : 7
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 69
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 37
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 37
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 9
# Generated clauses                    : 15
# ...of the previous two non-trivial   : 18
# Contextual simplify-reflections      : 1
# Paramodulations                      : 15
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 20
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 9
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 50
# Rec. Clause-clause subsumption calls : 16
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2153

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45376 pages
