# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3)))))), file('out/t36_waybel_0__7_premises__rerun__rerun.E_input', t36_waybel_0)).
fof(t16_waybel_0, axiom, (![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,k3_waybel_0(X1,X2))&r1_tarski(X2,k4_waybel_0(X1,X2)))))), file('out/t36_waybel_0__7_premises__rerun__rerun.E_input', t16_waybel_0)).
fof(d8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4))))))), file('out/t36_waybel_0__7_premises__rerun__rerun.E_input', d8_lattice3)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t36_waybel_0__7_premises__rerun__rerun.E_input', d3_tarski)).
fof(t3_orders_2, axiom, (![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X4))=>r1_orders_2(X1,X2,X4))))))), file('out/t36_waybel_0__7_premises__rerun__rerun.E_input', t3_orders_2)).
fof(d16_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k4_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X5,X4)&r2_hidden(X5,X2)))))))))), file('out/t36_waybel_0__7_premises__rerun__rerun.E_input', d16_waybel_0)).
fof(dt_k4_waybel_0, axiom, (![X1]:![X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t36_waybel_0__7_premises__rerun__rerun.E_input', dt_k4_waybel_0)).
fof(c_0_7, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3))))))), inference(assume_negation,[status(cth)],[t36_waybel_0])).
fof(c_0_8, plain, (![X3]:![X4]:(((r1_tarski(X4,k3_waybel_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|(~v3_orders_2(X3)|~l1_orders_2(X3)))&((r1_tarski(X4,k4_waybel_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|(~v3_orders_2(X3)|~l1_orders_2(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_waybel_0])])])])])])).
fof(c_0_9, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0))&(r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])])])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X8]:((((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk4_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,X7,esk4_3(X5,X6,X7))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])])])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_12,plain,(r1_tarski(X2,k4_waybel_0(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,negated_conjecture,(r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)|r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,negated_conjecture,(r1_tarski(X1,k4_waybel_0(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_20,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|r1_lattice3(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,k4_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_14])])).
cnf(c_0_21,negated_conjecture,(r2_hidden(X1,k4_waybel_0(esk1_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|r1_lattice3(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
cnf(c_0_24,plain,(r1_lattice3(X1,X3,X2)|r2_hidden(esk4_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_25,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,esk4_3(X1,esk2_0,X2))|r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(X1,esk2_0,X2)|~m1_subset_1(esk4_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27,negated_conjecture,(r1_lattice3(esk1_0,esk2_0,esk3_0)|~m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_17]), c_0_14])])).
cnf(c_0_28,plain,(r1_lattice3(X1,X3,X2)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_29, plain, (![X5]:![X6]:![X7]:![X8]:((~v4_orders_2(X5)|~l1_orders_2(X5))|(~m1_subset_1(X6,u1_struct_0(X5))|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|((~r1_orders_2(X5,X6,X7)|~r1_orders_2(X5,X7,X8))|r1_orders_2(X5,X6,X8))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_orders_2])])])])])).
fof(c_0_30, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:![X13]:(((((((((m1_subset_1(esk6_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((((r1_orders_2(X6,esk6_4(X6,X7,X8,X9),X9)|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((r2_hidden(esk6_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))))&(((((((~m1_subset_1(X11,u1_struct_0(X6))|(~r1_orders_2(X6,X11,X9)|~r2_hidden(X11,X7)))|r2_hidden(X9,X8))|~m1_subset_1(X9,u1_struct_0(X6)))|X8!=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))&(((((m1_subset_1(esk7_3(X6,X7,X8),u1_struct_0(X6))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((~r2_hidden(esk7_3(X6,X7,X8),X8)|(~m1_subset_1(X13,u1_struct_0(X6))|(~r1_orders_2(X6,X13,esk7_3(X6,X7,X8))|~r2_hidden(X13,X7))))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((m1_subset_1(esk8_3(X6,X7,X8),u1_struct_0(X6))|r2_hidden(esk7_3(X6,X7,X8),X8))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((r1_orders_2(X6,esk8_3(X6,X7,X8),esk7_3(X6,X7,X8))|r2_hidden(esk7_3(X6,X7,X8),X8))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((r2_hidden(esk8_3(X6,X7,X8),X7)|r2_hidden(esk7_3(X6,X7,X8),X8))|X8=k4_waybel_0(X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_waybel_0])])])])])])])).
cnf(c_0_31,negated_conjecture,(r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17]), c_0_14])])).
cnf(c_0_32,plain,(r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33,plain,(r1_orders_2(X1,esk6_4(X1,X2,X3,X4),X4)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_34,plain,(m1_subset_1(esk6_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_31]), c_0_17]), c_0_14])])).
cnf(c_0_36,plain,(r2_hidden(esk6_4(X1,X2,X3,X4),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37,plain,(r1_orders_2(X1,X2,X3)|X4!=k4_waybel_0(X1,X5)|~r1_orders_2(X1,X2,esk6_4(X1,X5,X4,X3))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_38,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,esk6_4(X1,esk2_0,X2,X3))|X2!=k4_waybel_0(X1,esk2_0)|~r2_hidden(X3,X2)|~m1_subset_1(esk6_4(X1,esk2_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|X2!=k4_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk6_4(esk1_0,esk2_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_22]), c_0_17]), c_0_14]), c_0_39])])).
fof(c_0_41, plain, (![X3]:![X4]:((~l1_orders_2(X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|m1_subset_1(k4_waybel_0(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_waybel_0])])).
cnf(c_0_42,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|X2!=k4_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_34]), c_0_22]), c_0_14])])).
cnf(c_0_43,plain,(m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_44,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|k4_waybel_0(esk1_0,X2)!=k4_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,k4_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_14])])).
cnf(c_0_45,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,esk4_3(X1,k4_waybel_0(esk1_0,X2),X3))|r1_lattice3(X1,k4_waybel_0(esk1_0,X2),X3)|k4_waybel_0(esk1_0,X2)!=k4_waybel_0(esk1_0,esk2_0)|~m1_subset_1(esk4_3(X1,k4_waybel_0(esk1_0,X2),X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_24])).
cnf(c_0_46,negated_conjecture,(r1_lattice3(esk1_0,k4_waybel_0(esk1_0,X1),esk3_0)|k4_waybel_0(esk1_0,X1)!=k4_waybel_0(esk1_0,esk2_0)|~m1_subset_1(esk4_3(esk1_0,k4_waybel_0(esk1_0,X1),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_45]), c_0_17]), c_0_14])])).
cnf(c_0_47,negated_conjecture,(~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_48,negated_conjecture,(r1_lattice3(esk1_0,k4_waybel_0(esk1_0,X1),esk3_0)|k4_waybel_0(esk1_0,X1)!=k4_waybel_0(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_28]), c_0_17]), c_0_14])])).
cnf(c_0_49,negated_conjecture,(~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_31])])).
cnf(c_0_50,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_48]), c_0_22])]), c_0_49]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 36
# Proof object formula steps           : 15
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 190
# ...of these trivial                  : 2
# ...subsumed                          : 25
# ...remaining for further processing  : 163
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 12
# Generated clauses                    : 399
# ...of the previous two non-trivial   : 391
# Contextual simplify-reflections      : 46
# Paramodulations                      : 396
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 118
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 109
# Current number of unprocessed clauses: 239
# ...number of literals in the above   : 2661
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 8936
# Rec. Clause-clause subsumption calls : 655
# Non-unit clause-clause subsumptions  : 76
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 25936

# -------------------------------------------------
# User time                : 0.036 s
# System time              : 0.004 s
# Total time               : 0.040 s
# Maximum resident set size: 45360 pages
