# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_wellord1, conjecture, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X1)=>(X4=X5|(r2_hidden(k4_tarski(k1_funct_1(X3,X4),k1_funct_1(X3,X5)),X2)&~(k1_funct_1(X3,X4)=k1_funct_1(X3,X5)))))))))), file('out/t36_wellord1__2_premises__rerun.E_input', t36_wellord1)).
fof(d7_wellord1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)<=>(k9_xtuple_0(X3)=k1_relat_1(X1)&(k10_xtuple_0(X3)=k1_relat_1(X2)&(v2_funct_1(X3)&![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X1)<=>(r2_hidden(X4,k1_relat_1(X1))&(r2_hidden(X5,k1_relat_1(X1))&r2_hidden(k4_tarski(k1_funct_1(X3,X4),k1_funct_1(X3,X5)),X2)))))))))))), file('out/t36_wellord1__2_premises__rerun.E_input', d7_wellord1)).
fof(d4_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2]:![X3]:((r2_hidden(X2,k9_xtuple_0(X1))&(r2_hidden(X3,k9_xtuple_0(X1))&k1_funct_1(X1,X2)=k1_funct_1(X1,X3)))=>X2=X3)))), file('out/t36_wellord1__2_premises__rerun.E_input', d4_funct_1)).
fof(c_0_3, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X1)=>(X4=X5|(r2_hidden(k4_tarski(k1_funct_1(X3,X4),k1_funct_1(X3,X5)),X2)&~(k1_funct_1(X3,X4)=k1_funct_1(X3,X5))))))))))), inference(assume_negation,[status(cth)],[t36_wellord1])).
fof(c_0_4, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X9]:![X10]:((((((k9_xtuple_0(X8)=k1_relat_1(X6)|~r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&(((((k10_xtuple_0(X8)=k1_relat_1(X7)|~r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&(((((v2_funct_1(X8)|~r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&(((((((r2_hidden(X9,k1_relat_1(X6))|~r2_hidden(k4_tarski(X9,X10),X6))|~r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&((((((r2_hidden(X10,k1_relat_1(X6))|~r2_hidden(k4_tarski(X9,X10),X6))|~r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&(((((r2_hidden(k4_tarski(k1_funct_1(X8,X9),k1_funct_1(X8,X10)),X7)|~r2_hidden(k4_tarski(X9,X10),X6))|~r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))))&((((((~r2_hidden(X9,k1_relat_1(X6))|(~r2_hidden(X10,k1_relat_1(X6))|~r2_hidden(k4_tarski(k1_funct_1(X8,X9),k1_funct_1(X8,X10)),X7)))|r2_hidden(k4_tarski(X9,X10),X6))|~r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))))))&(((((((((~r2_hidden(k4_tarski(esk6_3(X6,X7,X8),esk7_3(X6,X7,X8)),X6)|(~r2_hidden(esk6_3(X6,X7,X8),k1_relat_1(X6))|(~r2_hidden(esk7_3(X6,X7,X8),k1_relat_1(X6))|~r2_hidden(k4_tarski(k1_funct_1(X8,esk6_3(X6,X7,X8)),k1_funct_1(X8,esk7_3(X6,X7,X8))),X7))))|~v2_funct_1(X8))|k10_xtuple_0(X8)!=k1_relat_1(X7))|k9_xtuple_0(X8)!=k1_relat_1(X6))|r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&(((((((((r2_hidden(esk6_3(X6,X7,X8),k1_relat_1(X6))|r2_hidden(k4_tarski(esk6_3(X6,X7,X8),esk7_3(X6,X7,X8)),X6))|~v2_funct_1(X8))|k10_xtuple_0(X8)!=k1_relat_1(X7))|k9_xtuple_0(X8)!=k1_relat_1(X6))|r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&(((((((((r2_hidden(esk7_3(X6,X7,X8),k1_relat_1(X6))|r2_hidden(k4_tarski(esk6_3(X6,X7,X8),esk7_3(X6,X7,X8)),X6))|~v2_funct_1(X8))|k10_xtuple_0(X8)!=k1_relat_1(X7))|k9_xtuple_0(X8)!=k1_relat_1(X6))|r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))&((((((((r2_hidden(k4_tarski(k1_funct_1(X8,esk6_3(X6,X7,X8)),k1_funct_1(X8,esk7_3(X6,X7,X8))),X7)|r2_hidden(k4_tarski(esk6_3(X6,X7,X8),esk7_3(X6,X7,X8)),X6))|~v2_funct_1(X8))|k10_xtuple_0(X8)!=k1_relat_1(X7))|k9_xtuple_0(X8)!=k1_relat_1(X6))|r3_wellord1(X6,X7,X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))|~v1_relat_1(X7))|~v1_relat_1(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_wellord1])])])])])])])).
fof(c_0_5, negated_conjecture, ((v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&(r3_wellord1(esk1_0,esk2_0,esk3_0)&(r2_hidden(k4_tarski(esk4_0,esk5_0),esk1_0)&(esk4_0!=esk5_0&(~r2_hidden(k4_tarski(k1_funct_1(esk3_0,esk4_0),k1_funct_1(esk3_0,esk5_0)),esk2_0)|k1_funct_1(esk3_0,esk4_0)=k1_funct_1(esk3_0,esk5_0))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])).
cnf(c_0_6,plain,(r2_hidden(k4_tarski(k1_funct_1(X3,X4),k1_funct_1(X3,X5)),X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r3_wellord1(X1,X2,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,negated_conjecture,(r2_hidden(k4_tarski(esk4_0,esk5_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(r2_hidden(X5,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r3_wellord1(X1,X2,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:(((~v2_funct_1(X4)|((~r2_hidden(X5,k9_xtuple_0(X4))|(~r2_hidden(X6,k9_xtuple_0(X4))|k1_funct_1(X4,X5)!=k1_funct_1(X4,X6)))|X5=X6))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&((((r2_hidden(esk8_1(X4),k9_xtuple_0(X4))|v2_funct_1(X4))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&(((r2_hidden(esk9_1(X4),k9_xtuple_0(X4))|v2_funct_1(X4))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&((k1_funct_1(X4,esk8_1(X4))=k1_funct_1(X4,esk9_1(X4))|v2_funct_1(X4))|(~v1_relat_1(X4)|~v1_funct_1(X4)))))&((esk8_1(X4)!=esk9_1(X4)|v2_funct_1(X4))|(~v1_relat_1(X4)|~v1_funct_1(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_funct_1])])])])])])])).
cnf(c_0_11,negated_conjecture,(k1_funct_1(esk3_0,esk4_0)=k1_funct_1(esk3_0,esk5_0)|~r2_hidden(k4_tarski(k1_funct_1(esk3_0,esk4_0),k1_funct_1(esk3_0,esk5_0)),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,negated_conjecture,(r2_hidden(k4_tarski(k1_funct_1(X1,esk4_0),k1_funct_1(X1,esk5_0)),X2)|~r3_wellord1(esk1_0,X2,X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_13,negated_conjecture,(r3_wellord1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17,plain,(v2_funct_1(X3)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r3_wellord1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18,plain,(k9_xtuple_0(X3)=k1_relat_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r3_wellord1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_19,negated_conjecture,(r2_hidden(esk5_0,k1_relat_1(esk1_0))|~r3_wellord1(esk1_0,X1,X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_20,plain,(r2_hidden(X4,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r3_wellord1(X1,X2,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_21,plain,(X2=X3|~v1_funct_1(X1)|~v1_relat_1(X1)|k1_funct_1(X1,X2)!=k1_funct_1(X1,X3)|~r2_hidden(X3,k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22,negated_conjecture,(k1_funct_1(esk3_0,esk5_0)=k1_funct_1(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_23,negated_conjecture,(v2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_8])])).
cnf(c_0_24,negated_conjecture,(k9_xtuple_0(esk3_0)=k1_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_8])])).
cnf(c_0_25,negated_conjecture,(r2_hidden(esk5_0,k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_13]), c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_26,negated_conjecture,(r2_hidden(esk4_0,k1_relat_1(esk1_0))|~r3_wellord1(esk1_0,X1,X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_7]), c_0_8])])).
cnf(c_0_27,negated_conjecture,(X1=esk5_0|k1_funct_1(esk3_0,X1)!=k1_funct_1(esk3_0,esk4_0)|~r2_hidden(X1,k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_24]), c_0_14]), c_0_15])])).
cnf(c_0_28,negated_conjecture,(r2_hidden(esk4_0,k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_13]), c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_29,negated_conjecture,(esk4_0!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_28])]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 24
# Proof object formula steps           : 7
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 58
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 58
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 29
# ...of the previous two non-trivial   : 22
# Contextual simplify-reflections      : 0
# Paramodulations                      : 27
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 31
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 18
# Current number of unprocessed clauses: 12
# ...number of literals in the above   : 89
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 782
# Rec. Clause-clause subsumption calls : 14
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3641

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45288 pages
