# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d21_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4))))))))), file('out/t37_lattice3__7_premises__rerun.E_input', d21_lattice3)).
fof(dt_k15_lattice3, axiom, (![X1]:![X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1)))), file('out/t37_lattice3__7_premises__rerun.E_input', dt_k15_lattice3)).
fof(fraenkel_a_2_2_lattice3, axiom, (![X1]:![X2]:![X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_2_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r4_lattice3(X2,X4,X3)))))), file('out/t37_lattice3__7_premises__rerun.E_input', fraenkel_a_2_2_lattice3)).
fof(d16_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4))))))), file('out/t37_lattice3__7_premises__rerun.E_input', d16_lattice3)).
fof(t34_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2)))))))), file('out/t37_lattice3__7_premises__rerun.E_input', t34_lattice3)).
fof(redefinition_r3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3)))), file('out/t37_lattice3__7_premises__rerun.E_input', redefinition_r3_lattices)).
fof(cc1_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1)))))))))), file('out/t37_lattice3__7_premises__rerun.E_input', cc1_lattices)).
fof(t37_lattice3, conjecture, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2)))), file('out/t37_lattice3__7_premises__rerun.E_input', t37_lattice3)).
fof(c_0_8, plain, (![X5]:![X6]:![X7]:![X8]:((((((r4_lattice3(X5,X7,X6)|X7!=k15_lattice3(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((((~m1_subset_1(X8,u1_struct_0(X5))|(~r4_lattice3(X5,X8,X6)|r1_lattices(X5,X7,X8)))|X7!=k15_lattice3(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))|(v2_struct_0(X5)|~l3_lattices(X5))))&((((((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r4_lattice3(X5,X7,X6))|X7=k15_lattice3(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((((r4_lattice3(X5,esk1_3(X5,X6,X7),X6)|~r4_lattice3(X5,X7,X6))|X7=k15_lattice3(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((((~r1_lattices(X5,X7,esk1_3(X5,X6,X7))|~r4_lattice3(X5,X7,X6))|X7=k15_lattice3(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))|(v2_struct_0(X5)|~l3_lattices(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])])])).
cnf(c_0_9,plain,(v2_struct_0(X1)|v2_struct_0(X1)|r1_lattices(X1,X2,X4)|~l3_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X2!=k15_lattice3(X1,X3)|~r4_lattice3(X1,X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_10, plain, (![X3]:![X4]:((v2_struct_0(X3)|~l3_lattices(X3))|m1_subset_1(k15_lattice3(X3,X4),u1_struct_0(X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])])])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:((((m1_subset_1(esk2_3(X5,X6,X7),u1_struct_0(X6))|~r2_hidden(X5,a_2_2_lattice3(X6,X7)))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&(((X5=esk2_3(X5,X6,X7)|~r2_hidden(X5,a_2_2_lattice3(X6,X7)))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&((r4_lattice3(X6,esk2_3(X5,X6,X7),X7)|~r2_hidden(X5,a_2_2_lattice3(X6,X7)))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))))&(((~m1_subset_1(X9,u1_struct_0(X6))|(X5!=X9|~r4_lattice3(X6,X9,X7)))|r2_hidden(X5,a_2_2_lattice3(X6,X7)))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_2_lattice3])])])])])])])])).
fof(c_0_12, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:((((~r3_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|r1_lattices(X5,X6,X8))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((m1_subset_1(esk3_3(X5,X6,X7),u1_struct_0(X5))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((r2_hidden(esk3_3(X5,X6,X7),X7)|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((~r1_lattices(X5,X6,esk3_3(X5,X6,X7))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])])).
fof(c_0_13, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r3_lattice3(X5,X6,X7)|X6!=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r3_lattice3(X5,X8,X7)|r3_lattices(X5,X8,X6)))|X6!=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))))&(((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&(((((r3_lattice3(X5,esk4_3(X5,X6,X7),X7)|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&((((~r3_lattices(X5,esk4_3(X5,X6,X7),X6)|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(((~r3_lattices(X4,X5,X6)|r1_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_lattices(X4,X5,X6)|r3_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
fof(c_0_15, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v4_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v5_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v6_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v7_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v8_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&((v9_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_16,plain,(v2_struct_0(X1)|r1_lattices(X1,X2,X4)|X2!=k15_lattice3(X1,X3)|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~r4_lattice3(X1,X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_9])).
cnf(c_0_17,plain,(m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,plain,(v2_struct_0(X1)|X2=esk2_3(X2,X1,X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_2_lattice3(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|r2_hidden(esk3_3(X1,X2,X3),X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(v2_struct_0(X1)|X2=k16_lattice3(X1,X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r3_lattice3(X1,X2,X3)|~r3_lattices(X1,esk4_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,plain,(v2_struct_0(X2)|r3_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r1_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,plain,(v2_struct_0(X1)|X2=k16_lattice3(X1,X3)|m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r3_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,plain,(v2_struct_0(X1)|v6_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24,plain,(v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26,plain,(v2_struct_0(X1)|r1_lattices(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r3_lattice3(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27,plain,(v2_struct_0(X1)|X2=k16_lattice3(X1,X3)|r3_lattice3(X1,esk4_3(X1,X2,X3),X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r3_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_28,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,X2,esk3_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_29,plain,(r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_16]), c_0_17])).
cnf(c_0_30,plain,(v2_struct_0(X1)|r3_lattice3(X1,X2,X3)|m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31,plain,(v2_struct_0(X1)|r4_lattice3(X1,esk2_3(X2,X1,X3),X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_2_lattice3(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_32,plain,(esk2_3(esk3_3(X1,X2,a_2_2_lattice3(X3,X4)),X3,X4)=esk3_3(X1,X2,a_2_2_lattice3(X3,X4))|r3_lattice3(X1,X2,a_2_2_lattice3(X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_lattice3(X3)|~v10_lattices(X3)|~l3_lattices(X3)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_33,plain,(X1=k16_lattice3(X2,X3)|v2_struct_0(X2)|~r3_lattice3(X2,X1,X3)|~r1_lattices(X2,esk4_3(X2,X1,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_34,plain,(X1=k16_lattice3(X2,X3)|r1_lattices(X2,esk4_3(X2,X1,X3),X4)|v2_struct_0(X2)|~r3_lattice3(X2,X1,X3)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_22])).
cnf(c_0_35,plain,(r3_lattice3(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,esk3_3(X1,k15_lattice3(X1,X2),X3),X2)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_17])).
cnf(c_0_36,plain,(r3_lattice3(X1,X2,a_2_2_lattice3(X3,X4))|r4_lattice3(X3,esk3_3(X1,X2,a_2_2_lattice3(X3,X4)),X4)|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_lattice3(X3)|~v10_lattices(X3)|~l3_lattices(X3)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_19])).
cnf(c_0_37,plain,(v2_struct_0(X1)|v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|~l3_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X2!=k15_lattice3(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_38, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))))), inference(assume_negation,[status(cth)],[t37_lattice3])).
cnf(c_0_39,plain,(X1=k16_lattice3(X2,X3)|v2_struct_0(X2)|~r3_lattice3(X2,X1,X3)|~r2_hidden(X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40,plain,(r3_lattice3(X1,k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2))|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_17])).
cnf(c_0_41,plain,(v2_struct_0(X1)|r2_hidden(X2,a_2_2_lattice3(X1,X3))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r4_lattice3(X1,X4,X3)|X2!=X4|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42,plain,(v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_37])).
fof(c_0_43, negated_conjecture, (((~v2_struct_0(esk5_0)&(v10_lattices(esk5_0)&(v4_lattice3(esk5_0)&l3_lattices(esk5_0))))&k15_lattice3(esk5_0,esk6_0)!=k16_lattice3(esk5_0,a_2_2_lattice3(esk5_0,esk6_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_38])])])])])])).
cnf(c_0_44,plain,(k16_lattice3(X1,a_2_2_lattice3(X1,X2))=k15_lattice3(X1,X2)|v2_struct_0(X1)|~r2_hidden(k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_17])).
cnf(c_0_45,plain,(r2_hidden(X1,a_2_2_lattice3(X2,X3))|v2_struct_0(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_46,plain,(r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_42]), c_0_17])).
cnf(c_0_47,negated_conjecture,(k15_lattice3(esk5_0,esk6_0)!=k16_lattice3(esk5_0,a_2_2_lattice3(esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48,plain,(k16_lattice3(X1,a_2_2_lattice3(X1,X2))=k15_lattice3(X1,X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_17]), c_0_46])).
cnf(c_0_49,negated_conjecture,(v4_lattice3(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50,negated_conjecture,(v10_lattices(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51,negated_conjecture,(l3_lattices(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52,negated_conjecture,(~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 37
# Proof object formula steps           : 17
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 91
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 89
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 46
# ...of the previous two non-trivial   : 33
# Contextual simplify-reflections      : 24
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 5
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 46
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 70
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 1997
# Rec. Clause-clause subsumption calls : 79
# Non-unit clause-clause subsumptions  : 27
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5716

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45224 pages
