# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_waybel_9, conjecture, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>(![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X2),X1,X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(v10_waybel_0(X2,X1)=>(r1_waybel_9(X1,X2)&r2_hidden(k1_waybel_2(X1,X2),k10_yellow_6(X1,X2)))))))), file('out/t38_waybel_9__10_premises__rerun.E_input', t38_waybel_9)).
fof(t35_waybel_9, axiom, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v10_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_2(X1,X3)))))), file('out/t38_waybel_9__10_premises__rerun.E_input', t35_waybel_9)).
fof(cc1_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t38_waybel_9__10_premises__rerun.E_input', cc1_lattice3)).
fof(dt_l1_waybel_9, axiom, (![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1)))), file('out/t38_waybel_9__10_premises__rerun.E_input', dt_l1_waybel_9)).
fof(t30_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&(v1_compts_1(X1)&l1_pre_topc(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3))))), file('out/t38_waybel_9__10_premises__rerun.E_input', t30_waybel_9)).
fof(t33_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&(v1_compts_1(X1)&l1_pre_topc(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r3_waybel_9(X1,X2,X3)&r3_waybel_9(X1,X2,X4))=>X3=X4)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)=>r2_hidden(X3,k10_yellow_6(X1,X2)))))))), file('out/t38_waybel_9__10_premises__rerun.E_input', t33_waybel_9)).
fof(l48_waybel_9, axiom, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&(v10_waybel_0(X2,X1)&r3_waybel_9(X1,X2,X3))))=>r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4))))))), file('out/t38_waybel_9__10_premises__rerun.E_input', l48_waybel_9)).
fof(l49_waybel_9, axiom, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&r3_waybel_9(X1,X2,X3)))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)=>r3_orders_2(X1,X4,X5))))))))), file('out/t38_waybel_9__10_premises__rerun.E_input', l49_waybel_9)).
fof(t15_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('out/t38_waybel_9__10_premises__rerun.E_input', t15_yellow_0)).
fof(d3_waybel_9, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(l1_waybel_0(X2,X1)=>(r1_waybel_9(X1,X2)<=>r1_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2))))))), file('out/t38_waybel_9__10_premises__rerun.E_input', d3_waybel_9)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t38_waybel_9__10_premises__rerun.E_input', redefinition_r3_orders_2)).
fof(c_0_11, negated_conjecture, (~(![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>(![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X2),X1,X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(v10_waybel_0(X2,X1)=>(r1_waybel_9(X1,X2)&r2_hidden(k1_waybel_2(X1,X2),k10_yellow_6(X1,X2))))))))), inference(assume_negation,[status(cth)],[t38_waybel_9])).
fof(c_0_12, plain, (![X5]:![X6]:![X7]:((((((m1_subset_1(esk3_3(X5,X6,X7),u1_struct_0(X5))|(~v10_waybel_0(X7,X5)|~r3_waybel_9(X5,X7,X6)))|X6=k1_waybel_2(X5,X7))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X5)))))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v1_compts_1(X5)|~l1_waybel_9(X5))))))))))&(((((~v5_pre_topc(k4_waybel_1(X5,esk3_3(X5,X6,X7)),X5,X5)|(~v10_waybel_0(X7,X5)|~r3_waybel_9(X5,X7,X6)))|X6=k1_waybel_2(X5,X7))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X5)))))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v1_compts_1(X5)|~l1_waybel_9(X5)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_waybel_9])])])])])])])])).
fof(c_0_13, negated_conjecture, (![X4]:((v2_pre_topc(esk1_0)&(v8_pre_topc(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&(v1_compts_1(esk1_0)&l1_waybel_9(esk1_0)))))))))&((~m1_subset_1(X4,u1_struct_0(esk1_0))|v5_pre_topc(k4_waybel_1(esk1_0,X4),esk1_0,esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&(v10_waybel_0(esk2_0,esk1_0)&(~r1_waybel_9(esk1_0,esk2_0)|~r2_hidden(k1_waybel_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0)))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])])).
fof(c_0_14, plain, (![X2]:(~l1_orders_2(X2)|(~v1_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
fof(c_0_15, plain, (![X2]:((l1_pre_topc(X2)|~l1_waybel_9(X2))&(l1_orders_2(X2)|~l1_waybel_9(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
cnf(c_0_16,plain,(v2_struct_0(X3)|X2=k1_waybel_2(X1,X3)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~r3_waybel_9(X1,X3,X2)|~v10_waybel_0(X3,X1)|~v5_pre_topc(k4_waybel_1(X1,esk3_3(X1,X2,X3)),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(v5_pre_topc(k4_waybel_1(esk1_0,X1),esk1_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,(l1_waybel_9(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,negated_conjecture,(v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,negated_conjecture,(v8_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_26,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_27,plain,(~v2_struct_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_28,plain,(l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29,negated_conjecture,(X1=k1_waybel_2(esk1_0,X2)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X2,X1)|~v10_waybel_0(X2,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~m1_subset_1(esk3_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_30,plain,(v2_struct_0(X3)|X2=k1_waybel_2(X1,X3)|m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~r3_waybel_9(X1,X3,X2)|~v10_waybel_0(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_31, plain, (![X4]:![X5]:(((m1_subset_1(esk4_2(X4,X5),u1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v8_pre_topc(X4)|(~v1_compts_1(X4)|~l1_pre_topc(X4))))))&((r3_waybel_9(X4,X5,esk4_2(X4,X5))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v8_pre_topc(X4)|(~v1_compts_1(X4)|~l1_pre_topc(X4)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_waybel_9])])])])])])])])).
cnf(c_0_32,negated_conjecture,(~l1_orders_2(esk1_0)|~v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_21])).
cnf(c_0_33,negated_conjecture,(l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_18])).
cnf(c_0_34,negated_conjecture,(X1=k1_waybel_2(esk1_0,X2)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X2,X1)|~v10_waybel_0(X2,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_35,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,esk4_2(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36,negated_conjecture,(~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33])])).
cnf(c_0_37,negated_conjecture,(esk4_2(esk1_0,X1)=k1_waybel_2(esk1_0,X1)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_38,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, (![X5]:![X6]:![X9]:((((m1_subset_1(esk5_2(X5,X6),u1_struct_0(X5))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5))))))&((((m1_subset_1(esk6_2(X5,X6),u1_struct_0(X5))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5))))))&(((((r3_waybel_9(X5,X6,esk5_2(X5,X6))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5))))))&(((r3_waybel_9(X5,X6,esk6_2(X5,X6))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5)))))))&(((esk5_2(X5,X6)!=esk6_2(X5,X6)|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t33_waybel_9])])])])])])])])).
cnf(c_0_40,negated_conjecture,(esk4_2(esk1_0,X1)=k1_waybel_2(esk1_0,X1)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_41,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|r3_waybel_9(X1,X2,esk5_2(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_42,negated_conjecture,(v2_struct_0(X1)|m1_subset_1(k1_waybel_2(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_40]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_43,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|r3_waybel_9(X1,X2,esk6_2(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,X1))|r2_hidden(k1_waybel_2(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,k1_waybel_2(esk1_0,X2))|~v10_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_45,plain,(l1_pre_topc(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_46,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(k1_waybel_2(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,k1_waybel_2(esk1_0,X2))|~v10_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_42]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_47,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,X1))|r2_hidden(k1_waybel_2(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X1,k1_waybel_2(esk1_0,X2))|~v10_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_18])])).
cnf(c_0_49,negated_conjecture,(r3_waybel_9(esk1_0,X1,k1_waybel_2(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_40]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_50,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(k1_waybel_2(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X1,k1_waybel_2(esk1_0,X2))|~v10_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_45]), c_0_18])])).
cnf(c_0_51,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_52,negated_conjecture,(r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk5_2(esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X2,k1_waybel_2(esk1_0,X1))|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_42]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_53,negated_conjecture,(~r2_hidden(k1_waybel_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~r1_waybel_9(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,X1))|r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_55,negated_conjecture,(v10_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_56,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_57,negated_conjecture,(v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_58,negated_conjecture,(v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_59,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_60,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_50, c_0_49])).
cnf(c_0_61,negated_conjecture,(r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk6_2(esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X2,k1_waybel_2(esk1_0,X1))|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_42]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_62,negated_conjecture,(r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk5_2(esk1_0,X2),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,X2,k1_waybel_2(esk1_0,X1))|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_45]), c_0_18])])).
cnf(c_0_63,negated_conjecture,(r3_waybel_9(esk1_0,esk2_0,esk5_2(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_64,negated_conjecture,(r3_waybel_9(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_60]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_65,negated_conjecture,(r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk6_2(esk1_0,X2),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,X2,k1_waybel_2(esk1_0,X1))|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_45]), c_0_18])])).
fof(c_0_66, plain, (![X6]:![X7]:![X8]:![X9]:((((((((m1_subset_1(esk7_4(X6,X7,X8,X9),u1_struct_0(X6))|(~v10_waybel_0(X7,X6)|~r3_waybel_9(X6,X7,X8)))|X8!=X9)|r2_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X9))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6))))))))))&(((((((~v5_pre_topc(k4_waybel_1(X6,esk7_4(X6,X7,X8,X9)),X6,X6)|(~v10_waybel_0(X7,X6)|~r3_waybel_9(X6,X7,X8)))|X8!=X9)|r2_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X9))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l48_waybel_9])])])])])])])])).
cnf(c_0_67,negated_conjecture,(r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_62, c_0_49])).
cnf(c_0_68,negated_conjecture,(k1_waybel_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_63]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_69,negated_conjecture,(k1_waybel_2(esk1_0,esk2_0)=esk6_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_64]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_70,negated_conjecture,(r2_hidden(k1_waybel_2(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|m1_subset_1(esk6_2(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_65, c_0_49])).
fof(c_0_71, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:((((((((m1_subset_1(esk8_4(X6,X7,X8,X9),u1_struct_0(X6))|~r3_waybel_9(X6,X7,X8))|X8!=X9)|(~m1_subset_1(X11,u1_struct_0(X6))|(~r2_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X11)|r3_orders_2(X6,X9,X11))))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6))))))))))&(((((((~v5_pre_topc(k4_waybel_1(X6,esk8_4(X6,X7,X8,X9)),X6,X6)|~r3_waybel_9(X6,X7,X8))|X8!=X9)|(~m1_subset_1(X11,u1_struct_0(X6))|(~r2_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X11)|r3_orders_2(X6,X9,X11))))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l49_waybel_9])])])])])])])])).
cnf(c_0_72,plain,(v2_struct_0(X2)|r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)|~v5_pre_topc(k4_waybel_1(X1,esk7_4(X1,X2,X3,X4)),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_73,negated_conjecture,(m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_67]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_74,negated_conjecture,(esk6_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75,negated_conjecture,(m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_70]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_76,plain,(v2_struct_0(X2)|r3_orders_2(X1,X4,X5)|m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X5,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_77, plain, (![X5]:![X6]:![X8]:![X6]:![X9]:((((m1_subset_1(esk9_2(X5,X6),u1_struct_0(X5))|~r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_lattice3(X5,X6,esk9_2(X5,X6))|~r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X6,X8)|r1_orders_2(X5,esk9_2(X5,X6),X8)))|~r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))&(((((m1_subset_1(esk10_3(X5,X6,X9),u1_struct_0(X5))|~r2_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((((r2_lattice3(X5,X6,esk10_3(X5,X6,X9))|~r2_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~r1_orders_2(X5,X9,esk10_3(X5,X6,X9))|~r2_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r1_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])])).
cnf(c_0_78,plain,(r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk7_4(X1,X2,X3,X3)),X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_79,plain,(v2_struct_0(X2)|r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|m1_subset_1(esk7_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_80,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,esk2_0))|~r1_waybel_9(esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_73]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_81,negated_conjecture,(r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|m1_subset_1(esk6_2(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,esk2_0))|~r1_waybel_9(esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_73]), c_0_19]), c_0_25]), c_0_26])]), c_0_36])).
cnf(c_0_82,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|esk5_2(X1,X2)!=esk6_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_83,negated_conjecture,(esk6_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_73])).
fof(c_0_84, plain, (![X3]:![X4]:((((~r1_waybel_9(X3,X4)|r1_yellow_0(X3,k2_relset_1(u1_struct_0(X3),u1_waybel_0(X3,X4))))|~l1_waybel_0(X4,X3))|~l1_orders_2(X3))&(((~r1_yellow_0(X3,k2_relset_1(u1_struct_0(X3),u1_waybel_0(X3,X4)))|r1_waybel_9(X3,X4))|~l1_waybel_0(X4,X3))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_waybel_9])])])])])])).
cnf(c_0_85,plain,(r3_orders_2(X1,X2,X3)|v2_struct_0(X4)|m1_subset_1(esk8_4(X1,X4,X2,X2),u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X4)),X3)|~r3_waybel_9(X1,X4,X2)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_86,plain,(r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk10_3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_87,negated_conjecture,(r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X2)|v2_struct_0(X1)|~r3_waybel_9(esk1_0,X1,X2)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(esk7_4(esk1_0,X1,X2,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_88,plain,(r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|m1_subset_1(esk7_4(X1,X2,X3,X3),u1_struct_0(X1))|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_89,negated_conjecture,(k1_waybel_2(esk1_0,X1)=esk6_2(esk1_0,X1)|r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,esk2_0))|~r1_waybel_9(esk1_0,esk2_0)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_80]), c_0_81])).
cnf(c_0_90,negated_conjecture,(r2_hidden(X1,k10_yellow_6(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~r1_waybel_9(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_56]), c_0_57]), c_0_19]), c_0_58]), c_0_25]), c_0_26])]), c_0_59]), c_0_36])).
cnf(c_0_91,plain,(r1_waybel_9(X1,X2)|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~r1_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_92,negated_conjecture,(~l1_pre_topc(esk1_0)|~r2_hidden(esk6_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~r1_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_53, c_0_69])).
cnf(c_0_93,plain,(r1_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))|r3_orders_2(X1,X3,esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4))|v2_struct_0(X2)|m1_subset_1(esk8_4(X1,X2,X3,X3),u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|~r3_waybel_9(X1,X2,X3)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_28])).
cnf(c_0_94,plain,(r1_yellow_0(X1,X2)|m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_95,negated_conjecture,(r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X2)|v2_struct_0(X1)|~r3_waybel_9(esk1_0,X1,X2)|~v10_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_96,negated_conjecture,(r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~r2_hidden(esk6_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~r1_waybel_9(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_89]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59]), c_0_63])).
cnf(c_0_97,negated_conjecture,(r2_hidden(X1,k10_yellow_6(esk1_0,esk2_0))|~r1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_33]), c_0_56])])).
cnf(c_0_98,negated_conjecture,(~l1_pre_topc(esk1_0)|~r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~r1_waybel_9(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_83]), c_0_73])).
cnf(c_0_99,plain,(r1_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))|r3_orders_2(X1,X3,esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4))|v2_struct_0(X2)|m1_subset_1(esk8_4(X1,X2,X3,X3),u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|~r3_waybel_9(X1,X2,X3)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_28])).
cnf(c_0_100,negated_conjecture,(r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_101,negated_conjecture,(~r1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~r1_waybel_9(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_75]), c_0_64]), c_0_98])).
cnf(c_0_102,negated_conjecture,(r1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|r3_orders_2(esk1_0,X1,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X2))|m1_subset_1(esk8_4(esk1_0,esk2_0,X1,X1),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X1)|~r3_waybel_9(esk1_0,esk2_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_56]), c_0_57]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_58]), c_0_23]), c_0_24]), c_0_25]), c_0_26])]), c_0_59])).
cnf(c_0_103,negated_conjecture,(~r1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_91]), c_0_33]), c_0_56])])).
cnf(c_0_104,plain,(v2_struct_0(X2)|r3_orders_2(X1,X4,X5)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X5,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)|~v5_pre_topc(k4_waybel_1(X1,esk8_4(X1,X2,X3,X4)),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_105,negated_conjecture,(r3_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|m1_subset_1(esk8_4(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_35]), c_0_56]), c_0_57]), c_0_19]), c_0_58]), c_0_25]), c_0_26])]), c_0_59]), c_0_36]), c_0_103])).
cnf(c_0_106,plain,(r3_orders_2(X1,X2,X3)|v2_struct_0(X4)|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X4)),X3)|~r3_waybel_9(X1,X4,X2)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~v5_pre_topc(k4_waybel_1(X1,esk8_4(X1,X4,X2,X2)),X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_104])).
cnf(c_0_107,negated_conjecture,(r3_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|m1_subset_1(esk8_4(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_38]), c_0_56]), c_0_57]), c_0_19]), c_0_58]), c_0_25]), c_0_26])]), c_0_59]), c_0_36])).
cnf(c_0_108,plain,(r1_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))|r3_orders_2(X1,X3,esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4))|v2_struct_0(X2)|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|~r3_waybel_9(X1,X2,X3)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk8_4(X1,X2,X3,X3)),X1,X1)|~m1_subset_1(esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_86]), c_0_28])).
cnf(c_0_109,negated_conjecture,(r3_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|m1_subset_1(esk8_4(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_45]), c_0_18])])).
cnf(c_0_110,negated_conjecture,(r1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)))|r3_orders_2(esk1_0,X2,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X3))|v2_struct_0(X1)|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X3)|~r3_waybel_9(esk1_0,X1,X2)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X3),u1_struct_0(esk1_0))|~m1_subset_1(esk8_4(esk1_0,X1,X2,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_111,negated_conjecture,(r3_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|m1_subset_1(esk8_4(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_49]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_112,negated_conjecture,(r3_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|r3_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_56]), c_0_57]), c_0_58])]), c_0_59]), c_0_103])).
cnf(c_0_113,negated_conjecture,(r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,k1_waybel_2(esk1_0,esk2_0))|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_40]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_114,negated_conjecture,(r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~l1_pre_topc(esk1_0)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_49]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_115,negated_conjecture,(r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~l1_pre_topc(esk1_0)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_42]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_116,negated_conjecture,(r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_45]), c_0_18])])).
cnf(c_0_117,negated_conjecture,(r1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_94]), c_0_33]), c_0_22])])).
fof(c_0_118, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_119,negated_conjecture,(r1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_117, c_0_100])).
cnf(c_0_120,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_121,negated_conjecture,(r3_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_49]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59]), c_0_103])).
cnf(c_0_122,plain,(r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,esk10_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_123,negated_conjecture,(r1_orders_2(esk1_0,k1_waybel_2(esk1_0,esk2_0),esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_33]), c_0_24])]), c_0_36])).
cnf(c_0_124,negated_conjecture,(~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_33]), c_0_22])]), c_0_103])).
cnf(c_0_125,negated_conjecture,(~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_2(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_94]), c_0_33]), c_0_22])]), c_0_103])).
cnf(c_0_126,negated_conjecture,(~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,k1_waybel_2(esk1_0,esk2_0))|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_125, c_0_100])).
cnf(c_0_127,negated_conjecture,(~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_49]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_128,negated_conjecture,(~l1_pre_topc(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_42]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_129,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_45]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 130
# Proof object clause steps            : 107
# Proof object formula steps           : 23
# Proof object conjectures             : 82
# Proof object clause conjectures      : 79
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 37
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 65
# Proof object simplifying inferences  : 289
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 42
# Processed clauses                    : 384
# ...of these trivial                  : 0
# ...subsumed                          : 96
# ...remaining for further processing  : 288
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 69
# Backward-rewritten                   : 1
# Generated clauses                    : 497
# ...of the previous two non-trivial   : 481
# Contextual simplify-reflections      : 243
# Paramodulations                      : 493
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 172
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 155
# Current number of unprocessed clauses: 104
# ...number of literals in the above   : 1115
# Current number of archived formulas  : 0
# Current number of archived clauses   : 112
# Clause-clause subsumption calls (NU) : 34705
# Rec. Clause-clause subsumption calls : 1504
# Non-unit clause-clause subsumptions  : 408
# Unit Clause-clause subsumption calls : 103
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 39157

# -------------------------------------------------
# User time                : 0.064 s
# System time              : 0.004 s
# Total time               : 0.068 s
# Maximum resident set size: 43272 pages
