# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_relat_1, conjecture, ((k9_xtuple_0(k1_xboole_0)=k1_xboole_0&k10_xtuple_0(k1_xboole_0)=k1_xboole_0)), file('out/t38_relat_1__4_premises__rerun.E_input', t38_relat_1)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t38_relat_1__4_premises__rerun.E_input', l13_xboole_0)).
fof(fc5_xtuple_0, axiom, (![X1]:(v1_xboole_0(X1)=>v1_xboole_0(k10_xtuple_0(X1)))), file('out/t38_relat_1__4_premises__rerun.E_input', fc5_xtuple_0)).
fof(fc4_xtuple_0, axiom, (![X1]:(v1_xboole_0(X1)=>v1_xboole_0(k9_xtuple_0(X1)))), file('out/t38_relat_1__4_premises__rerun.E_input', fc4_xtuple_0)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t38_relat_1__4_premises__rerun.E_input', fc1_xboole_0)).
fof(c_0_5, negated_conjecture, (~((k9_xtuple_0(k1_xboole_0)=k1_xboole_0&k10_xtuple_0(k1_xboole_0)=k1_xboole_0))), inference(assume_negation,[status(cth)],[t38_relat_1])).
fof(c_0_6, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
fof(c_0_7, plain, (![X2]:(~v1_xboole_0(X2)|v1_xboole_0(k10_xtuple_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc5_xtuple_0])])).
fof(c_0_8, negated_conjecture, ((k9_xtuple_0(k1_xboole_0)!=k1_xboole_0|k10_xtuple_0(k1_xboole_0)!=k1_xboole_0)), inference(fof_nnf,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(v1_xboole_0(k10_xtuple_0(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X2]:(~v1_xboole_0(X2)|v1_xboole_0(k9_xtuple_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_xtuple_0])])).
cnf(c_0_12,negated_conjecture,(k10_xtuple_0(k1_xboole_0)!=k1_xboole_0|k9_xtuple_0(k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(k10_xtuple_0(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_15,plain,(v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(k9_xtuple_0(k1_xboole_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_17,plain,(k9_xtuple_0(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_9, c_0_15])).
cnf(c_0_18,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_14])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 19
# Proof object clause steps            : 9
# Proof object formula steps           : 10
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 13
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 13
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 6
# ...of the previous two non-trivial   : 3
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 8
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 6
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 2
# Rec. Clause-clause subsumption calls : 2
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 272

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.008 s
# Total time               : 0.012 s
# Maximum resident set size: 43028 pages
