# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_xboole_1, axiom, (![X1]:![X2]:r1_tarski(k4_xboole_0(X1,X2),X1)), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', t36_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', t3_boole)).
fof(t38_subset_1, conjecture, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X2))<=>X2=k1_subset_1(X1)))), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', t38_subset_1)).
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', d5_xboole_0)).
fof(d3_subset_1, axiom, (![X1]:k1_subset_1(X1)=k1_xboole_0), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', d3_subset_1)).
fof(d5_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2))), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', d5_subset_1)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', t28_xboole_1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', t2_xboole_1)).
fof(l32_xboole_1, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('out/t38_subset_1__10_premises__rerun__rerun.E_input', l32_xboole_1)).
fof(c_0_10, plain, (![X3]:![X4]:r1_tarski(k4_xboole_0(X3,X4),X3)), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_11, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_12, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_13, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X2))<=>X2=k1_subset_1(X1))))), inference(assume_negation,[status(cth)],[t38_subset_1])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk3_3(X5,X6,X7),X7)|(~r2_hidden(esk3_3(X5,X6,X7),X5)|r2_hidden(esk3_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk3_3(X5,X6,X7),X5)|r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk3_3(X5,X6,X7),X6)|r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
cnf(c_0_15,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, negated_conjecture, ((m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&((~r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))|esk2_0!=k1_subset_1(esk1_0))&(r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))|esk2_0=k1_subset_1(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_19, plain, (![X2]:k1_subset_1(X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[d3_subset_1])).
fof(c_0_20, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(X3))|k3_subset_1(X3,X4)=k4_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_21,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_22, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_23,plain,(r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_24,plain,(k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_17, c_0_16])).
fof(c_0_25, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_26,negated_conjecture,(esk2_0=k1_subset_1(esk1_0)|r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(k1_subset_1(X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|k4_xboole_0(X3,X4)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])])])).
cnf(c_0_30,plain,(X1!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X4,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[c_0_21, c_0_16])).
cnf(c_0_31,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(r2_hidden(X4,X1)|r2_hidden(X4,X3)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35,negated_conjecture,(esk2_0=k1_xboole_0|r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36,plain,(k3_subset_1(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_28, c_0_16])).
cnf(c_0_37,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_38,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39,plain,(~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_40,plain,(k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41,plain,(k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_33])).
cnf(c_0_42,plain,(r2_hidden(X4,X3)|r2_hidden(X4,X1)|X1!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X4,X2)), inference(rw,[status(thm)],[c_0_34, c_0_16])).
cnf(c_0_43,negated_conjecture,(k3_xboole_0(esk2_0,k3_subset_1(esk1_0,esk2_0))=esk2_0|k1_xboole_0=esk2_0), inference(spm,[status(thm)],[c_0_31, c_0_35])).
cnf(c_0_44,negated_conjecture,(k3_subset_1(esk1_0,esk2_0)=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_45,plain,(k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_38, c_0_16])).
cnf(c_0_46,plain,(~r2_hidden(X1,k5_xboole_0(X2,X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_47,plain,(k5_xboole_0(k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_41])).
cnf(c_0_48,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X1)|r2_hidden(esk3_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49,plain,(r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_50,negated_conjecture,(k3_xboole_0(esk2_0,k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0)))=esk2_0|k1_xboole_0=esk2_0), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_40]), c_0_32])])).
cnf(c_0_52,plain,(~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53,plain,(X1=k5_xboole_0(X2,k3_xboole_0(X2,X3))|r2_hidden(esk3_3(X2,X3,X1),X2)|r2_hidden(esk3_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[c_0_48, c_0_16])).
cnf(c_0_54,negated_conjecture,(esk2_0!=k1_subset_1(esk1_0)|~r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_55,negated_conjecture,(k1_xboole_0=esk2_0|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52]), c_0_39])).
cnf(c_0_56,plain,(X1=k1_xboole_0|r2_hidden(esk3_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_41]), c_0_47])).
cnf(c_0_57,negated_conjecture,(esk2_0!=k1_xboole_0|~r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_54, c_0_27])).
cnf(c_0_58,negated_conjecture,(k1_xboole_0=esk2_0), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_59,negated_conjecture,(k1_xboole_0!=esk2_0|~r1_tarski(esk2_0,k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0)))), inference(rw,[status(thm)],[c_0_57, c_0_44])).
cnf(c_0_60,plain,(r1_tarski(esk2_0,X1)), inference(rw,[status(thm)],[c_0_33, c_0_58])).
cnf(c_0_61,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_58])]), c_0_60])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 41
# Proof object formula steps           : 21
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 16
# Processed clauses                    : 90
# ...of these trivial                  : 3
# ...subsumed                          : 18
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 27
# Generated clauses                    : 129
# ...of the previous two non-trivial   : 102
# Contextual simplify-reflections      : 12
# Paramodulations                      : 123
# Factorizations                       : 2
# Equation resolutions                 : 4
# Current number of processed clauses  : 22
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 50
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 121
# Rec. Clause-clause subsumption calls : 88
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 22
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2587

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45416 pages
