# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2))))))))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', t7_yellow_0)).
fof(t30_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))))))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', t30_yellow_0)).
fof(t38_yellow_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)))))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', t38_yellow_0)).
fof(t31_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))))))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', t31_yellow_0)).
fof(t24_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2)))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', t24_orders_2)).
fof(redefinition_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', redefinition_k6_domain_1)).
fof(t8_boole, axiom, (![X1]:![X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2))))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', t8_boole)).
fof(rc1_xboole_0, axiom, (?[X1]:v1_xboole_0(X1)), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', rc1_xboole_0)).
fof(cc3_relset_1, axiom, (![X1]:![X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3)))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', cc3_relset_1)).
fof(dt_u1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', dt_u1_orders_2)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', d1_xboole_0)).
fof(d9_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))))))), file('out/t38_yellow_0__12_premises__rerun__rerun.E_input', d9_orders_2)).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:(((((~r1_lattice3(X4,k1_tarski(X6),X5)|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&(((((~r1_orders_2(X4,X5,X6)|r1_lattice3(X4,k1_tarski(X6),X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&(((((~r2_lattice3(X4,k1_tarski(X6),X5)|r1_orders_2(X4,X6,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&((((~r1_orders_2(X4,X6,X5)|r2_lattice3(X4,k1_tarski(X6),X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])])])).
fof(c_0_13, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r2_lattice3(X5,X7,X6)|(X6!=k1_yellow_0(X5,X7)|~r1_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X7,X8)|r1_orders_2(X5,X6,X8)))|(X6!=k1_yellow_0(X5,X7)|~r1_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k1_yellow_0(X5,X7)|(m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_yellow_0(X5,X7)|(m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k1_yellow_0(X5,X7)|(r2_lattice3(X5,X7,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_yellow_0(X5,X7)|(r2_lattice3(X5,X7,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&((((X6=k1_yellow_0(X5,X7)|(~r1_orders_2(X5,X6,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_yellow_0(X5,X7)|(~r1_orders_2(X5,X6,esk1_3(X5,X6,X7))|~r2_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])])).
cnf(c_0_14,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,k1_tarski(X3),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15,plain,(r2_lattice3(X1,X3,esk1_3(X1,X2,X3))|r1_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16,plain,(m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))))), inference(assume_negation,[status(cth)],[t38_yellow_0])).
fof(c_0_18, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r1_lattice3(X5,X7,X6)|(X6!=k2_yellow_0(X5,X7)|~r2_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X7,X8)|r1_orders_2(X5,X8,X6)))|(X6!=k2_yellow_0(X5,X7)|~r2_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k2_yellow_0(X5,X7)|(m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k2_yellow_0(X5,X7)|(r1_lattice3(X5,X7,esk4_3(X5,X6,X7))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(r1_lattice3(X5,X7,esk4_3(X5,X6,X7))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&((((X6=k2_yellow_0(X5,X7)|(~r1_orders_2(X5,esk4_3(X5,X6,X7),X6)|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(~r1_orders_2(X5,esk4_3(X5,X6,X7),X6)|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])])).
cnf(c_0_19,plain,(r1_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk1_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(r1_yellow_0(X1,k1_tarski(X2))|r1_orders_2(X1,X2,esk1_3(X1,X3,k1_tarski(X2)))|~v5_orders_2(X1)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
fof(c_0_21, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v3_orders_2(X3)|~l1_orders_2(X3)))|(~m1_subset_1(X4,u1_struct_0(X3))|r1_orders_2(X3,X4,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])])])).
fof(c_0_22, negated_conjecture, (((~v2_struct_0(esk2_0)&(v3_orders_2(esk2_0)&(v5_orders_2(esk2_0)&l1_orders_2(esk2_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(~r1_yellow_0(esk2_0,k6_domain_1(u1_struct_0(esk2_0),esk3_0))|~r2_yellow_0(esk2_0,k6_domain_1(u1_struct_0(esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
cnf(c_0_23,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,k1_tarski(X3),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,plain,(r1_lattice3(X1,X3,esk4_3(X1,X2,X3))|r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,plain,(m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|k6_domain_1(X3,X4)=k1_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_27,plain,(r1_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r2_lattice3(X1,k1_tarski(X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28,plain,(r2_lattice3(X1,k1_tarski(X3),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_29,plain,(r1_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31,negated_conjecture,(v3_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34,plain,(r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,esk4_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_35,plain,(r2_yellow_0(X1,k1_tarski(X2))|r1_orders_2(X1,esk4_3(X1,X3,k1_tarski(X2)),X2)|~v5_orders_2(X1)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
fof(c_0_36, plain, (![X3]:![X4]:(~v1_xboole_0(X3)|(X3=X4|~v1_xboole_0(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])])])).
fof(c_0_37, plain, (v1_xboole_0(esk5_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_xboole_0])])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:(~v1_xboole_0(X4)|(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])])])).
fof(c_0_39, plain, (![X2]:(~l1_orders_2(X2)|m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_40,negated_conjecture,(~r2_yellow_0(esk2_0,k6_domain_1(u1_struct_0(esk2_0),esk3_0))|~r1_yellow_0(esk2_0,k6_domain_1(u1_struct_0(esk2_0),esk3_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_41,plain,(k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_42,plain,(r1_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_43,negated_conjecture,(r1_orders_2(esk2_0,esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])]), c_0_33])).
cnf(c_0_44,negated_conjecture,(v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_45,plain,(r2_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r1_lattice3(X1,k1_tarski(X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_46,plain,(r1_lattice3(X1,k1_tarski(X3),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_47, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk6_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
fof(c_0_48, plain, (![X4]:![X5]:![X6]:(((((~r1_orders_2(X4,X5,X6)|r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&((((~r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4))|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])])])).
cnf(c_0_49,plain,(X2=X1|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50,plain,(v1_xboole_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51,plain,(v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_52,plain,(m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|~r2_yellow_0(esk2_0,k1_tarski(esk3_0))|~r1_yellow_0(esk2_0,k1_tarski(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_30])])).
cnf(c_0_54,negated_conjecture,(r1_yellow_0(esk2_0,k1_tarski(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_30]), c_0_32])])).
cnf(c_0_55,plain,(r2_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57,plain,(r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58,plain,(X1=esk5_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59,plain,(v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|~r2_yellow_0(esk2_0,k1_tarski(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_61,negated_conjecture,(r2_yellow_0(esk2_0,k1_tarski(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_43]), c_0_44]), c_0_30]), c_0_32])])).
cnf(c_0_62,plain,(~v1_xboole_0(u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63,plain,(u1_orders_2(X1)=esk5_0|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_64,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_65,plain,(~v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_50])])).
cnf(c_0_66,negated_conjecture,(u1_struct_0(esk2_0)=esk5_0), inference(spm,[status(thm)],[c_0_58, c_0_64])).
cnf(c_0_67,negated_conjecture,(~r1_orders_2(esk2_0,X1,X2)|~m1_subset_1(X2,esk5_0)|~m1_subset_1(X1,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_50]), c_0_32])])).
cnf(c_0_68,negated_conjecture,(m1_subset_1(esk3_0,esk5_0)), inference(rw,[status(thm)],[c_0_30, c_0_66])).
cnf(c_0_69,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_43]), c_0_68])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 70
# Proof object clause steps            : 45
# Proof object formula steps           : 25
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 137
# ...of these trivial                  : 0
# ...subsumed                          : 25
# ...remaining for further processing  : 112
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 112
# ...of the previous two non-trivial   : 106
# Contextual simplify-reflections      : 55
# Paramodulations                      : 108
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 67
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 57
# Current number of unprocessed clauses: 38
# ...number of literals in the above   : 222
# Current number of archived formulas  : 0
# Current number of archived clauses   : 41
# Clause-clause subsumption calls (NU) : 2038
# Rec. Clause-clause subsumption calls : 336
# Non-unit clause-clause subsumptions  : 80
# Unit Clause-clause subsumption calls : 21
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6549

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45368 pages
