# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&(v1_borsuk_1(X3,X1)&m1_pre_topc(X3,X1)))=>(~(r1_tsep_1(X3,X2))=>(v1_borsuk_1(k2_tsep_1(X1,X3,X2),X2)&m1_pre_topc(k2_tsep_1(X1,X3,X2),X2))))))), file('out/t38_tmap_1__13_premises__rerun.E_input', t38_tmap_1)).
fof(t13_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(v4_pre_topc(X3,X2)<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(v4_pre_topc(X4,X1)&k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2))=X3))))))), file('out/t38_tmap_1__13_premises__rerun.E_input', t13_pre_topc)).
fof(d4_tsep_1, axiom, (![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k2_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))))))), file('out/t38_tmap_1__13_premises__rerun.E_input', d4_tsep_1)).
fof(dt_k2_tsep_1, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k2_tsep_1(X1,X2,X3)))&(v1_pre_topc(k2_tsep_1(X1,X2,X3))&m1_pre_topc(k2_tsep_1(X1,X2,X3),X1))))), file('out/t38_tmap_1__13_premises__rerun.E_input', dt_k2_tsep_1)).
fof(t27_tsep_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>(m1_pre_topc(k2_tsep_1(X1,X2,X3),X2)&m1_pre_topc(k2_tsep_1(X1,X2,X3),X3))))))), file('out/t38_tmap_1__13_premises__rerun.E_input', t27_tsep_1)).
fof(redefinition_k9_subset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3))), file('out/t38_tmap_1__13_premises__rerun.E_input', redefinition_k9_subset_1)).
fof(d3_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1))), file('out/t38_tmap_1__13_premises__rerun.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t38_tmap_1__13_premises__rerun.E_input', dt_l1_pre_topc)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t38_tmap_1__13_premises__rerun.E_input', t3_subset)).
fof(reflexivity_r1_tarski, axiom, (![X1]:![X2]:r1_tarski(X1,X1)), file('out/t38_tmap_1__13_premises__rerun.E_input', reflexivity_r1_tarski)).
fof(t1_tsep_1, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t38_tmap_1__13_premises__rerun.E_input', t1_tsep_1)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t38_tmap_1__13_premises__rerun.E_input', dt_m1_pre_topc)).
fof(t11_tsep_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>((v1_borsuk_1(X2,X1)&m1_pre_topc(X2,X1))<=>v4_pre_topc(X3,X1))))))), file('out/t38_tmap_1__13_premises__rerun.E_input', t11_tsep_1)).
fof(cc1_pre_topc, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2)))), file('out/t38_tmap_1__13_premises__rerun.E_input', cc1_pre_topc)).
fof(c_0_14, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&(v1_borsuk_1(X3,X1)&m1_pre_topc(X3,X1)))=>(~(r1_tsep_1(X3,X2))=>(v1_borsuk_1(k2_tsep_1(X1,X3,X2),X2)&m1_pre_topc(k2_tsep_1(X1,X3,X2),X2)))))))), inference(assume_negation,[status(cth)],[t38_tmap_1])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:![X9]:((((((m1_subset_1(esk4_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(X5)))|~v4_pre_topc(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X5))&(((((v4_pre_topc(esk4_3(X5,X6,X7),X5)|~v4_pre_topc(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X5))&((((k9_subset_1(u1_struct_0(X6),esk4_3(X5,X6,X7),k2_struct_0(X6))=X7|~v4_pre_topc(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X5))))&(((((~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X5)))|(~v4_pre_topc(X9,X5)|k9_subset_1(u1_struct_0(X6),X9,k2_struct_0(X6))!=X7))|v4_pre_topc(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_pre_topc])])])])])])])).
fof(c_0_16, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X8!=k2_tsep_1(X5,X6,X7)|u1_struct_0(X8)=k3_xboole_0(u1_struct_0(X6),u1_struct_0(X7)))|(v2_struct_0(X8)|(~v1_pre_topc(X8)|~m1_pre_topc(X8,X5))))|r1_tsep_1(X6,X7))|(v2_struct_0(X7)|~m1_pre_topc(X7,X5)))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|~l1_pre_topc(X5)))&((((((u1_struct_0(X8)!=k3_xboole_0(u1_struct_0(X6),u1_struct_0(X7))|X8=k2_tsep_1(X5,X6,X7))|(v2_struct_0(X8)|(~v1_pre_topc(X8)|~m1_pre_topc(X8,X5))))|r1_tsep_1(X6,X7))|(v2_struct_0(X7)|~m1_pre_topc(X7,X5)))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|~l1_pre_topc(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tsep_1])])])])])])])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:((~v2_struct_0(k2_tsep_1(X4,X5,X6))|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))&((v1_pre_topc(k2_tsep_1(X4,X5,X6))|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))&(m1_pre_topc(k2_tsep_1(X4,X5,X6),X4)|((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tsep_1])])])])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:(((((m1_pre_topc(k2_tsep_1(X4,X5,X6),X5)|r1_tsep_1(X5,X6))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((((m1_pre_topc(k2_tsep_1(X4,X5,X6),X6)|r1_tsep_1(X5,X6))|(v2_struct_0(X6)|~m1_pre_topc(X6,X4)))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_tsep_1])])])])])])])).
fof(c_0_19, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&(v1_borsuk_1(esk3_0,esk1_0)&m1_pre_topc(esk3_0,esk1_0)))&(~r1_tsep_1(esk3_0,esk2_0)&(~v1_borsuk_1(k2_tsep_1(esk1_0,esk3_0,esk2_0),esk2_0)|~m1_pre_topc(k2_tsep_1(esk1_0,esk3_0,esk2_0),esk2_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])])])).
cnf(c_0_20,plain,(v4_pre_topc(X3,X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2))!=X3|~v4_pre_topc(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(X4))|k9_subset_1(X4,X5,X6)=k3_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
fof(c_0_22, plain, (![X2]:(~l1_struct_0(X2)|k2_struct_0(X2)=u1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_23, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_24, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_25, plain, (![X3]:r1_tarski(X3,X3)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[reflexivity_r1_tarski])])).
fof(c_0_26, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(X3)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])])])).
cnf(c_0_27,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_tsep_1(X2,X3)|v2_struct_0(X4)|u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_pre_topc(X4)|X4!=k2_tsep_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28,plain,(v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|m1_pre_topc(k2_tsep_1(X2,X3,X1),X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29,plain,(v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v1_pre_topc(k2_tsep_1(X2,X3,X1))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30,plain,(v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_struct_0(k2_tsep_1(X2,X3,X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_tsep_1(X2,X3)|m1_pre_topc(k2_tsep_1(X1,X2,X3),X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32,negated_conjecture,(m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_35,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_37, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
cnf(c_0_38,plain,(v4_pre_topc(k9_subset_1(u1_struct_0(X1),X2,k2_struct_0(X1)),X1)|~v4_pre_topc(X2,X3)|~m1_subset_1(k9_subset_1(u1_struct_0(X1),X2,k2_struct_0(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_39,plain,(k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_40,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_41,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_42,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_43,plain,(r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_44,plain,(m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_45,plain,(u1_struct_0(k2_tsep_1(X1,X2,X3))=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_46,negated_conjecture,(~r1_tsep_1(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47,negated_conjecture,(r1_tsep_1(X1,esk2_0)|m1_pre_topc(k2_tsep_1(esk1_0,X1,esk2_0),esk2_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])]), c_0_35]), c_0_36])).
cnf(c_0_48,negated_conjecture,(m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_49,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_50,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_51, plain, (![X4]:![X5]:![X6]:(((((((~v1_borsuk_1(X5,X4)|~m1_pre_topc(X5,X4))|v4_pre_topc(X6,X4))|X6!=u1_struct_0(X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4))))|~m1_pre_topc(X5,X4))|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))&((((((v1_borsuk_1(X5,X4)|~v4_pre_topc(X6,X4))|X6!=u1_struct_0(X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4))))|~m1_pre_topc(X5,X4))|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))&(((((m1_pre_topc(X5,X4)|~v4_pre_topc(X6,X4))|X6!=u1_struct_0(X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4))))|~m1_pre_topc(X5,X4))|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_tsep_1])])])])])])).
cnf(c_0_52,plain,(v4_pre_topc(k3_xboole_0(X1,k2_struct_0(X2)),X2)|~v4_pre_topc(X1,X3)|~m1_subset_1(k3_xboole_0(X1,k2_struct_0(X2)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(k2_struct_0(X2),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X2,X3)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_53,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_54,plain,(m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_55,plain,(m1_subset_1(k3_xboole_0(u1_struct_0(X1),u1_struct_0(X2)),k1_zfmisc_1(u1_struct_0(X3)))|r1_tsep_1(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|~m1_pre_topc(k2_tsep_1(X4,X1,X2),X3)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X3)|~l1_pre_topc(X4)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_56,negated_conjecture,(m1_pre_topc(k2_tsep_1(esk1_0,esk3_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_57,negated_conjecture,(l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_32]), c_0_33])])).
cnf(c_0_58,plain,(v4_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=u1_struct_0(X2)|~m1_pre_topc(X2,X1)|~v1_borsuk_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59,plain,(v4_pre_topc(k3_xboole_0(X1,u1_struct_0(X2)),X2)|~v4_pre_topc(X1,X3)|~m1_subset_1(k3_xboole_0(X1,u1_struct_0(X2)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X2,X3)|~l1_pre_topc(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])]), c_0_50])).
cnf(c_0_60,negated_conjecture,(m1_subset_1(k3_xboole_0(u1_struct_0(esk3_0),u1_struct_0(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_32]), c_0_48]), c_0_57]), c_0_33])]), c_0_46]), c_0_35]), c_0_49]), c_0_36])).
cnf(c_0_61,plain,(v4_pre_topc(X3,X1)|X3!=u1_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v1_borsuk_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(cn,[status(thm)],[c_0_58])).
fof(c_0_62, plain, (![X3]:![X4]:((~v2_pre_topc(X3)|~l1_pre_topc(X3))|(~m1_pre_topc(X4,X3)|v2_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])])])).
cnf(c_0_63,negated_conjecture,(v4_pre_topc(k3_xboole_0(u1_struct_0(esk3_0),u1_struct_0(esk2_0)),esk2_0)|~v4_pre_topc(u1_struct_0(esk3_0),X1)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_64,plain,(v4_pre_topc(u1_struct_0(X1),X2)|~v1_borsuk_1(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_61]), c_0_44])).
cnf(c_0_65,negated_conjecture,(v1_borsuk_1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_66,negated_conjecture,(~m1_pre_topc(k2_tsep_1(esk1_0,esk3_0,esk2_0),esk2_0)|~v1_borsuk_1(k2_tsep_1(esk1_0,esk3_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_67,plain,(v1_borsuk_1(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=u1_struct_0(X2)|~v4_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68,plain,(v2_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69,negated_conjecture,(v4_pre_topc(k3_xboole_0(u1_struct_0(esk3_0),u1_struct_0(esk2_0)),esk2_0)|~v4_pre_topc(u1_struct_0(esk3_0),X1)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_63, c_0_44])).
cnf(c_0_70,negated_conjecture,(v4_pre_topc(u1_struct_0(esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_48]), c_0_33]), c_0_34])])).
cnf(c_0_71,negated_conjecture,(~v1_borsuk_1(k2_tsep_1(esk1_0,esk3_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_56])])).
cnf(c_0_72,plain,(v1_borsuk_1(X1,X2)|~v4_pre_topc(u1_struct_0(X1),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_67]), c_0_44])).
cnf(c_0_73,negated_conjecture,(v2_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_74,negated_conjecture,(v4_pre_topc(k3_xboole_0(u1_struct_0(esk3_0),u1_struct_0(esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_32]), c_0_48]), c_0_33])])).
cnf(c_0_75,negated_conjecture,(~v4_pre_topc(u1_struct_0(k2_tsep_1(esk1_0,esk3_0,esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_56]), c_0_57]), c_0_73])])).
cnf(c_0_76,negated_conjecture,(v4_pre_topc(u1_struct_0(k2_tsep_1(X1,esk3_0,esk2_0)),esk2_0)|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_45]), c_0_46]), c_0_49]), c_0_35])).
cnf(c_0_77,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_32]), c_0_48]), c_0_33])]), c_0_36]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 78
# Proof object clause steps            : 49
# Proof object formula steps           : 29
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 131
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 153
# ...of the previous two non-trivial   : 142
# Contextual simplify-reflections      : 13
# Paramodulations                      : 149
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 89
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 64
# Current number of unprocessed clauses: 73
# ...number of literals in the above   : 818
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 2850
# Rec. Clause-clause subsumption calls : 174
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 19
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10484

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.004 s
# Total time               : 0.024 s
# Maximum resident set size: 45368 pages
