# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(cc1_relat_1, axiom, (![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1))), file('out/t39_relat_1__5_premises__rerun.E_input', cc1_relat_1)).
fof(dt_o_0_0_xboole_0, axiom, (v1_xboole_0(o_0_0_xboole_0)), file('out/t39_relat_1__5_premises__rerun.E_input', dt_o_0_0_xboole_0)).
fof(d2_xboole_0, axiom, (k1_xboole_0=o_0_0_xboole_0), file('out/t39_relat_1__5_premises__rerun.E_input', d2_xboole_0)).
fof(d8_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(X3=k3_relat_1(X1,X2)<=>![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2)))))), file('out/t39_relat_1__5_premises__rerun.E_input', d8_relat_1)).
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t39_relat_1__5_premises__rerun.E_input', t7_boole)).
fof(t39_relat_1, conjecture, (![X1]:(v1_relat_1(X1)=>(k3_relat_1(k1_xboole_0,X1)=k1_xboole_0&k3_relat_1(X1,k1_xboole_0)=k1_xboole_0))), file('out/t39_relat_1__5_premises__rerun.E_input', t39_relat_1)).
fof(c_0_6, plain, (![X2]:(~v1_xboole_0(X2)|v1_relat_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
cnf(c_0_7,plain,(v1_xboole_0(o_0_0_xboole_0)), inference(split_conjunct,[status(thm)],[dt_o_0_0_xboole_0])).
cnf(c_0_8,plain,(k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
fof(c_0_9, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X10]:![X11]:![X13]:![X16]:((((((r2_hidden(k4_tarski(X10,esk1_5(X7,X8,X9,X10,X11)),X7)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9))&(((r2_hidden(k4_tarski(esk1_5(X7,X8,X9,X10,X11),X11),X8)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9)))&((((~r2_hidden(k4_tarski(X10,X13),X7)|~r2_hidden(k4_tarski(X13,X11),X8))|r2_hidden(k4_tarski(X10,X11),X9))|X9!=k3_relat_1(X7,X8))|~v1_relat_1(X9)))&((((~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9)|(~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),X16),X7)|~r2_hidden(k4_tarski(X16,esk3_3(X7,X8,X9)),X8)))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9))&((((r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk4_3(X7,X8,X9)),X7)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9))&(((r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk3_3(X7,X8,X9)),X8)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9))|X9=k3_relat_1(X7,X8))|~v1_relat_1(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])])).
cnf(c_0_10,plain,(v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_7, c_0_8])).
fof(c_0_12, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_13,plain,(X1=k3_relat_1(X2,X3)|r2_hidden(k4_tarski(esk2_3(X2,X3,X1),esk3_3(X2,X3,X1)),X1)|r2_hidden(k4_tarski(esk2_3(X2,X3,X1),esk4_3(X2,X3,X1)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(k3_relat_1(X1,X2)=k1_xboole_0|r2_hidden(k4_tarski(esk2_3(X1,X2,k1_xboole_0),esk3_3(X1,X2,k1_xboole_0)),k1_xboole_0)|r2_hidden(k4_tarski(esk2_3(X1,X2,k1_xboole_0),esk4_3(X1,X2,k1_xboole_0)),X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17,plain,(r2_hidden(k4_tarski(X4,esk1_5(X2,X3,X1,X4,X5)),X2)|~v1_relat_1(X1)|X1!=k3_relat_1(X2,X3)|~r2_hidden(k4_tarski(X4,X5),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,plain,(k3_relat_1(X1,X2)=k1_xboole_0|r2_hidden(k4_tarski(esk2_3(X1,X2,k1_xboole_0),esk3_3(X1,X2,k1_xboole_0)),k1_xboole_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(r2_hidden(k4_tarski(X1,esk1_5(X2,X3,k3_relat_1(X2,X3),X1,X4)),X2)|~r2_hidden(k4_tarski(X1,X4),k3_relat_1(X2,X3))|~v1_relat_1(k3_relat_1(X2,X3))), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_20,plain,(k3_relat_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_18]), c_0_11])])).
fof(c_0_21, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>(k3_relat_1(k1_xboole_0,X1)=k1_xboole_0&k3_relat_1(X1,k1_xboole_0)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t39_relat_1])).
cnf(c_0_22,plain,(~r2_hidden(k4_tarski(X1,X2),k3_relat_1(X3,X4))|~v1_relat_1(k3_relat_1(X3,X4))|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_15, c_0_19])).
cnf(c_0_23,plain,(k3_relat_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_11])).
fof(c_0_24, negated_conjecture, ((v1_relat_1(esk5_0)&(k3_relat_1(k1_xboole_0,esk5_0)!=k1_xboole_0|k3_relat_1(esk5_0,k1_xboole_0)!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
cnf(c_0_25,plain,(X1=k3_relat_1(X2,X3)|r2_hidden(k4_tarski(esk2_3(X2,X3,X1),esk3_3(X2,X3,X1)),X1)|r2_hidden(k4_tarski(esk4_3(X2,X3,X1),esk3_3(X2,X3,X1)),X3)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26,plain,(~r2_hidden(k4_tarski(X1,X2),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_14]), c_0_11])])).
cnf(c_0_27,negated_conjecture,(k3_relat_1(esk5_0,k1_xboole_0)!=k1_xboole_0|k3_relat_1(k1_xboole_0,esk5_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28,plain,(k3_relat_1(X1,X2)=k1_xboole_0|r2_hidden(k4_tarski(esk4_3(X1,X2,k1_xboole_0),esk3_3(X1,X2,k1_xboole_0)),X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_14]), c_0_26])).
cnf(c_0_29,negated_conjecture,(k3_relat_1(esk5_0,k1_xboole_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_23])])).
cnf(c_0_30,plain,(k3_relat_1(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_28])).
cnf(c_0_31,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 32
# Proof object clause steps            : 21
# Proof object formula steps           : 11
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 700
# ...of these trivial                  : 0
# ...subsumed                          : 565
# ...remaining for further processing  : 135
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 10
# Generated clauses                    : 3338
# ...of the previous two non-trivial   : 3094
# Contextual simplify-reflections      : 716
# Paramodulations                      : 3334
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 107
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 100
# Current number of unprocessed clauses: 1900
# ...number of literals in the above   : 13252
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 14706
# Rec. Clause-clause subsumption calls : 6389
# Non-unit clause-clause subsumptions  : 1004
# Unit Clause-clause subsumption calls : 73
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 92397

# -------------------------------------------------
# User time                : 0.080 s
# System time              : 0.008 s
# Total time               : 0.088 s
# Maximum resident set size: 45280 pages
