# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4))))))), file('out/t39_waybel_0__4_premises__rerun.E_input', d8_lattice3)).
fof(t31_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))))))), file('out/t39_waybel_0__4_premises__rerun.E_input', t31_yellow_0)).
fof(t18_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3)))))), file('out/t39_waybel_0__4_premises__rerun.E_input', t18_waybel_0)).
fof(t39_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2)))), file('out/t39_waybel_0__4_premises__rerun.E_input', t39_waybel_0)).
fof(t24_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2)))), file('out/t39_waybel_0__4_premises__rerun.E_input', t24_orders_2)).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:((((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk1_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,X7,esk1_3(X5,X6,X7))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])])])).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r1_lattice3(X5,X7,X6)|(X6!=k2_yellow_0(X5,X7)|~r2_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X7,X8)|r1_orders_2(X5,X8,X6)))|(X6!=k2_yellow_0(X5,X7)|~r2_yellow_0(X5,X7)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k2_yellow_0(X5,X7)|(m1_subset_1(esk2_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(m1_subset_1(esk2_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&(((((X6=k2_yellow_0(X5,X7)|(r1_lattice3(X5,X7,esk2_3(X5,X6,X7))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(r1_lattice3(X5,X7,esk2_3(X5,X6,X7))|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5))))&((((X6=k2_yellow_0(X5,X7)|(~r1_orders_2(X5,esk2_3(X5,X6,X7),X6)|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r2_yellow_0(X5,X7)|(~r1_orders_2(X5,esk2_3(X5,X6,X7),X6)|~r1_lattice3(X5,X7,X6)))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])])).
fof(c_0_7, plain, (![X4]:![X5]:![X6]:(((((~r2_hidden(X6,k6_waybel_0(X4,X5))|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((~r1_orders_2(X4,X5,X6)|r2_hidden(X6,k6_waybel_0(X4,X5)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])])])).
cnf(c_0_8,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(r1_lattice3(X1,X3,esk2_3(X1,X2,X3))|r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k6_waybel_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(r1_lattice3(X1,X3,X2)|r2_hidden(esk1_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13,plain,(r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,plain,(r2_yellow_0(X1,X2)|r1_orders_2(X1,esk2_3(X1,X3,X2),X4)|~v5_orders_2(X1)|~r2_hidden(X4,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])).
cnf(c_0_15,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk1_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16,plain,(v2_struct_0(X1)|r1_orders_2(X1,X2,esk1_3(X3,k6_waybel_0(X1,X2),X4))|r1_lattice3(X3,k6_waybel_0(X1,X2),X4)|~m1_subset_1(esk1_3(X3,k6_waybel_0(X1,X2),X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X1)|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17,plain,(r1_lattice3(X1,X3,X2)|m1_subset_1(esk1_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_18, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2))))), inference(assume_negation,[status(cth)],[t39_waybel_0])).
cnf(c_0_19,plain,(r1_lattice3(X1,X3,esk2_3(X1,X2,X3))|X2=k2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_20,plain,(m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|X2=k2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21,plain,(r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~r2_hidden(X3,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_22,plain,(v2_struct_0(X1)|r1_lattice3(X1,k6_waybel_0(X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
fof(c_0_23, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v3_orders_2(X3)|~l1_orders_2(X3)))|(~m1_subset_1(X4,u1_struct_0(X3))|r1_orders_2(X3,X4,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])])])).
fof(c_0_24, negated_conjecture, (((~v2_struct_0(esk3_0)&(v3_orders_2(esk3_0)&(v4_orders_2(esk3_0)&(v5_orders_2(esk3_0)&l1_orders_2(esk3_0)))))&(m1_subset_1(esk4_0,u1_struct_0(esk3_0))&(~r2_yellow_0(esk3_0,k6_waybel_0(esk3_0,esk4_0))|k2_yellow_0(esk3_0,k6_waybel_0(esk3_0,esk4_0))!=esk4_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])])])).
cnf(c_0_25,plain,(X2=k2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26,plain,(X1=k2_yellow_0(X2,X3)|r1_orders_2(X2,esk2_3(X2,X1,X3),X4)|~v5_orders_2(X2)|~r2_hidden(X4,X3)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_19]), c_0_20])).
cnf(c_0_27,plain,(v2_struct_0(X1)|r2_yellow_0(X1,k6_waybel_0(X1,X2))|~v5_orders_2(X1)|~r2_hidden(X2,k6_waybel_0(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28,plain,(v2_struct_0(X1)|r2_hidden(X3,k6_waybel_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_29,plain,(r1_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,negated_conjecture,(v3_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,negated_conjecture,(l1_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,plain,(X1=k2_yellow_0(X2,X3)|~v5_orders_2(X2)|~r2_hidden(X1,X3)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34,plain,(v2_struct_0(X1)|r2_yellow_0(X1,k6_waybel_0(X1,X2))|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35,negated_conjecture,(r1_orders_2(esk3_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_36,negated_conjecture,(v5_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37,plain,(k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2|v2_struct_0(X1)|~v5_orders_2(X1)|~r2_hidden(X2,k6_waybel_0(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_33, c_0_22])).
cnf(c_0_38,negated_conjecture,(k2_yellow_0(esk3_0,k6_waybel_0(esk3_0,esk4_0))!=esk4_0|~r2_yellow_0(esk3_0,k6_waybel_0(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_39,negated_conjecture,(r2_yellow_0(esk3_0,k6_waybel_0(esk3_0,X1))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_31])]), c_0_32])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41,plain,(k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2|v2_struct_0(X1)|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_28])).
cnf(c_0_42,negated_conjecture,(k2_yellow_0(esk3_0,k6_waybel_0(esk3_0,esk4_0))!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
cnf(c_0_43,negated_conjecture,(k2_yellow_0(esk3_0,k6_waybel_0(esk3_0,X1))=X1|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_35]), c_0_36]), c_0_31])]), c_0_32])).
cnf(c_0_44,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_40])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 34
# Proof object formula steps           : 11
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 71
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 67
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 52
# ...of the previous two non-trivial   : 47
# Contextual simplify-reflections      : 11
# Paramodulations                      : 50
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 42
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 138
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 333
# Rec. Clause-clause subsumption calls : 38
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3939

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45300 pages
