# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_waybel_9, conjecture, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>(![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X2),X1,X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(v11_waybel_0(X2,X1)=>(r2_waybel_9(X1,X2)&r2_hidden(k1_waybel_9(X1,X2),k10_yellow_6(X1,X2)))))))), file('out/t39_waybel_9__9_premises__rerun.E_input', t39_waybel_9)).
fof(t36_waybel_9, axiom, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v11_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_9(X1,X3)))))), file('out/t39_waybel_9__9_premises__rerun.E_input', t36_waybel_9)).
fof(cc1_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1))))), file('out/t39_waybel_9__9_premises__rerun.E_input', cc1_lattice3)).
fof(dt_l1_waybel_9, axiom, (![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1)))), file('out/t39_waybel_9__9_premises__rerun.E_input', dt_l1_waybel_9)).
fof(t30_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&(v1_compts_1(X1)&l1_pre_topc(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3))))), file('out/t39_waybel_9__9_premises__rerun.E_input', t30_waybel_9)).
fof(t33_waybel_9, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&(v1_compts_1(X1)&l1_pre_topc(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r3_waybel_9(X1,X2,X3)&r3_waybel_9(X1,X2,X4))=>X3=X4)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)=>r2_hidden(X3,k10_yellow_6(X1,X2)))))))), file('out/t39_waybel_9__9_premises__rerun.E_input', t33_waybel_9)).
fof(l51_waybel_9, axiom, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&(v11_waybel_0(X2,X1)&r3_waybel_9(X1,X2,X3))))=>r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4))))))), file('out/t39_waybel_9__9_premises__rerun.E_input', l51_waybel_9)).
fof(d4_waybel_9, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(l1_waybel_0(X2,X1)=>(r2_waybel_9(X1,X2)<=>r2_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2))))))), file('out/t39_waybel_9__9_premises__rerun.E_input', d4_waybel_9)).
fof(l52_waybel_9, axiom, (![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&r3_waybel_9(X1,X2,X3)))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)=>r1_orders_2(X1,X5,X4))))))))), file('out/t39_waybel_9__9_premises__rerun.E_input', l52_waybel_9)).
fof(t16_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('out/t39_waybel_9__9_premises__rerun.E_input', t16_yellow_0)).
fof(c_0_10, negated_conjecture, (~(![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>(![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X2),X1,X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(v11_waybel_0(X2,X1)=>(r2_waybel_9(X1,X2)&r2_hidden(k1_waybel_9(X1,X2),k10_yellow_6(X1,X2))))))))), inference(assume_negation,[status(cth)],[t39_waybel_9])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:((((((m1_subset_1(esk3_3(X5,X6,X7),u1_struct_0(X5))|(~v11_waybel_0(X7,X5)|~r3_waybel_9(X5,X7,X6)))|X6=k1_waybel_9(X5,X7))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X5)))))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v1_compts_1(X5)|~l1_waybel_9(X5))))))))))&(((((~v5_pre_topc(k4_waybel_1(X5,esk3_3(X5,X6,X7)),X5,X5)|(~v11_waybel_0(X7,X5)|~r3_waybel_9(X5,X7,X6)))|X6=k1_waybel_9(X5,X7))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X5)))))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v1_compts_1(X5)|~l1_waybel_9(X5)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_waybel_9])])])])])])])])).
fof(c_0_12, negated_conjecture, (![X4]:((v2_pre_topc(esk1_0)&(v8_pre_topc(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&(v1_compts_1(esk1_0)&l1_waybel_9(esk1_0)))))))))&((~m1_subset_1(X4,u1_struct_0(esk1_0))|v5_pre_topc(k4_waybel_1(esk1_0,X4),esk1_0,esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&(v11_waybel_0(esk2_0,esk1_0)&(~r2_waybel_9(esk1_0,esk2_0)|~r2_hidden(k1_waybel_9(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0)))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])])).
fof(c_0_13, plain, (![X2]:(~l1_orders_2(X2)|(~v1_lattice3(X2)|~v2_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
fof(c_0_14, plain, (![X2]:((l1_pre_topc(X2)|~l1_waybel_9(X2))&(l1_orders_2(X2)|~l1_waybel_9(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
cnf(c_0_15,plain,(v2_struct_0(X3)|X2=k1_waybel_9(X1,X3)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~r3_waybel_9(X1,X3,X2)|~v11_waybel_0(X3,X1)|~v5_pre_topc(k4_waybel_1(X1,esk3_3(X1,X2,X3)),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(v5_pre_topc(k4_waybel_1(esk1_0,X1),esk1_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(l1_waybel_9(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,negated_conjecture,(v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,negated_conjecture,(v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,negated_conjecture,(v8_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_26,plain,(~v2_struct_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_27,plain,(l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_28,negated_conjecture,(X1=k1_waybel_9(esk1_0,X2)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X2,X1)|~v11_waybel_0(X2,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~m1_subset_1(esk3_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_29,plain,(v2_struct_0(X3)|X2=k1_waybel_9(X1,X3)|m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~r3_waybel_9(X1,X3,X2)|~v11_waybel_0(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_30, plain, (![X4]:![X5]:(((m1_subset_1(esk4_2(X4,X5),u1_struct_0(X4))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v8_pre_topc(X4)|(~v1_compts_1(X4)|~l1_pre_topc(X4))))))&((r3_waybel_9(X4,X5,esk4_2(X4,X5))|(v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v8_pre_topc(X4)|(~v1_compts_1(X4)|~l1_pre_topc(X4)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_waybel_9])])])])])])])])).
cnf(c_0_31,negated_conjecture,(~l1_orders_2(esk1_0)|~v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_26, c_0_20])).
cnf(c_0_32,negated_conjecture,(l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_17])).
cnf(c_0_33,negated_conjecture,(X1=k1_waybel_9(esk1_0,X2)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X2,X1)|~v11_waybel_0(X2,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_34,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,esk4_2(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35,negated_conjecture,(~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32])])).
cnf(c_0_36,negated_conjecture,(esk4_2(esk1_0,X1)=k1_waybel_9(esk1_0,X1)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_37,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_38, plain, (![X5]:![X6]:![X9]:((((m1_subset_1(esk5_2(X5,X6),u1_struct_0(X5))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5))))))&((((m1_subset_1(esk6_2(X5,X6),u1_struct_0(X5))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5))))))&(((((r3_waybel_9(X5,X6,esk5_2(X5,X6))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5))))))&(((r3_waybel_9(X5,X6,esk6_2(X5,X6))|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5)))))))&(((esk5_2(X5,X6)!=esk6_2(X5,X6)|(~m1_subset_1(X9,u1_struct_0(X5))|(~r3_waybel_9(X5,X6,X9)|r2_hidden(X9,k10_yellow_6(X5,X6)))))|(v2_struct_0(X6)|(~v4_orders_2(X6)|(~v7_waybel_0(X6)|~l1_waybel_0(X6,X5)))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|(~v8_pre_topc(X5)|(~v1_compts_1(X5)|~l1_pre_topc(X5)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t33_waybel_9])])])])])])])])).
cnf(c_0_39,negated_conjecture,(esk4_2(esk1_0,X1)=k1_waybel_9(esk1_0,X1)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_40,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|r3_waybel_9(X1,X2,esk5_2(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_41,negated_conjecture,(v2_struct_0(X1)|m1_subset_1(k1_waybel_9(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_39]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_42,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|r3_waybel_9(X1,X2,esk6_2(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,X1))|r2_hidden(k1_waybel_9(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,k1_waybel_9(esk1_0,X2))|~v11_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_44,plain,(l1_pre_topc(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_45,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(k1_waybel_9(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,k1_waybel_9(esk1_0,X2))|~v11_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_41]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_46,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,X1))|r2_hidden(k1_waybel_9(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X1,k1_waybel_9(esk1_0,X2))|~v11_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_17])])).
cnf(c_0_47,negated_conjecture,(r3_waybel_9(esk1_0,X1,k1_waybel_9(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_39]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_48,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(k1_waybel_9(esk1_0,X2),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(esk1_0,X1,k1_waybel_9(esk1_0,X2))|~v11_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_44]), c_0_17])])).
cnf(c_0_49,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51,negated_conjecture,(~r2_hidden(k1_waybel_9(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~r2_waybel_9(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_52,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,X1))|r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53,negated_conjecture,(v11_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_54,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_55,negated_conjecture,(v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_56,negated_conjecture,(v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_57,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_58,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_48, c_0_47])).
cnf(c_0_59,negated_conjecture,(r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk6_2(esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X2,k1_waybel_9(esk1_0,X1))|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_41]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_60,negated_conjecture,(r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk5_2(esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X2,k1_waybel_9(esk1_0,X1))|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_41]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_61,negated_conjecture,(r3_waybel_9(esk1_0,esk2_0,esk5_2(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_62,negated_conjecture,(r3_waybel_9(esk1_0,esk2_0,esk6_2(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_58]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_63,negated_conjecture,(r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk6_2(esk1_0,X2),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,X2,k1_waybel_9(esk1_0,X1))|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_44]), c_0_17])])).
cnf(c_0_64,negated_conjecture,(r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X2))|v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk5_2(esk1_0,X2),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,X2,k1_waybel_9(esk1_0,X1))|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X2)|~v7_waybel_0(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_44]), c_0_17])])).
cnf(c_0_65,negated_conjecture,(k1_waybel_9(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_61]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_66,negated_conjecture,(k1_waybel_9(esk1_0,esk2_0)=esk6_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_62]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_67,negated_conjecture,(r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|m1_subset_1(esk6_2(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_63, c_0_47])).
cnf(c_0_68,negated_conjecture,(r2_hidden(k1_waybel_9(esk1_0,X1),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_64, c_0_47])).
cnf(c_0_69,negated_conjecture,(esk6_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_70,negated_conjecture,(m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_67]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_71,negated_conjecture,(m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_68]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
fof(c_0_72, plain, (![X6]:![X7]:![X8]:![X9]:((((((((m1_subset_1(esk7_4(X6,X7,X8,X9),u1_struct_0(X6))|(~v11_waybel_0(X7,X6)|~r3_waybel_9(X6,X7,X8)))|X8!=X9)|r1_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X9))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6))))))))))&(((((((~v5_pre_topc(k4_waybel_1(X6,esk7_4(X6,X7,X8,X9)),X6,X6)|(~v11_waybel_0(X7,X6)|~r3_waybel_9(X6,X7,X8)))|X8!=X9)|r1_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X9))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l51_waybel_9])])])])])])])])).
cnf(c_0_73,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X3,k10_yellow_6(X1,X2))|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|esk5_2(X1,X2)!=esk6_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_74,negated_conjecture,(esk6_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
fof(c_0_75, plain, (![X3]:![X4]:((((~r2_waybel_9(X3,X4)|r2_yellow_0(X3,k2_relset_1(u1_struct_0(X3),u1_waybel_0(X3,X4))))|~l1_waybel_0(X4,X3))|~l1_orders_2(X3))&(((~r2_yellow_0(X3,k2_relset_1(u1_struct_0(X3),u1_waybel_0(X3,X4)))|r2_waybel_9(X3,X4))|~l1_waybel_0(X4,X3))|~l1_orders_2(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_waybel_9])])])])])])).
fof(c_0_76, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:((((((((m1_subset_1(esk8_4(X6,X7,X8,X9),u1_struct_0(X6))|~r3_waybel_9(X6,X7,X8))|X8!=X9)|(~m1_subset_1(X11,u1_struct_0(X6))|(~r1_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X11)|r1_orders_2(X6,X11,X9))))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6))))))))))&(((((((~v5_pre_topc(k4_waybel_1(X6,esk8_4(X6,X7,X8,X9)),X6,X6)|~r3_waybel_9(X6,X7,X8))|X8!=X9)|(~m1_subset_1(X11,u1_struct_0(X6))|(~r1_lattice3(X6,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,X7)),X11)|r1_orders_2(X6,X11,X9))))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l52_waybel_9])])])])])])])])).
cnf(c_0_77,plain,(v2_struct_0(X2)|r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)|~v11_waybel_0(X2,X1)|~v5_pre_topc(k4_waybel_1(X1,esk7_4(X1,X2,X3,X4)),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_78,negated_conjecture,(r2_hidden(X1,k10_yellow_6(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~r2_waybel_9(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_54]), c_0_55]), c_0_18]), c_0_56]), c_0_24]), c_0_25])]), c_0_57]), c_0_35])).
cnf(c_0_79,plain,(r2_waybel_9(X1,X2)|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~r2_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_80,plain,(v2_struct_0(X2)|r1_orders_2(X1,X5,X4)|m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X5,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_76])).
fof(c_0_81, plain, (![X5]:![X6]:![X8]:![X6]:![X9]:((((m1_subset_1(esk9_2(X5,X6),u1_struct_0(X5))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_lattice3(X5,X6,esk9_2(X5,X6))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X6,X8)|r1_orders_2(X5,X8,esk9_2(X5,X6))))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))&(((((m1_subset_1(esk10_3(X5,X6,X9),u1_struct_0(X5))|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((((r1_lattice3(X5,X6,esk10_3(X5,X6,X9))|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~r1_orders_2(X5,esk10_3(X5,X6,X9),X9)|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])])).
cnf(c_0_82,plain,(r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~v11_waybel_0(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk7_4(X1,X2,X3,X3)),X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_77])).
cnf(c_0_83,plain,(v2_struct_0(X2)|r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|m1_subset_1(esk7_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)|~v11_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_84,negated_conjecture,(r2_hidden(X1,k10_yellow_6(esk1_0,esk2_0))|~r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_32]), c_0_54])])).
cnf(c_0_85,plain,(r1_orders_2(X1,X2,X3)|v2_struct_0(X4)|m1_subset_1(esk8_4(X1,X4,X3,X3),u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X4)),X2)|~r3_waybel_9(X1,X4,X3)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_80])).
cnf(c_0_86,plain,(r2_yellow_0(X1,X2)|r1_lattice3(X1,X2,esk10_3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_87,negated_conjecture,(r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X2)|v2_struct_0(X1)|~r3_waybel_9(esk1_0,X1,X2)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(esk7_4(esk1_0,X1,X2,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_88,plain,(r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|m1_subset_1(esk7_4(X1,X2,X3,X3),u1_struct_0(X1))|~r3_waybel_9(X1,X2,X3)|~v11_waybel_0(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_83])).
cnf(c_0_89,negated_conjecture,(~r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,k1_waybel_9(esk1_0,esk2_0))|~r2_waybel_9(esk1_0,esk2_0)|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_84])).
cnf(c_0_90,negated_conjecture,(r3_waybel_9(esk1_0,X1,esk6_2(esk1_0,X1))|r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,esk2_0))|~r2_waybel_9(esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_71]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_91,negated_conjecture,(r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|m1_subset_1(esk6_2(esk1_0,X1),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,esk2_0))|~r2_waybel_9(esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_71]), c_0_18]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_92,negated_conjecture,(~l1_pre_topc(esk1_0)|~r2_hidden(esk6_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~r2_waybel_9(esk1_0,esk2_0)|~m1_subset_1(esk6_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_66])).
cnf(c_0_93,plain,(r1_orders_2(X1,esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3),X4)|r2_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))|v2_struct_0(X2)|m1_subset_1(esk8_4(X1,X2,X4,X4),u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|~r3_waybel_9(X1,X2,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_27])).
cnf(c_0_94,plain,(r2_yellow_0(X1,X2)|m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_95,negated_conjecture,(r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X2)|v2_struct_0(X1)|~r3_waybel_9(esk1_0,X1,X2)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_96,negated_conjecture,(~r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_47]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_97,negated_conjecture,(k1_waybel_9(esk1_0,X1)=esk6_2(esk1_0,X1)|r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,X1,esk5_2(esk1_0,esk2_0))|~r2_waybel_9(esk1_0,esk2_0)|~v11_waybel_0(X1,esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_90]), c_0_91])).
cnf(c_0_98,negated_conjecture,(~l1_pre_topc(esk1_0)|~r2_hidden(esk5_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))|~r2_waybel_9(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_74]), c_0_71])).
cnf(c_0_99,plain,(r1_orders_2(X1,esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3),X4)|r2_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))|v2_struct_0(X2)|m1_subset_1(esk8_4(X1,X2,X4,X4),u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|~r3_waybel_9(X1,X2,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_27])).
cnf(c_0_100,negated_conjecture,(r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_101,negated_conjecture,(~r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)|~r2_waybel_9(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57]), c_0_61]), c_0_70]), c_0_98])).
cnf(c_0_102,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),X2)|r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|m1_subset_1(esk8_4(esk1_0,esk2_0,X2,X2),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X2)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_54]), c_0_55]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_56]), c_0_22]), c_0_23]), c_0_24]), c_0_25])]), c_0_57])).
cnf(c_0_103,negated_conjecture,(~r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~l1_pre_topc(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_79]), c_0_32]), c_0_54])])).
cnf(c_0_104,plain,(v2_struct_0(X2)|r1_orders_2(X1,X5,X4)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X5,u1_struct_0(X1))|X3!=X4|~r3_waybel_9(X1,X2,X3)|~v5_pre_topc(k4_waybel_1(X1,esk8_4(X1,X2,X3,X4)),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_105,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|m1_subset_1(esk8_4(esk1_0,esk2_0,k1_waybel_9(esk1_0,esk2_0),k1_waybel_9(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_47]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57]), c_0_103])).
cnf(c_0_106,plain,(r1_orders_2(X1,X2,X3)|v2_struct_0(X4)|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X4)),X2)|~r3_waybel_9(X1,X4,X3)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~v5_pre_topc(k4_waybel_1(X1,esk8_4(X1,X4,X3,X3)),X1,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_104])).
cnf(c_0_107,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|m1_subset_1(esk8_4(esk1_0,esk2_0,k1_waybel_9(esk1_0,esk2_0),k1_waybel_9(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_41]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_108,plain,(r1_orders_2(X1,esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3),X4)|r2_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)))|v2_struct_0(X2)|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|~r3_waybel_9(X1,X2,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk8_4(X1,X2,X4,X4)),X1,X1)|~m1_subset_1(esk10_3(X1,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,X2)),X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_86]), c_0_27])).
cnf(c_0_109,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|m1_subset_1(esk8_4(esk1_0,esk2_0,k1_waybel_9(esk1_0,esk2_0),k1_waybel_9(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_44]), c_0_17])])).
cnf(c_0_110,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X2),X3)|r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)))|v2_struct_0(X1)|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X2)|~r3_waybel_9(esk1_0,X1,X3)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,X1)),X2),u1_struct_0(esk1_0))|~m1_subset_1(esk8_4(esk1_0,X1,X3,X3),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_111,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|m1_subset_1(esk8_4(esk1_0,esk2_0,k1_waybel_9(esk1_0,esk2_0),k1_waybel_9(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_47]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_112,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,k1_waybel_9(esk1_0,esk2_0))|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_54]), c_0_55]), c_0_56])]), c_0_57]), c_0_103])).
cnf(c_0_113,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~l1_pre_topc(esk1_0)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_47]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_114,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~l1_pre_topc(esk1_0)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_41]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_115,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~m1_subset_1(esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_44]), c_0_17])])).
cnf(c_0_116,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_94]), c_0_32]), c_0_21])])).
cnf(c_0_117,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),X1),k1_waybel_9(esk1_0,esk2_0))|r2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_116, c_0_100])).
cnf(c_0_118,plain,(r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk10_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_119,negated_conjecture,(r1_orders_2(esk1_0,esk10_3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_47]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57]), c_0_103])).
cnf(c_0_120,negated_conjecture,(~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)),k1_waybel_9(esk1_0,esk2_0))|~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_32]), c_0_21])]), c_0_103])).
cnf(c_0_121,negated_conjecture,(~l1_pre_topc(esk1_0)|~r3_waybel_9(esk1_0,esk2_0,k1_waybel_9(esk1_0,esk2_0))|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_120, c_0_100])).
cnf(c_0_122,negated_conjecture,(~l1_pre_topc(esk1_0)|~m1_subset_1(k1_waybel_9(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_47]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_123,negated_conjecture,(~l1_pre_topc(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_41]), c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_124,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_44]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 125
# Proof object clause steps            : 104
# Proof object formula steps           : 21
# Proof object conjectures             : 80
# Proof object clause conjectures      : 77
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 274
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 40
# Processed clauses                    : 338
# ...of these trivial                  : 0
# ...subsumed                          : 81
# ...remaining for further processing  : 257
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 99
# Backward-rewritten                   : 1
# Generated clauses                    : 433
# ...of the previous two non-trivial   : 418
# Contextual simplify-reflections      : 224
# Paramodulations                      : 429
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 113
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 96
# Current number of unprocessed clauses: 18
# ...number of literals in the above   : 218
# Current number of archived formulas  : 0
# Current number of archived clauses   : 140
# Clause-clause subsumption calls (NU) : 29863
# Rec. Clause-clause subsumption calls : 1395
# Non-unit clause-clause subsumptions  : 404
# Unit Clause-clause subsumption calls : 44
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 33607

# -------------------------------------------------
# User time                : 0.040 s
# System time              : 0.000 s
# Total time               : 0.040 s
# Maximum resident set size: 43276 pages
