# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_pre_topc, conjecture, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))), file('out/t39_pre_topc__7_premises__rerun.E_input', t39_pre_topc)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t39_pre_topc__7_premises__rerun.E_input', t3_subset)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t39_pre_topc__7_premises__rerun.E_input', t12_xboole_1)).
fof(d9_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2))))))))))), file('out/t39_pre_topc__7_premises__rerun.E_input', d9_pre_topc)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t39_pre_topc__7_premises__rerun.E_input', dt_m1_pre_topc)).
fof(t11_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3))), file('out/t39_pre_topc__7_premises__rerun.E_input', t11_xboole_1)).
fof(d3_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1))), file('out/t39_pre_topc__7_premises__rerun.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t39_pre_topc__7_premises__rerun.E_input', dt_l1_pre_topc)).
fof(c_0_8, negated_conjecture, (~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))))), inference(assume_negation,[status(cth)],[t39_pre_topc])).
fof(c_0_9, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_10, negated_conjecture, ((l1_pre_topc(esk1_0)&(m1_pre_topc(esk2_0,esk1_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
fof(c_0_11, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_12,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X9]:![X11]:(((((r1_tarski(k2_struct_0(X6),k2_struct_0(X5))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&(((((((m1_subset_1(esk4_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(X5)))|~r2_hidden(X7,u1_pre_topc(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&((((((r2_hidden(esk4_3(X5,X6,X7),u1_pre_topc(X5))|~r2_hidden(X7,u1_pre_topc(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&(((((X7=k9_subset_1(u1_struct_0(X6),esk4_3(X5,X6,X7),k2_struct_0(X6))|~r2_hidden(X7,u1_pre_topc(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))))&((((((~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X5)))|(~r2_hidden(X9,u1_pre_topc(X5))|X7!=k9_subset_1(u1_struct_0(X6),X9,k2_struct_0(X6))))|r2_hidden(X7,u1_pre_topc(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))))&(((((m1_subset_1(esk5_2(X5,X6),k1_zfmisc_1(u1_struct_0(X6)))|~r1_tarski(k2_struct_0(X6),k2_struct_0(X5)))|m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&((((((~r2_hidden(esk5_2(X5,X6),u1_pre_topc(X6))|(~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X5)))|(~r2_hidden(X11,u1_pre_topc(X5))|esk5_2(X5,X6)!=k9_subset_1(u1_struct_0(X6),X11,k2_struct_0(X6)))))|~r1_tarski(k2_struct_0(X6),k2_struct_0(X5)))|m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&((((((m1_subset_1(esk6_2(X5,X6),k1_zfmisc_1(u1_struct_0(X5)))|r2_hidden(esk5_2(X5,X6),u1_pre_topc(X6)))|~r1_tarski(k2_struct_0(X6),k2_struct_0(X5)))|m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&((((((r2_hidden(esk6_2(X5,X6),u1_pre_topc(X5))|r2_hidden(esk5_2(X5,X6),u1_pre_topc(X6)))|~r1_tarski(k2_struct_0(X6),k2_struct_0(X5)))|m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&(((((esk5_2(X5,X6)=k9_subset_1(u1_struct_0(X6),esk6_2(X5,X6),k2_struct_0(X6))|r2_hidden(esk5_2(X5,X6),u1_pre_topc(X6)))|~r1_tarski(k2_struct_0(X6),k2_struct_0(X5)))|m1_pre_topc(X6,X5))|~l1_pre_topc(X6))|~l1_pre_topc(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])])])).
fof(c_0_15, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(~r1_tarski(k2_xboole_0(X4,X5),X6)|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
cnf(c_0_17,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(r1_tarski(esk3_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19,plain,(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, (![X2]:(~l1_struct_0(X2)|k2_struct_0(X2)=u1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_22,plain,(r1_tarski(X1,X2)|~r1_tarski(k2_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,negated_conjecture,(k2_xboole_0(esk3_0,u1_struct_0(esk2_0))=u1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24,plain,(r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26,negated_conjecture,(r1_tarski(esk3_0,X1)|~r1_tarski(u1_struct_0(esk2_0),X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27,plain,(r1_tarski(u1_struct_0(X1),k2_struct_0(X2))|~l1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_28, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_29,negated_conjecture,(~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_30,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31,negated_conjecture,(r1_tarski(esk3_0,k2_struct_0(X1))|~l1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33,negated_conjecture,(~r1_tarski(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34,negated_conjecture,(r1_tarski(esk3_0,u1_struct_0(X1))|~l1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_25]), c_0_32])).
cnf(c_0_35,negated_conjecture,(m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_37,negated_conjecture,(~l1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_38,negated_conjecture,(l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_35]), c_0_36])])).
cnf(c_0_39,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_32]), c_0_38])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 23
# Proof object formula steps           : 17
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 34
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 34
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 37
# ...of the previous two non-trivial   : 32
# Contextual simplify-reflections      : 7
# Paramodulations                      : 36
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 108
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 423
# Rec. Clause-clause subsumption calls : 85
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 69
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2540

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.008 s
# Total time               : 0.012 s
# Maximum resident set size: 45328 pages
