# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_lattices, conjecture, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k3_lattices(X1,k5_lattices(X1),X2)=X2))), file('out/t39_lattices__7_premises__rerun.E_input', t39_lattices)).
fof(redefinition_k3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k3_lattices(X1,X2,X3)=k1_lattices(X1,X2,X3))), file('out/t39_lattices__7_premises__rerun.E_input', redefinition_k3_lattices)).
fof(d3_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&l2_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k1_lattices(X1,X2,X3)=X3))))), file('out/t39_lattices__7_premises__rerun.E_input', d3_lattices)).
fof(t21_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2))))), file('out/t39_lattices__7_premises__rerun.E_input', t21_lattices)).
fof(d16_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))))), file('out/t39_lattices__7_premises__rerun.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1)))), file('out/t39_lattices__7_premises__rerun.E_input', dt_k5_lattices)).
fof(cc1_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1)))))))))), file('out/t39_lattices__7_premises__rerun.E_input', cc1_lattices)).
fof(dt_l3_lattices, axiom, (![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1)))), file('out/t39_lattices__7_premises__rerun.E_input', dt_l3_lattices)).
fof(c_0_8, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k3_lattices(X1,k5_lattices(X1),X2)=X2)))), inference(assume_negation,[status(cth)],[t39_lattices])).
fof(c_0_9, negated_conjecture, (((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v13_lattices(esk1_0)&l3_lattices(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&k3_lattices(esk1_0,k5_lattices(esk1_0),esk2_0)!=esk2_0))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v4_lattices(X4)|~l2_lattices(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|k3_lattices(X4,X5,X6)=k1_lattices(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_lattices])])])).
cnf(c_0_11,negated_conjecture,(k3_lattices(esk1_0,k5_lattices(esk1_0),esk2_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(k3_lattices(X1,X2,X3)=k1_lattices(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)|~v4_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(((((~r1_lattices(X4,X5,X6)|k1_lattices(X4,X5,X6)=X6)|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l2_lattices(X4)))&((((k1_lattices(X4,X5,X6)!=X6|r1_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l2_lattices(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattices])])])])])])])).
cnf(c_0_16,negated_conjecture,(k1_lattices(esk1_0,k5_lattices(esk1_0),esk2_0)!=esk2_0|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])]), c_0_14])).
cnf(c_0_17,plain,(v2_struct_0(X1)|k1_lattices(X1,X2,X3)=X3|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:(((((~r1_lattices(X4,X5,X6)|k2_lattices(X4,X5,X6)=X5)|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))&((((k2_lattices(X4,X5,X6)!=X5|r1_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])])])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:(((((((k2_lattices(X4,X5,X6)=X5|~m1_subset_1(X6,u1_struct_0(X4)))|X5!=k5_lattices(X4))|~m1_subset_1(X5,u1_struct_0(X4)))|~v13_lattices(X4))|(v2_struct_0(X4)|~l1_lattices(X4)))&(((((k2_lattices(X4,X6,X5)=X5|~m1_subset_1(X6,u1_struct_0(X4)))|X5!=k5_lattices(X4))|~m1_subset_1(X5,u1_struct_0(X4)))|~v13_lattices(X4))|(v2_struct_0(X4)|~l1_lattices(X4))))&(((((m1_subset_1(esk3_2(X4,X5),u1_struct_0(X4))|X5=k5_lattices(X4))|~m1_subset_1(X5,u1_struct_0(X4)))|~v13_lattices(X4))|(v2_struct_0(X4)|~l1_lattices(X4)))&(((((k2_lattices(X4,X5,esk3_2(X4,X5))!=X5|k2_lattices(X4,esk3_2(X4,X5),X5)!=X5)|X5=k5_lattices(X4))|~m1_subset_1(X5,u1_struct_0(X4)))|~v13_lattices(X4))|(v2_struct_0(X4)|~l1_lattices(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])])])).
fof(c_0_20, plain, (![X2]:((v2_struct_0(X2)|~l1_lattices(X2))|m1_subset_1(k5_lattices(X2),u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
cnf(c_0_21,negated_conjecture,(~r1_lattices(esk1_0,k5_lattices(esk1_0),esk2_0)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_13])]), c_0_14])).
cnf(c_0_22,plain,(v2_struct_0(X1)|r1_lattices(X1,X2,X3)|~l3_lattices(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|k2_lattices(X1,X2,X3)!=X2), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23,negated_conjecture,(l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24,plain,(v2_struct_0(X1)|k2_lattices(X1,X2,X3)=X2|~l1_lattices(X1)|~v13_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X2!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,plain,(m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,negated_conjecture,(k2_lattices(esk1_0,k5_lattices(esk1_0),esk2_0)!=k5_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_13]), c_0_23])]), c_0_14])).
cnf(c_0_27,plain,(k2_lattices(X1,k5_lattices(X1),X2)=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_24]), c_0_25])).
cnf(c_0_28,negated_conjecture,(v13_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_29, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v4_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v5_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v6_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v7_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v8_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&((v9_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_30,negated_conjecture,(~l1_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_13]), c_0_28])]), c_0_14])).
cnf(c_0_31,plain,(v2_struct_0(X1)|v4_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_32,negated_conjecture,(v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_33,negated_conjecture,(~l1_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~l2_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_23]), c_0_32])]), c_0_14])).
cnf(c_0_34,plain,(v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35,negated_conjecture,(~l1_lattices(esk1_0)|~v8_lattices(esk1_0)|~l2_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_23]), c_0_32])]), c_0_14])).
cnf(c_0_36,plain,(v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,negated_conjecture,(~l1_lattices(esk1_0)|~l2_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_23]), c_0_32])]), c_0_14])).
fof(c_0_38, plain, (![X2]:((l1_lattices(X2)|~l3_lattices(X2))&(l2_lattices(X2)|~l3_lattices(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_39,negated_conjecture,(~l1_lattices(esk1_0)|~l2_lattices(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_25]), c_0_14])).
cnf(c_0_40,plain,(l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_41,negated_conjecture,(~l1_lattices(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_23])])).
cnf(c_0_42,plain,(l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_23])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 27
# Proof object formula steps           : 17
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 59
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 59
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 14
# ...of the previous two non-trivial   : 11
# Contextual simplify-reflections      : 2
# Paramodulations                      : 12
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 404
# Rec. Clause-clause subsumption calls : 74
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2564

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45280 pages
