# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', t3_subset)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', t2_xboole_1)).
fof(d12_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>![X3]:((v6_orders_2(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>(m2_orders_2(X3,X1,X2)<=>(~(X3=k1_xboole_0)&(r2_wellord1(u1_orders_2(X1),X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>k1_funct_1(X2,k1_orders_2(X1,k3_orders_2(X1,X3,X4)))=X4))))))))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', d12_orders_2)).
fof(dt_m2_orders_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_orders_1(X2,k1_orders_1(u1_struct_0(X1))))=>![X3]:(m2_orders_2(X3,X1,X2)=>(v6_orders_2(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', dt_m2_orders_2)).
fof(t39_orders_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_orders_1(X3,k1_orders_1(u1_struct_0(X1)))=>![X4]:(m2_orders_2(X4,X1,X3)=>(X2=k1_funct_1(X3,u1_struct_0(X1))=>k3_orders_2(X1,X4,X2)=k1_xboole_0)))))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', t39_orders_2)).
fof(t34_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((r1_tarski(X2,X3)&(r2_wellord1(u1_orders_2(X1),X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r2_hidden(X5,X2)&r2_orders_2(X1,X4,X5))=>r2_hidden(X4,X2))))&(~(X2=X3)&~(m1_orders_2(X2,X1,X3))))))))))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', t34_orders_2)).
fof(t5_subset, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3))))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', t5_subset)).
fof(d11_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((~(X2=k1_xboole_0)=>(m1_orders_2(X3,X1,X2)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r2_hidden(X4,X2)&X3=k3_orders_2(X1,X2,X4)))))&(X2=k1_xboole_0=>(m1_orders_2(X3,X1,X2)<=>X3=k1_xboole_0))))))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', d11_orders_2)).
fof(dt_m1_orders_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>![X3]:(m1_orders_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', dt_m1_orders_2)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', fc1_xboole_0)).
fof(d2_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k1_struct_0(X1)=k1_xboole_0)), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', d2_struct_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', dt_l1_orders_2)).
fof(t14_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k1_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1))), file('out/t39_orders_2__13_premises__rerun__rerun.E_input', t14_orders_2)).
fof(c_0_13, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_14, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:![X8]:((((((X7!=k1_xboole_0|~m2_orders_2(X7,X5,X6))|(~v6_orders_2(X7,X5)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))))|~m1_orders_1(X6,k1_orders_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(((((r2_wellord1(u1_orders_2(X5),X7)|~m2_orders_2(X7,X5,X6))|(~v6_orders_2(X7,X5)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))))|~m1_orders_1(X6,k1_orders_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(((((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|k1_funct_1(X6,k1_orders_2(X5,k3_orders_2(X5,X7,X8)))=X8))|~m2_orders_2(X7,X5,X6))|(~v6_orders_2(X7,X5)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))))|~m1_orders_1(X6,k1_orders_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))))&(((((((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r2_wellord1(u1_orders_2(X5),X7))|X7=k1_xboole_0)|m2_orders_2(X7,X5,X6))|(~v6_orders_2(X7,X5)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))))|~m1_orders_1(X6,k1_orders_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(((((((r2_hidden(esk1_3(X5,X6,X7),X7)|~r2_wellord1(u1_orders_2(X5),X7))|X7=k1_xboole_0)|m2_orders_2(X7,X5,X6))|(~v6_orders_2(X7,X5)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))))|~m1_orders_1(X6,k1_orders_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&((((((k1_funct_1(X6,k1_orders_2(X5,k3_orders_2(X5,X7,esk1_3(X5,X6,X7))))!=esk1_3(X5,X6,X7)|~r2_wellord1(u1_orders_2(X5),X7))|X7=k1_xboole_0)|m2_orders_2(X7,X5,X6))|(~v6_orders_2(X7,X5)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))))|~m1_orders_1(X6,k1_orders_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_orders_2])])])])])])])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(((v6_orders_2(X6,X4)|~m2_orders_2(X6,X4,X5))|((v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|~l1_orders_2(X4)))))|~m1_orders_1(X5,k1_orders_1(u1_struct_0(X4)))))&((m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|~m2_orders_2(X6,X4,X5))|((v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|~l1_orders_2(X4)))))|~m1_orders_1(X5,k1_orders_1(u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_orders_2])])])])])])])).
fof(c_0_17, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_orders_1(X3,k1_orders_1(u1_struct_0(X1)))=>![X4]:(m2_orders_2(X4,X1,X3)=>(X2=k1_funct_1(X3,u1_struct_0(X1))=>k3_orders_2(X1,X4,X2)=k1_xboole_0))))))), inference(assume_negation,[status(cth)],[t39_orders_2])).
fof(c_0_18, plain, (![X6]:![X7]:![X8]:(((((((m1_subset_1(esk6_3(X6,X7,X8),u1_struct_0(X6))|(X7=X8|m1_orders_2(X7,X6,X8)))|~r2_wellord1(u1_orders_2(X6),X8))|~r1_tarski(X7,X8))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|~l1_orders_2(X6))))))&(((((((m1_subset_1(esk7_3(X6,X7,X8),u1_struct_0(X6))|(X7=X8|m1_orders_2(X7,X6,X8)))|~r2_wellord1(u1_orders_2(X6),X8))|~r1_tarski(X7,X8))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|~l1_orders_2(X6))))))&((((((((r2_hidden(esk7_3(X6,X7,X8),X7)|(X7=X8|m1_orders_2(X7,X6,X8)))|~r2_wellord1(u1_orders_2(X6),X8))|~r1_tarski(X7,X8))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|~l1_orders_2(X6))))))&((((((r2_orders_2(X6,esk6_3(X6,X7,X8),esk7_3(X6,X7,X8))|(X7=X8|m1_orders_2(X7,X6,X8)))|~r2_wellord1(u1_orders_2(X6),X8))|~r1_tarski(X7,X8))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|~l1_orders_2(X6)))))))&((((((~r2_hidden(esk6_3(X6,X7,X8),X7)|(X7=X8|m1_orders_2(X7,X6,X8)))|~r2_wellord1(u1_orders_2(X6),X8))|~r1_tarski(X7,X8))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|~l1_orders_2(X6)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_orders_2])])])])])])])])).
cnf(c_0_19,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(v2_struct_0(X1)|r2_wellord1(u1_orders_2(X1),X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v6_orders_2(X3,X1)|~m2_orders_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(v2_struct_0(X2)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_orders_1(X1,k1_orders_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m2_orders_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(v2_struct_0(X2)|v6_orders_2(X3,X2)|~m1_orders_1(X1,k1_orders_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m2_orders_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, negated_conjecture, (((~v2_struct_0(esk2_0)&(v3_orders_2(esk2_0)&(v4_orders_2(esk2_0)&(v5_orders_2(esk2_0)&l1_orders_2(esk2_0)))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_orders_1(esk4_0,k1_orders_1(u1_struct_0(esk2_0)))&(m2_orders_2(esk5_0,esk2_0,esk4_0)&(esk3_0=k1_funct_1(esk4_0,u1_struct_0(esk2_0))&k3_orders_2(esk2_0,esk5_0,esk3_0)!=k1_xboole_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
cnf(c_0_25,plain,(v2_struct_0(X1)|m1_orders_2(X2,X1,X3)|X2=X3|r2_hidden(esk7_3(X1,X2,X3),X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(X2,X3)|~r2_wellord1(u1_orders_2(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27,plain,(r2_wellord1(u1_orders_2(X1),X2)|v2_struct_0(X1)|~m2_orders_2(X2,X1,X3)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_28,negated_conjecture,(m2_orders_2(esk5_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29,negated_conjecture,(m1_orders_1(esk4_0,k1_orders_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,negated_conjecture,(v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,negated_conjecture,(v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,negated_conjecture,(v3_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_35, plain, (![X4]:![X5]:![X6]:(~r2_hidden(X4,X5)|(~m1_subset_1(X5,k1_zfmisc_1(X6))|~v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])])])).
cnf(c_0_36,plain,(k1_xboole_0=X1|m1_orders_2(k1_xboole_0,X2,X1)|r2_hidden(esk7_3(X2,k1_xboole_0,X1),k1_xboole_0)|v2_struct_0(X2)|~r2_wellord1(u1_orders_2(X2),X1)|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_20]), c_0_26])])).
cnf(c_0_37,negated_conjecture,(r2_wellord1(u1_orders_2(esk2_0),esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_38,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34])).
fof(c_0_39, plain, (![X5]:![X6]:![X7]:![X9]:((((((((m1_subset_1(esk8_3(X5,X6,X7),u1_struct_0(X5))|~m1_orders_2(X7,X5,X6))|X6=k1_xboole_0)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&((((((r2_hidden(esk8_3(X5,X6,X7),X6)|~m1_orders_2(X7,X5,X6))|X6=k1_xboole_0)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(((((X7=k3_orders_2(X5,X6,esk8_3(X5,X6,X7))|~m1_orders_2(X7,X5,X6))|X6=k1_xboole_0)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))))&((((((~m1_subset_1(X9,u1_struct_0(X5))|(~r2_hidden(X9,X6)|X7!=k3_orders_2(X5,X6,X9)))|m1_orders_2(X7,X5,X6))|X6=k1_xboole_0)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))&((((((~m1_orders_2(X7,X5,X6)|X7=k1_xboole_0)|X6!=k1_xboole_0)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(((((X7!=k1_xboole_0|m1_orders_2(X7,X5,X6))|X6!=k1_xboole_0)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_orders_2])])])])])])])])).
fof(c_0_40, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|~l1_orders_2(X4)))))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(~m1_orders_2(X6,X4,X5)|m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_orders_2])])])])])])).
cnf(c_0_41,plain,(~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42,negated_conjecture,(esk5_0=k1_xboole_0|m1_orders_2(k1_xboole_0,esk2_0,esk5_0)|r2_hidden(esk7_3(esk2_0,k1_xboole_0,esk5_0),k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_38])]), c_0_34])).
cnf(c_0_43,plain,(v2_struct_0(X1)|X2=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_orders_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44,plain,(m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m1_orders_2(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45,negated_conjecture,(esk5_0=k1_xboole_0|m1_orders_2(k1_xboole_0,esk2_0,esk5_0)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_26])])).
cnf(c_0_46,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_47,plain,(v2_struct_0(X1)|k1_funct_1(X2,k1_orders_2(X1,k3_orders_2(X1,X3,X4)))=X4|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v6_orders_2(X3,X1)|~m2_orders_2(X3,X1,X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_48,plain,(X1=k1_xboole_0|r2_hidden(esk8_3(X2,X1,X3),X1)|v2_struct_0(X2)|~m1_orders_2(X3,X2,X1)|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(csr,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49,negated_conjecture,(esk5_0=k1_xboole_0|m1_orders_2(k1_xboole_0,esk2_0,esk5_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50,plain,(v2_struct_0(X1)|X2=k1_xboole_0|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_orders_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51,plain,(k1_funct_1(X1,k1_orders_2(X2,k3_orders_2(X2,X3,X4)))=X4|v2_struct_0(X2)|~r2_hidden(X4,X3)|~m2_orders_2(X3,X2,X1)|~m1_orders_1(X1,k1_orders_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_47, c_0_22]), c_0_23])).
cnf(c_0_52,negated_conjecture,(esk5_0=k1_xboole_0|r2_hidden(esk8_3(esk2_0,esk5_0,k1_xboole_0),esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_38])]), c_0_34])).
cnf(c_0_53,plain,(X1=k1_xboole_0|v2_struct_0(X2)|m1_subset_1(esk8_3(X2,X1,X3),u1_struct_0(X2))|~m1_orders_2(X3,X2,X1)|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(csr,[status(thm)],[c_0_50, c_0_44])).
cnf(c_0_54,plain,(v2_struct_0(X1)|X2=k1_xboole_0|X3=k3_orders_2(X1,X2,esk8_3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_orders_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_55,negated_conjecture,(k1_funct_1(X1,k1_orders_2(X2,k3_orders_2(X2,esk5_0,esk8_3(esk2_0,esk5_0,k1_xboole_0))))=esk8_3(esk2_0,esk5_0,k1_xboole_0)|esk5_0=k1_xboole_0|v2_struct_0(X2)|~m2_orders_2(esk5_0,X2,X1)|~m1_orders_1(X1,k1_orders_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk8_3(esk2_0,esk5_0,k1_xboole_0),u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56,negated_conjecture,(esk5_0=k1_xboole_0|m1_subset_1(esk8_3(esk2_0,esk5_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_49]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_38])]), c_0_34])).
cnf(c_0_57,plain,(k3_orders_2(X1,X2,esk8_3(X1,X2,X3))=X3|X2=k1_xboole_0|v2_struct_0(X1)|~m1_orders_2(X3,X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[c_0_54, c_0_44])).
fof(c_0_58, plain, (![X2]:(~l1_struct_0(X2)|k1_struct_0(X2)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_struct_0])])).
fof(c_0_59, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_60,negated_conjecture,(k1_funct_1(esk4_0,k1_orders_2(esk2_0,k3_orders_2(esk2_0,esk5_0,esk8_3(esk2_0,esk5_0,k1_xboole_0))))=esk8_3(esk2_0,esk5_0,k1_xboole_0)|esk5_0=k1_xboole_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34]), c_0_56])).
cnf(c_0_61,negated_conjecture,(k3_orders_2(esk2_0,esk5_0,esk8_3(esk2_0,esk5_0,k1_xboole_0))=k1_xboole_0|esk5_0=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_49]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_38])]), c_0_34])).
fof(c_0_62, plain, (![X2]:((v2_struct_0(X2)|(~v3_orders_2(X2)|(~v4_orders_2(X2)|(~v5_orders_2(X2)|~l1_orders_2(X2)))))|k1_orders_2(X2,k1_struct_0(X2))=u1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_orders_2])])])).
cnf(c_0_63,plain,(k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65,negated_conjecture,(esk8_3(esk2_0,esk5_0,k1_xboole_0)=k1_funct_1(esk4_0,k1_orders_2(esk2_0,k1_xboole_0))|esk5_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66,plain,(k1_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_67,plain,(k1_struct_0(X1)=k1_xboole_0|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68,negated_conjecture,(k3_orders_2(esk2_0,esk5_0,k1_funct_1(esk4_0,k1_orders_2(esk2_0,k1_xboole_0)))=k1_xboole_0|esk5_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_65])).
cnf(c_0_69,plain,(k1_orders_2(X1,k1_xboole_0)=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_70,negated_conjecture,(esk3_0=k1_funct_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_71,negated_conjecture,(k3_orders_2(esk2_0,esk5_0,esk3_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_72,plain,(v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v6_orders_2(X3,X1)|~m2_orders_2(X3,X1,X2)|X3!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_73,negated_conjecture,(esk5_0=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_71]), c_0_34])).
cnf(c_0_74,plain,(v2_struct_0(X1)|~m2_orders_2(k1_xboole_0,X1,X2)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_72, c_0_22]), c_0_23])])).
cnf(c_0_75,negated_conjecture,(m2_orders_2(k1_xboole_0,esk2_0,esk4_0)), inference(rw,[status(thm)],[c_0_28, c_0_73])).
cnf(c_0_76,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 77
# Proof object clause steps            : 51
# Proof object formula steps           : 26
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 80
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 115
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 113
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 24
# Generated clauses                    : 63
# ...of the previous two non-trivial   : 59
# Contextual simplify-reflections      : 13
# Paramodulations                      : 59
# Factorizations                       : 0
# Equation resolutions                 : 5
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 8
# ...number of literals in the above   : 46
# Current number of archived formulas  : 0
# Current number of archived clauses   : 65
# Clause-clause subsumption calls (NU) : 3127
# Rec. Clause-clause subsumption calls : 97
# Non-unit clause-clause subsumptions  : 17
# Unit Clause-clause subsumption calls : 24
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7509

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 47428 pages
