# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X1))&X3=k1_funct_1(X1,X4)))))), file('out/t3_funct_2__5_premises__rerun.E_input', d3_funct_1)).
fof(t3_funct_2, conjecture, (![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k9_xtuple_0(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))))), file('out/t3_funct_2__5_premises__rerun.E_input', t3_funct_2)).
fof(t4_relset_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>((r1_tarski(k9_xtuple_0(X3),X1)&r1_tarski(k10_xtuple_0(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('out/t3_funct_2__5_premises__rerun.E_input', t4_relset_1)).
fof(reflexivity_r1_tarski, axiom, (![X1]:![X2]:r1_tarski(X1,X1)), file('out/t3_funct_2__5_premises__rerun.E_input', reflexivity_r1_tarski)).
fof(t2_funct_2, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k10_xtuple_0(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k9_xtuple_0(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(X2),X1)))))))), file('out/t3_funct_2__5_premises__rerun.E_input', t2_funct_2)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t3_funct_2__5_premises__rerun.E_input', d3_tarski)).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X6]:![X11]:((((((r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X5))|~r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((X7=k1_funct_1(X5,esk1_3(X5,X6,X7))|~r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(X9,k9_xtuple_0(X5))|X7!=k1_funct_1(X5,X9))|r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(esk2_2(X5,X6),X6)|(~r2_hidden(X11,k9_xtuple_0(X5))|esk2_2(X5,X6)!=k1_funct_1(X5,X11)))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&((((r2_hidden(esk3_2(X5,X6),k9_xtuple_0(X5))|r2_hidden(esk2_2(X5,X6),X6))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((esk2_2(X5,X6)=k1_funct_1(X5,esk3_2(X5,X6))|r2_hidden(esk2_2(X5,X6),X6))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_1])])])])])])])).
fof(c_0_7, negated_conjecture, (~(![X1]:![X2]:![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k9_xtuple_0(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))))), inference(assume_negation,[status(cth)],[t3_funct_2])).
cnf(c_0_8,plain,(r2_hidden(esk1_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, negated_conjecture, (![X8]:((v1_relat_1(esk6_0)&v1_funct_1(esk6_0))&((k9_xtuple_0(esk6_0)=esk4_0&(~r2_hidden(X8,esk4_0)|r2_hidden(k1_funct_1(esk6_0,X8),esk5_0)))&(~v1_funct_1(esk6_0)|(~v1_funct_2(esk6_0,esk4_0,esk5_0)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
cnf(c_0_10,plain,(X3=k1_funct_1(X1,esk1_3(X1,X2,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(r2_hidden(esk1_3(X1,k10_xtuple_0(X1),X2),k9_xtuple_0(X1))|~r2_hidden(X2,k10_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_12,negated_conjecture,(k9_xtuple_0(esk6_0)=esk4_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,negated_conjecture,(v1_relat_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,negated_conjecture,(~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))|~v1_funct_2(esk6_0,esk4_0,esk5_0)|~v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X6)|((~r1_tarski(k9_xtuple_0(X6),X4)|~r1_tarski(k10_xtuple_0(X6),X5))|m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_relset_1])])).
fof(c_0_17, plain, (![X3]:r1_tarski(X3,X3)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[reflexivity_r1_tarski])])).
fof(c_0_18, plain, (![X3]:![X4]:(((v1_funct_1(X4)|~r1_tarski(k10_xtuple_0(X4),X3))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&(((v1_funct_2(X4,k9_xtuple_0(X4),X3)|~r1_tarski(k10_xtuple_0(X4),X3))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&((m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(X4),X3)))|~r1_tarski(k10_xtuple_0(X4),X3))|(~v1_relat_1(X4)|~v1_funct_1(X4)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_funct_2])])])).
cnf(c_0_19,negated_conjecture,(r2_hidden(k1_funct_1(esk6_0,X1),esk5_0)|~r2_hidden(X1,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20,plain,(k1_funct_1(X1,esk1_3(X1,k10_xtuple_0(X1),X2))=X2|~r2_hidden(X2,k10_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_21,negated_conjecture,(r2_hidden(esk1_3(esk6_0,k10_xtuple_0(esk6_0),X1),esk4_0)|~r2_hidden(X1,k10_xtuple_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk7_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk7_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_23,negated_conjecture,(~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0)))|~v1_funct_2(esk6_0,esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_13])])).
cnf(c_0_24,plain,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(k10_xtuple_0(X1),X3)|~r1_tarski(k9_xtuple_0(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,plain,(r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26,plain,(v1_funct_2(X1,k9_xtuple_0(X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_tarski(k10_xtuple_0(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,negated_conjecture,(r2_hidden(X1,esk5_0)|~r2_hidden(X1,k10_xtuple_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_13]), c_0_14])]), c_0_21])).
cnf(c_0_28,plain,(r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,negated_conjecture,(~r1_tarski(k10_xtuple_0(esk6_0),esk5_0)|~v1_funct_2(esk6_0,esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_12]), c_0_25]), c_0_14])])).
cnf(c_0_30,negated_conjecture,(v1_funct_2(esk6_0,esk4_0,X1)|~r1_tarski(k10_xtuple_0(esk6_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_31,plain,(r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,negated_conjecture,(r1_tarski(k10_xtuple_0(esk6_0),X1)|r2_hidden(esk7_2(k10_xtuple_0(esk6_0),X1),esk5_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33,negated_conjecture,(~r1_tarski(k10_xtuple_0(esk6_0),esk5_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 22
# Proof object formula steps           : 13
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 18
# Processed clauses                    : 48
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 2
# Paramodulations                      : 15
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 6
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 97
# Rec. Clause-clause subsumption calls : 38
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2010

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45284 pages
