# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_waybel_7, conjecture, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_yellow_0(X1)&l1_orders_2(X1))))=>(v7_struct_0(X1)<=>k4_yellow_0(X1)=k3_yellow_0(X1)))), file('out/t3_waybel_7__7_premises__rerun.E_input', t3_waybel_7)).
fof(t45_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,k4_yellow_0(X1))))), file('out/t3_waybel_7__7_premises__rerun.E_input', t45_yellow_0)).
fof(d10_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>(v7_struct_0(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>X2=X3))))), file('out/t3_waybel_7__7_premises__rerun.E_input', d10_struct_0)).
fof(dt_k4_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>m1_subset_1(k4_yellow_0(X1),u1_struct_0(X1)))), file('out/t3_waybel_7__7_premises__rerun.E_input', dt_k4_yellow_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t3_waybel_7__7_premises__rerun.E_input', dt_l1_orders_2)).
fof(t19_yellow_5, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_orders_2(X1,X2,k3_yellow_0(X1))=>X2=k3_yellow_0(X1))))), file('out/t3_waybel_7__7_premises__rerun.E_input', t19_yellow_5)).
fof(dt_k3_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1)))), file('out/t3_waybel_7__7_premises__rerun.E_input', dt_k3_yellow_0)).
fof(cc4_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>(v3_yellow_0(X1)=>(v1_yellow_0(X1)&v2_yellow_0(X1))))), file('out/t3_waybel_7__7_premises__rerun.E_input', cc4_yellow_0)).
fof(c_0_8, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_yellow_0(X1)&l1_orders_2(X1))))=>(v7_struct_0(X1)<=>k4_yellow_0(X1)=k3_yellow_0(X1))))), inference(assume_negation,[status(cth)],[t3_waybel_7])).
fof(c_0_9, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v5_orders_2(X3)|(~v2_yellow_0(X3)|~l1_orders_2(X3))))|(~m1_subset_1(X4,u1_struct_0(X3))|r1_orders_2(X3,X4,k4_yellow_0(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t45_yellow_0])])])])])])).
fof(c_0_10, negated_conjecture, (((~v2_struct_0(esk3_0)&(v5_orders_2(esk3_0)&(v3_yellow_0(esk3_0)&l1_orders_2(esk3_0))))&((~v7_struct_0(esk3_0)|k4_yellow_0(esk3_0)!=k3_yellow_0(esk3_0))&(v7_struct_0(esk3_0)|k4_yellow_0(esk3_0)=k3_yellow_0(esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:(((~v7_struct_0(X4)|(~m1_subset_1(X5,u1_struct_0(X4))|(~m1_subset_1(X6,u1_struct_0(X4))|X5=X6)))|~l1_struct_0(X4))&(((m1_subset_1(esk1_1(X4),u1_struct_0(X4))|v7_struct_0(X4))|~l1_struct_0(X4))&(((m1_subset_1(esk2_1(X4),u1_struct_0(X4))|v7_struct_0(X4))|~l1_struct_0(X4))&((esk1_1(X4)!=esk2_1(X4)|v7_struct_0(X4))|~l1_struct_0(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_struct_0])])])])])])])).
fof(c_0_12, plain, (![X2]:(~l1_orders_2(X2)|m1_subset_1(k4_yellow_0(X2),u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_yellow_0])])).
fof(c_0_13, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_14, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v5_orders_2(X3)|(~v1_yellow_0(X3)|~l1_orders_2(X3))))|(~m1_subset_1(X4,u1_struct_0(X3))|(~r1_orders_2(X3,X4,k3_yellow_0(X3))|X4=k3_yellow_0(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t19_yellow_5])])])])])])).
cnf(c_0_15,plain,(r1_orders_2(X1,X2,k4_yellow_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,negated_conjecture,(k4_yellow_0(esk3_0)=k3_yellow_0(esk3_0)|v7_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,negated_conjecture,(v5_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,negated_conjecture,(l1_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20,plain,(X2=X3|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v7_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21,plain,(m1_subset_1(k4_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_23, plain, (![X2]:(~l1_orders_2(X2)|m1_subset_1(k3_yellow_0(X2),u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_yellow_0])])).
cnf(c_0_24,plain,(X1=k3_yellow_0(X2)|v2_struct_0(X2)|~r1_orders_2(X2,X1,k3_yellow_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25,negated_conjecture,(r1_orders_2(esk3_0,X1,k3_yellow_0(esk3_0))|v7_struct_0(esk3_0)|~v2_yellow_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
fof(c_0_26, plain, (![X2]:(((v1_yellow_0(X2)|~v3_yellow_0(X2))|~l1_orders_2(X2))&((v2_yellow_0(X2)|~v3_yellow_0(X2))|~l1_orders_2(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_yellow_0])])])).
cnf(c_0_27,plain,(X1=k4_yellow_0(X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v7_struct_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_28,plain,(m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29,negated_conjecture,(X1=k3_yellow_0(esk3_0)|v7_struct_0(esk3_0)|~v1_yellow_0(esk3_0)|~v2_yellow_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_30,plain,(v1_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31,negated_conjecture,(v3_yellow_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_32,negated_conjecture,(k4_yellow_0(esk3_0)!=k3_yellow_0(esk3_0)|~v7_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_33,plain,(k4_yellow_0(X1)=k3_yellow_0(X1)|~l1_orders_2(X1)|~v7_struct_0(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34,negated_conjecture,(X1=k3_yellow_0(esk3_0)|v7_struct_0(esk3_0)|~v2_yellow_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_18])])).
cnf(c_0_35,plain,(v2_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36,negated_conjecture,(~v7_struct_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_18])])).
cnf(c_0_37,negated_conjecture,(X1=k3_yellow_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_31]), c_0_18])]), c_0_36])).
cnf(c_0_38,plain,(v7_struct_0(X1)|m1_subset_1(esk2_1(X1),u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39,plain,(v7_struct_0(X1)|~l1_struct_0(X1)|esk1_1(X1)!=esk2_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_40,negated_conjecture,(esk2_1(esk3_0)=k3_yellow_0(esk3_0)|~l1_struct_0(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_36])).
cnf(c_0_41,plain,(v7_struct_0(X1)|m1_subset_1(esk1_1(X1),u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42,negated_conjecture,(esk1_1(esk3_0)!=k3_yellow_0(esk3_0)|~l1_struct_0(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_36])).
cnf(c_0_43,negated_conjecture,(~l1_struct_0(esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_41]), c_0_36]), c_0_42])).
cnf(c_0_44,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_22]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 28
# Proof object formula steps           : 17
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 49
# ...of these trivial                  : 1
# ...subsumed                          : 1
# ...remaining for further processing  : 47
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 30
# ...of the previous two non-trivial   : 20
# Contextual simplify-reflections      : 2
# Paramodulations                      : 28
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 22
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 126
# Rec. Clause-clause subsumption calls : 39
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1792

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45272 pages
