# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B02_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t30_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3)))))), file('out/t3_wellord2__5_premises__rerun.E_input', t30_relat_1)).
fof(l2_wellord1, axiom, (![X1]:(v1_relat_1(X1)=>(v8_relat_2(X1)<=>![X2]:![X3]:![X4]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X4),X1))=>r2_hidden(k4_tarski(X2,X4),X1))))), file('out/t3_wellord2__5_premises__rerun.E_input', l2_wellord1)).
fof(d1_wellord2, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3]:![X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4))))))), file('out/t3_wellord2__5_premises__rerun.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, (![X1]:v1_relat_1(k1_wellord2(X1))), file('out/t3_wellord2__5_premises__rerun.E_input', dt_k1_wellord2)).
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t3_wellord2__5_premises__rerun.E_input', t1_xboole_1)).
fof(t3_wellord2, conjecture, (![X1]:v8_relat_2(k1_wellord2(X1))), file('out/t3_wellord2__5_premises__rerun.E_input', t3_wellord2)).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:(((r2_hidden(X4,k3_relat_1(X6))|~r2_hidden(k4_tarski(X4,X5),X6))|~v1_relat_1(X6))&((r2_hidden(X5,k3_relat_1(X6))|~r2_hidden(k4_tarski(X4,X5),X6))|~v1_relat_1(X6)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_relat_1])])])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:(((~v8_relat_2(X5)|((~r2_hidden(k4_tarski(X6,X7),X5)|~r2_hidden(k4_tarski(X7,X8),X5))|r2_hidden(k4_tarski(X6,X8),X5)))|~v1_relat_1(X5))&((((r2_hidden(k4_tarski(esk1_1(X5),esk2_1(X5)),X5)|v8_relat_2(X5))|~v1_relat_1(X5))&((r2_hidden(k4_tarski(esk2_1(X5),esk3_1(X5)),X5)|v8_relat_2(X5))|~v1_relat_1(X5)))&((~r2_hidden(k4_tarski(esk1_1(X5),esk3_1(X5)),X5)|v8_relat_2(X5))|~v1_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l2_wellord1])])])])])])])).
fof(c_0_8, plain, (![X5]:![X6]:![X7]:![X8]:((((k3_relat_1(X6)=X5|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&(((((~r2_hidden(k4_tarski(X7,X8),X6)|r1_tarski(X7,X8))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&((((~r1_tarski(X7,X8)|r2_hidden(k4_tarski(X7,X8),X6))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))))&(((((r2_hidden(esk4_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&(((r2_hidden(esk5_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))&(((((~r2_hidden(k4_tarski(esk4_2(X5,X6),esk5_2(X5,X6)),X6)|~r1_tarski(esk4_2(X5,X6),esk5_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk4_2(X5,X6),esk5_2(X5,X6)),X6)|r1_tarski(esk4_2(X5,X6),esk5_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])])])).
fof(c_0_9, plain, (![X2]:v1_relat_1(k1_wellord2(X2))), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
cnf(c_0_10,plain,(r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(v8_relat_2(X1)|r2_hidden(k4_tarski(esk2_1(X1),esk3_1(X1)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(k3_relat_1(X1)=X2|~v1_relat_1(X1)|X1!=k1_wellord2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(v8_relat_2(X1)|r2_hidden(k4_tarski(esk1_1(X1),esk2_1(X1)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,plain,(r1_tarski(X4,X3)|~v1_relat_1(X1)|X1!=k1_wellord2(X2)|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)|~r2_hidden(k4_tarski(X4,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16,plain,(v8_relat_2(X1)|r2_hidden(esk3_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_17,plain,(k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_12]), c_0_13])])).
cnf(c_0_18,plain,(v8_relat_2(X1)|r2_hidden(esk2_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_10, c_0_14])).
cnf(c_0_19,plain,(r1_tarski(esk2_1(X1),esk3_1(X1))|v8_relat_2(X1)|X1!=k1_wellord2(X2)|~r2_hidden(esk2_1(X1),X2)|~r2_hidden(esk3_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_15, c_0_11])).
cnf(c_0_20,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk3_1(k1_wellord2(X1)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_13])])).
cnf(c_0_21,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk2_1(k1_wellord2(X1)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_17]), c_0_13])])).
cnf(c_0_22,plain,(r2_hidden(k4_tarski(X4,X3),X1)|~v1_relat_1(X1)|X1!=k1_wellord2(X2)|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)|~r1_tarski(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_23,plain,(r1_tarski(esk2_1(k1_wellord2(X1)),esk3_1(k1_wellord2(X1)))|v8_relat_2(k1_wellord2(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_13])]), c_0_21])).
cnf(c_0_24,plain,(r2_hidden(X2,k3_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(k4_tarski(esk2_1(k1_wellord2(X1)),esk3_1(k1_wellord2(X1))),X2)|X2!=k1_wellord2(X3)|~r2_hidden(esk2_1(k1_wellord2(X1)),X3)|~r2_hidden(esk3_1(k1_wellord2(X1)),X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26,plain,(r1_tarski(esk1_1(X1),esk2_1(X1))|v8_relat_2(X1)|X1!=k1_wellord2(X2)|~r2_hidden(esk1_1(X1),X2)|~r2_hidden(esk2_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_27,plain,(v8_relat_2(X1)|r2_hidden(esk1_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_24, c_0_14])).
cnf(c_0_28,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(k4_tarski(esk2_1(k1_wellord2(X1)),esk3_1(k1_wellord2(X1))),X2)|X2!=k1_wellord2(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_16]), c_0_17]), c_0_17]), c_0_13])]), c_0_21])).
cnf(c_0_29,plain,(r1_tarski(esk1_1(X1),esk2_1(X1))|v8_relat_2(X1)|k1_wellord2(k3_relat_1(X1))!=X1|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_18]), c_0_27])).
cnf(c_0_30,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk2_1(k1_wellord2(X1)),k3_relat_1(X2))|X2!=k1_wellord2(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_24, c_0_28])).
fof(c_0_31, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_32,plain,(v8_relat_2(X1)|r2_hidden(k4_tarski(esk1_1(X1),esk2_1(X1)),X2)|k1_wellord2(k3_relat_1(X1))!=X1|X2!=k1_wellord2(X3)|~r2_hidden(esk1_1(X1),X3)|~r2_hidden(esk2_1(X1),X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_22, c_0_29])).
cnf(c_0_33,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk2_1(k1_wellord2(X1)),X2)|k1_wellord2(X2)!=k1_wellord2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_17]), c_0_13])])).
cnf(c_0_34,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_35,plain,(r1_tarski(esk2_1(X1),esk3_1(X1))|v8_relat_2(X1)|k1_wellord2(k3_relat_1(X1))!=X1|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_16]), c_0_18])).
cnf(c_0_36,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(k4_tarski(esk1_1(k1_wellord2(X1)),esk2_1(k1_wellord2(X1))),X2)|k1_wellord2(X3)!=k1_wellord2(X1)|X2!=k1_wellord2(X3)|~r2_hidden(esk1_1(k1_wellord2(X1)),X3)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_17]), c_0_13])])).
cnf(c_0_37,plain,(r1_tarski(X1,esk3_1(X2))|v8_relat_2(X2)|k1_wellord2(k3_relat_1(X2))!=X2|~r1_tarski(X1,esk2_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(k4_tarski(esk1_1(k1_wellord2(X1)),esk2_1(k1_wellord2(X1))),X2)|X2!=k1_wellord2(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_27]), c_0_17]), c_0_17]), c_0_13])])).
cnf(c_0_39,plain,(v8_relat_2(X1)|r2_hidden(k4_tarski(X2,esk3_1(X1)),X3)|k1_wellord2(k3_relat_1(X1))!=X1|X3!=k1_wellord2(X4)|~r1_tarski(X2,esk2_1(X1))|~r2_hidden(esk3_1(X1),X4)|~r2_hidden(X2,X4)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_22, c_0_37])).
cnf(c_0_40,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk3_1(k1_wellord2(X1)),k3_relat_1(X2))|X2!=k1_wellord2(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_10, c_0_28])).
cnf(c_0_41,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk1_1(k1_wellord2(X1)),k3_relat_1(X2))|X2!=k1_wellord2(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_24, c_0_38])).
cnf(c_0_42,plain,(v8_relat_2(X1)|r2_hidden(k4_tarski(esk1_1(X1),esk3_1(X1)),X2)|k1_wellord2(k3_relat_1(X1))!=X1|X2!=k1_wellord2(X3)|~r2_hidden(esk3_1(X1),X3)|~r2_hidden(esk1_1(X1),X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_29])).
cnf(c_0_43,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk3_1(k1_wellord2(X1)),X2)|k1_wellord2(X2)!=k1_wellord2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_17]), c_0_13])])).
cnf(c_0_44,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(esk1_1(k1_wellord2(X1)),X2)|k1_wellord2(X2)!=k1_wellord2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_17]), c_0_13])])).
fof(c_0_45, negated_conjecture, (~(![X1]:v8_relat_2(k1_wellord2(X1)))), inference(assume_negation,[status(cth)],[t3_wellord2])).
cnf(c_0_46,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(k4_tarski(esk1_1(k1_wellord2(X1)),esk3_1(k1_wellord2(X1))),X2)|k1_wellord2(X3)!=k1_wellord2(X1)|X2!=k1_wellord2(X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_17]), c_0_13])]), c_0_44])).
fof(c_0_47, negated_conjecture, (~v8_relat_2(k1_wellord2(esk6_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_45])])])).
cnf(c_0_48,plain,(v8_relat_2(X1)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(esk1_1(X1),esk3_1(X1)),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_49,plain,(v8_relat_2(k1_wellord2(X1))|r2_hidden(k4_tarski(esk1_1(k1_wellord2(X1)),esk3_1(k1_wellord2(X1))),k1_wellord2(X2))|k1_wellord2(X2)!=k1_wellord2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_46]), c_0_13])])).
cnf(c_0_50,negated_conjecture,(~v8_relat_2(k1_wellord2(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_51,plain,(v8_relat_2(k1_wellord2(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_13])])).
cnf(c_0_52,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_51])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 40
# Proof object formula steps           : 13
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 107
# ...of these trivial                  : 0
# ...subsumed                          : 22
# ...remaining for further processing  : 85
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 30
# Generated clauses                    : 114
# ...of the previous two non-trivial   : 110
# Contextual simplify-reflections      : 30
# Paramodulations                      : 110
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 38
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 46
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 1744
# Rec. Clause-clause subsumption calls : 331
# Non-unit clause-clause subsumptions  : 53
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5045

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45296 pages
