# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_waybel_0, conjecture, (![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>((X3=X4&v1_waybel_0(X3,X1))=>v1_waybel_0(X4,X2)))))))), file('out/t3_waybel_0__3_premises__rerun.E_input', t3_waybel_0)).
fof(d1_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X3,X5)&r1_orders_2(X1,X4,X5)))))))))))))), file('out/t3_waybel_0__3_premises__rerun.E_input', d1_waybel_0)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t3_waybel_0__3_premises__rerun.E_input', t4_subset)).
fof(t1_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X3=X5&X4=X6)=>((r1_orders_2(X1,X3,X4)=>r1_orders_2(X2,X5,X6))&(r2_orders_2(X1,X3,X4)=>r2_orders_2(X2,X5,X6)))))))))))), file('out/t3_waybel_0__3_premises__rerun.E_input', t1_yellow_0)).
fof(c_0_4, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>((X3=X4&v1_waybel_0(X3,X1))=>v1_waybel_0(X4,X2))))))))), inference(assume_negation,[status(cth)],[t3_waybel_0])).
fof(c_0_5, plain, (![X6]:![X7]:![X8]:![X9]:![X13]:(((((((((m1_subset_1(esk5_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((((r2_hidden(esk5_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((((r1_orders_2(X6,X8,esk5_4(X6,X7,X8,X9))|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((((r1_orders_2(X6,X9,esk5_4(X6,X7,X8,X9))|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))&((((m1_subset_1(esk6_2(X6,X7),u1_struct_0(X6))|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((m1_subset_1(esk7_2(X6,X7),u1_struct_0(X6))|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((r2_hidden(esk6_2(X6,X7),X7)|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((r2_hidden(esk7_2(X6,X7),X7)|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((~m1_subset_1(X13,u1_struct_0(X6))|(~r2_hidden(X13,X7)|(~r1_orders_2(X6,esk6_2(X6,X7),X13)|~r1_orders_2(X6,esk7_2(X6,X7),X13))))|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel_0])])])])])])])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_7, negated_conjecture, ((l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&((esk3_0=esk4_0&v1_waybel_0(esk3_0,esk1_0))&~v1_waybel_0(esk4_0,esk2_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_8,plain,(r2_hidden(esk5_4(X1,X2,X3,X4),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,negated_conjecture,(esk3_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,negated_conjecture,(~v1_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(r2_hidden(esk5_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X2)|~r2_hidden(X3,X2)|~v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_8, c_0_9]), c_0_9])).
cnf(c_0_14,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,negated_conjecture,(v1_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,plain,(v1_waybel_0(X2,X1)|r2_hidden(esk7_2(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_19,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20,negated_conjecture,(~v1_waybel_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_12, c_0_11])).
cnf(c_0_21,negated_conjecture,(r2_hidden(esk5_4(esk1_0,esk3_0,X1,X2),esk3_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_22,negated_conjecture,(r2_hidden(esk7_2(esk2_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_23,plain,(v1_waybel_0(X2,X1)|r2_hidden(esk6_2(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_24,plain,(m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_25, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((((((((~r1_orders_2(X7,X9,X10)|r1_orders_2(X8,X11,X12))|(X9!=X11|X10!=X12))|~m1_subset_1(X12,u1_struct_0(X8)))|~m1_subset_1(X11,u1_struct_0(X8)))|~m1_subset_1(X10,u1_struct_0(X7)))|~m1_subset_1(X9,u1_struct_0(X7)))|g1_orders_2(u1_struct_0(X7),u1_orders_2(X7))!=g1_orders_2(u1_struct_0(X8),u1_orders_2(X8)))|~l1_orders_2(X8))|~l1_orders_2(X7))&(((((((((~r2_orders_2(X7,X9,X10)|r2_orders_2(X8,X11,X12))|(X9!=X11|X10!=X12))|~m1_subset_1(X12,u1_struct_0(X8)))|~m1_subset_1(X11,u1_struct_0(X8)))|~m1_subset_1(X10,u1_struct_0(X7)))|~m1_subset_1(X9,u1_struct_0(X7)))|g1_orders_2(u1_struct_0(X7),u1_orders_2(X7))!=g1_orders_2(u1_struct_0(X8),u1_orders_2(X8)))|~l1_orders_2(X8))|~l1_orders_2(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_yellow_0])])])])])])).
cnf(c_0_26,negated_conjecture,(r2_hidden(esk5_4(esk1_0,esk3_0,X1,esk7_2(esk2_0,esk3_0)),esk3_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27,negated_conjecture,(r2_hidden(esk6_2(esk2_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_28,plain,(m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X1))|~r2_hidden(X4,X2)|~r2_hidden(X3,X2)|~v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_24, c_0_9]), c_0_9])).
cnf(c_0_29,plain,(r1_orders_2(X1,X3,esk5_4(X1,X2,X3,X4))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30,plain,(r1_orders_2(X1,X4,esk5_4(X1,X2,X3,X4))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_31,plain,(r1_orders_2(X2,X5,X6)|~l1_orders_2(X1)|~l1_orders_2(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|X4!=X6|X3!=X5|~r1_orders_2(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_9, c_0_18])).
cnf(c_0_33,negated_conjecture,(r2_hidden(esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)),esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk3_0,X1,X2),u1_struct_0(esk1_0))|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_35,plain,(r1_orders_2(X1,X2,esk5_4(X1,X3,X2,X4))|~r2_hidden(X4,X3)|~r2_hidden(X2,X3)|~v1_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_29, c_0_9]), c_0_9])).
cnf(c_0_36,plain,(r1_orders_2(X1,X2,esk5_4(X1,X3,X4,X2))|~r2_hidden(X2,X3)|~r2_hidden(X4,X3)|~v1_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_30, c_0_9]), c_0_9])).
cnf(c_0_37,plain,(r1_orders_2(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_orders_2(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~l1_orders_2(X1)|~l1_orders_2(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
cnf(c_0_38,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39,negated_conjecture,(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk3_0,X1,esk7_2(esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_22])).
cnf(c_0_41,negated_conjecture,(r1_orders_2(esk1_0,X1,esk5_4(esk1_0,esk3_0,X1,X2))|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_42,negated_conjecture,(r1_orders_2(esk1_0,X1,esk5_4(esk1_0,esk3_0,X2,X1))|~r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_43,negated_conjecture,(r1_orders_2(esk2_0,X1,esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r1_orders_2(X2,X1,esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))|~m1_subset_1(esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_19])])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_27])).
cnf(c_0_45,negated_conjecture,(r1_orders_2(esk1_0,X1,esk5_4(esk1_0,esk3_0,X1,esk7_2(esk2_0,esk3_0)))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_41, c_0_22])).
cnf(c_0_46,plain,(v1_waybel_0(X2,X1)|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_47,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_9, c_0_14])).
cnf(c_0_48,negated_conjecture,(r1_orders_2(esk1_0,esk7_2(esk2_0,esk3_0),esk5_4(esk1_0,esk3_0,X1,esk7_2(esk2_0,esk3_0)))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_22])).
cnf(c_0_49,plain,(v1_waybel_0(X2,X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_50,plain,(v1_waybel_0(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,esk7_2(X1,X2),X3)|~r1_orders_2(X1,esk6_2(X1,X2),X3)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_51,negated_conjecture,(r1_orders_2(esk2_0,X1,esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))|~r1_orders_2(esk1_0,X1,esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_16])])).
cnf(c_0_52,negated_conjecture,(r1_orders_2(esk1_0,esk6_2(esk2_0,esk3_0),esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_45, c_0_27])).
cnf(c_0_53,negated_conjecture,(m1_subset_1(esk6_2(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_54,negated_conjecture,(m1_subset_1(esk6_2(esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_47, c_0_27])).
cnf(c_0_55,negated_conjecture,(r1_orders_2(esk1_0,esk7_2(esk2_0,esk3_0),esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_48, c_0_27])).
cnf(c_0_56,negated_conjecture,(m1_subset_1(esk7_2(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_57,negated_conjecture,(m1_subset_1(esk7_2(esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_47, c_0_22])).
cnf(c_0_58,plain,(v1_waybel_0(X1,X2)|~r1_orders_2(X2,esk6_2(X2,X1),X3)|~r1_orders_2(X2,esk7_2(X2,X1),X3)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_50, c_0_9])).
cnf(c_0_59,negated_conjecture,(r1_orders_2(esk2_0,esk6_2(esk2_0,esk3_0),esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_54])])).
cnf(c_0_60,negated_conjecture,(r1_orders_2(esk2_0,esk7_2(esk2_0,esk3_0),esk5_4(esk1_0,esk3_0,esk6_2(esk2_0,esk3_0),esk7_2(esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_55]), c_0_56]), c_0_57])])).
cnf(c_0_61,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_33]), c_0_18]), c_0_19])]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 53
# Proof object formula steps           : 9
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 54
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 274
# ...of these trivial                  : 2
# ...subsumed                          : 4
# ...remaining for further processing  : 268
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 2319
# ...of the previous two non-trivial   : 2268
# Contextual simplify-reflections      : 9
# Paramodulations                      : 2317
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 246
#    Positive orientable unit clauses  : 109
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 136
# Current number of unprocessed clauses: 2034
# ...number of literals in the above   : 2564
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 4231
# Rec. Clause-clause subsumption calls : 3243
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 247
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1593
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 99716

# -------------------------------------------------
# User time                : 0.040 s
# System time              : 0.004 s
# Total time               : 0.044 s
# Maximum resident set size: 45336 pages
