# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_yellow_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))))), file('out/t3_yellow_1__14_premises__rerun.E_input', t3_yellow_1)).
fof(d1_yellow_1, axiom, (![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), file('out/t3_yellow_1__14_premises__rerun.E_input', d1_yellow_1)).
fof(dt_k2_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1)))), file('out/t3_yellow_1__14_premises__rerun.E_input', dt_k2_yellow_1)).
fof(fc5_yellow_1, axiom, (![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1)))))), file('out/t3_yellow_1__14_premises__rerun.E_input', fc5_yellow_1)).
fof(free_g1_orders_2, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3]:![X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4)))), file('out/t3_yellow_1__14_premises__rerun.E_input', free_g1_orders_2)).
fof(abstractness_v1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))), file('out/t3_yellow_1__14_premises__rerun.E_input', abstractness_v1_orders_2)).
fof(dt_u1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))), file('out/t3_yellow_1__14_premises__rerun.E_input', dt_u1_orders_2)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t3_yellow_1__14_premises__rerun.E_input', redefinition_r3_orders_2)).
fof(fc1_struct_0, axiom, (![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1)))), file('out/t3_yellow_1__14_premises__rerun.E_input', fc1_struct_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t3_yellow_1__14_premises__rerun.E_input', dt_l1_orders_2)).
fof(d1_wellord2, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3]:![X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4))))))), file('out/t3_yellow_1__14_premises__rerun.E_input', d1_wellord2)).
fof(redefinition_k1_yellow_1, axiom, (![X1]:k1_yellow_1(X1)=k1_wellord2(X1)), file('out/t3_yellow_1__14_premises__rerun.E_input', redefinition_k1_yellow_1)).
fof(dt_k1_wellord2, axiom, (![X1]:v1_relat_1(k1_wellord2(X1))), file('out/t3_yellow_1__14_premises__rerun.E_input', dt_k1_wellord2)).
fof(d5_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))))))), file('out/t3_yellow_1__14_premises__rerun.E_input', d5_orders_2)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t3_yellow_1__14_premises__rerun.E_input', t2_subset)).
fof(c_0_15, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))))), inference(assume_negation,[status(cth)],[t3_yellow_1])).
fof(c_0_16, negated_conjecture, ((~v1_xboole_0(esk1_0)&(m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))&((~r3_orders_2(k2_yellow_1(esk1_0),esk2_0,esk3_0)|~r1_tarski(esk2_0,esk3_0))&(r3_orders_2(k2_yellow_1(esk1_0),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])])).
fof(c_0_17, plain, (![X2]:k2_yellow_1(X2)=g1_orders_2(X2,k1_yellow_1(X2))), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_18, plain, (![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&l1_orders_2(k2_yellow_1(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k2_yellow_1])])])).
fof(c_0_19, plain, (![X2]:![X2]:![X2]:![X2]:(v1_orders_2(k2_yellow_1(X2))&(v3_orders_2(k2_yellow_1(X2))&(v4_orders_2(k2_yellow_1(X2))&v5_orders_2(k2_yellow_1(X2)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc5_yellow_1])])])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X8]:(((X5=X7|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5))))&((X6=X8|g1_orders_2(X5,X6)!=g1_orders_2(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])])])).
fof(c_0_21, plain, (![X2]:(~l1_orders_2(X2)|(~v1_orders_2(X2)|X2=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
fof(c_0_22, plain, (![X2]:(~l1_orders_2(X2)|m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,plain,(k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26,plain,(l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(X2=X3|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|g1_orders_2(X2,X1)!=g1_orders_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,plain,(X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,plain,(m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31,plain,(X1=X4|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|g1_orders_2(X2,X1)!=g1_orders_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_32, plain, (![X2]:((~v2_struct_0(X2)|~l1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
fof(c_0_33, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_34,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_36,plain,(l1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_37,plain,(v3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_27, c_0_25])).
cnf(c_0_38,negated_conjecture,(r1_tarski(esk2_0,esk3_0)|r3_orders_2(k2_yellow_1(esk1_0),esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_40,plain,(u1_struct_0(X1)=X2|X1!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_41,plain,(v1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42,plain,(u1_orders_2(X1)=X2|X1!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_29]), c_0_30])).
cnf(c_0_43,plain,(v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,X1)|v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|~r3_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])])).
cnf(c_0_46,negated_conjecture,(r1_tarski(esk2_0,esk3_0)|r3_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_38, c_0_25])).
cnf(c_0_47,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(rw,[status(thm)],[c_0_39, c_0_25])).
fof(c_0_48, plain, (![X5]:![X6]:![X7]:![X8]:((((k3_relat_1(X6)=X5|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&(((((~r2_hidden(k4_tarski(X7,X8),X6)|r1_tarski(X7,X8))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))&((((~r1_tarski(X7,X8)|r2_hidden(k4_tarski(X7,X8),X6))|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5)))|X6!=k1_wellord2(X5))|~v1_relat_1(X6))))&(((((r2_hidden(esk4_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&(((r2_hidden(esk5_2(X5,X6),X5)|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))&(((((~r2_hidden(k4_tarski(esk4_2(X5,X6),esk5_2(X5,X6)),X6)|~r1_tarski(esk4_2(X5,X6),esk5_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6))&((((r2_hidden(k4_tarski(esk4_2(X5,X6),esk5_2(X5,X6)),X6)|r1_tarski(esk4_2(X5,X6),esk5_2(X5,X6)))|k3_relat_1(X6)!=X5)|X6=k1_wellord2(X5))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])])])).
fof(c_0_49, plain, (![X2]:k1_yellow_1(X2)=k1_wellord2(X2)), inference(variable_rename,[status(thm)],[redefinition_k1_yellow_1])).
fof(c_0_50, plain, (![X2]:v1_relat_1(k1_wellord2(X2))), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_51, plain, (![X4]:![X5]:![X6]:(((((~r1_orders_2(X4,X5,X6)|r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&((((~r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4))|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_orders_2])])])])])])).
cnf(c_0_52,plain,(u1_struct_0(g1_orders_2(X1,X2))=X1|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_53,plain,(v1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_41, c_0_25])).
cnf(c_0_54,plain,(u1_orders_2(g1_orders_2(X1,X2))=X2|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_55,plain,(v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_56,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)|v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
fof(c_0_57, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_58,plain,(r1_tarski(X4,X3)|~v1_relat_1(X1)|X1!=k1_wellord2(X2)|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)|~r2_hidden(k4_tarski(X4,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59,plain,(k1_yellow_1(X1)=k1_wellord2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60,plain,(v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61,plain,(r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62,plain,(u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_36]), c_0_53])])).
cnf(c_0_63,plain,(u1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))=k1_yellow_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_36]), c_0_53])])).
cnf(c_0_64,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_36])])).
cnf(c_0_65,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_66,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_67,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68,plain,(r2_hidden(k4_tarski(X4,X3),X1)|~v1_relat_1(X1)|X1!=k1_wellord2(X2)|~r2_hidden(X3,X2)|~r2_hidden(X4,X2)|~r1_tarski(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_69,plain,(r1_tarski(X4,X3)|X1!=k1_yellow_1(X2)|~v1_relat_1(X1)|~r2_hidden(X4,X2)|~r2_hidden(X3,X2)|~r2_hidden(k4_tarski(X4,X3),X1)), inference(rw,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_70,plain,(v1_relat_1(k1_yellow_1(X1))), inference(rw,[status(thm)],[c_0_60, c_0_59])).
cnf(c_0_71,plain,(r2_hidden(k4_tarski(X1,X2),k1_yellow_1(X3))|~r1_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_36])]), c_0_63])).
cnf(c_0_72,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_62]), c_0_65])).
cnf(c_0_73,negated_conjecture,(m1_subset_1(esk3_0,esk1_0)), inference(rw,[status(thm)],[c_0_47, c_0_62])).
cnf(c_0_74,negated_conjecture,(m1_subset_1(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_35, c_0_62])).
cnf(c_0_75,negated_conjecture,(r2_hidden(esk2_0,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|v1_xboole_0(u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(spm,[status(thm)],[c_0_66, c_0_35])).
cnf(c_0_76,negated_conjecture,(r2_hidden(esk3_0,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|v1_xboole_0(u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(spm,[status(thm)],[c_0_66, c_0_47])).
cnf(c_0_77,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,esk3_0)|~r2_hidden(k4_tarski(X1,esk3_0),u1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_47]), c_0_36])])).
cnf(c_0_78,plain,(r2_hidden(k4_tarski(X4,X3),X1)|X1!=k1_yellow_1(X2)|~v1_relat_1(X1)|~r1_tarski(X4,X3)|~r2_hidden(X4,X2)|~r2_hidden(X3,X2)), inference(rw,[status(thm)],[c_0_68, c_0_59])).
cnf(c_0_79,plain,(r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X1,X2),k1_yellow_1(X3))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_69]), c_0_70])])).
cnf(c_0_80,negated_conjecture,(r2_hidden(k4_tarski(esk2_0,esk3_0),k1_yellow_1(esk1_0))|r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_74])])).
cnf(c_0_81,negated_conjecture,(r2_hidden(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_62]), c_0_62]), c_0_65])).
cnf(c_0_82,negated_conjecture,(r2_hidden(esk3_0,esk1_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_62]), c_0_62]), c_0_65])).
cnf(c_0_83,plain,(v2_struct_0(X2)|r3_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r1_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_84,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,esk3_0)|~r2_hidden(k4_tarski(X1,esk3_0),u1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_77, c_0_62])).
cnf(c_0_85,plain,(r2_hidden(k4_tarski(X1,X2),k1_yellow_1(X3))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_78]), c_0_70])])).
cnf(c_0_86,negated_conjecture,(r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81]), c_0_82])])).
cnf(c_0_87,negated_conjecture,(v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|r3_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,X1)|~r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_35]), c_0_36]), c_0_37])])).
cnf(c_0_88,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,esk3_0)|~r2_hidden(k4_tarski(X1,esk3_0),k1_yellow_1(esk1_0))|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_84, c_0_63])).
cnf(c_0_89,negated_conjecture,(r2_hidden(k4_tarski(esk2_0,esk3_0),k1_yellow_1(X1))|~r2_hidden(esk2_0,X1)|~r2_hidden(esk3_0,X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_90,negated_conjecture,(~r1_tarski(esk2_0,esk3_0)|~r3_orders_2(k2_yellow_1(esk1_0),esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_91,negated_conjecture,(v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|r3_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,X1)|~r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,X1)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_87, c_0_62])).
cnf(c_0_92,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)|~r2_hidden(k4_tarski(esk2_0,esk3_0),k1_yellow_1(esk1_0))), inference(spm,[status(thm)],[c_0_88, c_0_74])).
cnf(c_0_93,negated_conjecture,(r2_hidden(k4_tarski(esk2_0,esk3_0),k1_yellow_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_81]), c_0_82])])).
cnf(c_0_94,negated_conjecture,(~r1_tarski(esk2_0,esk3_0)|~r3_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_90, c_0_25])).
cnf(c_0_95,negated_conjecture,(v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|r3_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)|~r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_91, c_0_73])).
cnf(c_0_96,negated_conjecture,(r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_97,negated_conjecture,(~r3_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_86])])).
cnf(c_0_98,negated_conjecture,(v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_96])]), c_0_97])).
cnf(c_0_99,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_98]), c_0_62]), c_0_36])]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 100
# Proof object clause steps            : 69
# Proof object formula steps           : 31
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 67
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 30
# Processed clauses                    : 134
# ...of these trivial                  : 2
# ...subsumed                          : 9
# ...remaining for further processing  : 123
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 33
# Generated clauses                    : 100
# ...of the previous two non-trivial   : 91
# Contextual simplify-reflections      : 6
# Paramodulations                      : 79
# Factorizations                       : 0
# Equation resolutions                 : 21
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 37
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 5
# Current number of archived formulas  : 0
# Current number of archived clauses   : 68
# Clause-clause subsumption calls (NU) : 493
# Rec. Clause-clause subsumption calls : 212
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 109
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5107

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 43288 pages
