# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_filter_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>((v14_lattices(X1)&v14_lattices(X2))<=>v14_lattices(k7_filter_1(X1,X2)))))), file('out/t41_filter_1__6_premises__rerun__rerun.E_input', t41_filter_1)).
fof(t40_filter_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>((v14_lattices(X1)&v14_lattices(X2))<=>v14_lattices(k7_filter_1(X1,X2)))))), file('out/t41_filter_1__6_premises__rerun__rerun.E_input', t40_filter_1)).
fof(c_0_2, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>((v14_lattices(X1)&v14_lattices(X2))<=>v14_lattices(k7_filter_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t41_filter_1])).
fof(c_0_3, plain, (![X3]:![X4]:(((((~v14_lattices(X3)|~v14_lattices(X4))|v14_lattices(k7_filter_1(X3,X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4))))|(v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3))))&((((v14_lattices(X3)|~v14_lattices(k7_filter_1(X3,X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4))))|(v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3))))&(((v14_lattices(X4)|~v14_lattices(k7_filter_1(X3,X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4))))|(v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t40_filter_1])])])])])])])).
fof(c_0_4, negated_conjecture, (((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&l3_lattices(esk1_0)))&((~v2_struct_0(esk2_0)&(v10_lattices(esk2_0)&l3_lattices(esk2_0)))&(((~v14_lattices(esk1_0)|~v14_lattices(esk2_0))|~v14_lattices(k7_filter_1(esk1_0,esk2_0)))&((v14_lattices(esk1_0)|v14_lattices(k7_filter_1(esk1_0,esk2_0)))&(v14_lattices(esk2_0)|v14_lattices(k7_filter_1(esk1_0,esk2_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_2])])])])])])])).
cnf(c_0_5,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v14_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X1)|~l3_lattices(X2)|~v10_lattices(X2)|~v14_lattices(k7_filter_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(v14_lattices(k7_filter_1(esk1_0,esk2_0))|v14_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,negated_conjecture,(l3_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,negated_conjecture,(l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(v10_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13,negated_conjecture,(~v14_lattices(k7_filter_1(esk1_0,esk2_0))|~v14_lattices(esk2_0)|~v14_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14,negated_conjecture,(v14_lattices(esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7]), c_0_8]), c_0_9]), c_0_10])]), c_0_11]), c_0_12])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v14_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)|~l3_lattices(X2)|~v10_lattices(X2)|~v14_lattices(k7_filter_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_16,negated_conjecture,(v14_lattices(k7_filter_1(esk1_0,esk2_0))|v14_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(~v14_lattices(k7_filter_1(esk1_0,esk2_0))|~v14_lattices(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14])])).
cnf(c_0_18,negated_conjecture,(v14_lattices(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_7]), c_0_8]), c_0_9]), c_0_10])]), c_0_11]), c_0_12])).
cnf(c_0_19,negated_conjecture,(~v14_lattices(k7_filter_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18])])).
cnf(c_0_20,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v14_lattices(k7_filter_1(X1,X2))|~l3_lattices(X1)|~v10_lattices(X1)|~l3_lattices(X2)|~v10_lattices(X2)|~v14_lattices(X2)|~v14_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_21,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_14]), c_0_18]), c_0_7]), c_0_8]), c_0_9]), c_0_10])]), c_0_11]), c_0_12]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 22
# Proof object clause steps            : 17
# Proof object formula steps           : 5
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 29
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 28
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 9
# ...of the previous two non-trivial   : 5
# Contextual simplify-reflections      : 0
# Paramodulations                      : 9
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 12
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 4
# Rec. Clause-clause subsumption calls : 4
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1016

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45328 pages
