# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t42_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>(r1_tmap_1(X2,X1,X3,X4)<=>![X5]:(m1_connsp_2(X5,X1,k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4))=>m1_connsp_2(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X5),X2,X4)))))))), file('out/t42_tmap_1__14_premises__rerun.E_input', t42_tmap_1)).
fof(redefinition_k8_relset_1, axiom, (![X1]:![X2]:![X3]:![X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k8_relat_1(X3,X4))), file('out/t42_tmap_1__14_premises__rerun.E_input', redefinition_k8_relset_1)).
fof(cc1_relset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('out/t42_tmap_1__14_premises__rerun.E_input', cc1_relset_1)).
fof(redefinition_k7_relset_1, axiom, (![X1]:![X2]:![X3]:![X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k7_relat_1(X3,X4))), file('out/t42_tmap_1__14_premises__rerun.E_input', redefinition_k7_relset_1)).
fof(d2_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_tmap_1(X1,X2,X3,X4)<=>![X5]:(m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))=>?[X6]:(m1_connsp_2(X6,X1,X4)&r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X6),X5))))))))), file('out/t42_tmap_1__14_premises__rerun.E_input', d2_tmap_1)).
fof(t75_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>r1_tarski(k7_relat_1(X2,k8_relat_1(X2,X1)),X1))), file('out/t42_tmap_1__14_premises__rerun.E_input', t75_funct_1)).
fof(dt_k3_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>m1_subset_1(k3_funct_2(X1,X2,X3,X4),X2))), file('out/t42_tmap_1__14_premises__rerun.E_input', dt_k3_funct_2)).
fof(dt_m1_connsp_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>![X3]:(m1_connsp_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t42_tmap_1__14_premises__rerun.E_input', dt_m1_connsp_2)).
fof(t6_connsp_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(m1_connsp_2(X3,X1,X2)<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(v3_pre_topc(X4,X1)&(r1_tarski(X4,X3)&r2_hidden(X2,X4))))))))), file('out/t42_tmap_1__14_premises__rerun.E_input', t6_connsp_2)).
fof(t95_funct_2, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(X2))=>(r1_tarski(k7_relset_1(X1,X2,X3,X4),X5)<=>r1_tarski(X4,k8_relset_1(X1,X2,X3,X5))))))))), file('out/t42_tmap_1__14_premises__rerun.E_input', t95_funct_2)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t42_tmap_1__14_premises__rerun.E_input', t4_subset)).
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t42_tmap_1__14_premises__rerun.E_input', t1_xboole_1)).
fof(dt_k8_relset_1, axiom, (![X1]:![X2]:![X3]:![X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1)))), file('out/t42_tmap_1__14_premises__rerun.E_input', dt_k8_relset_1)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t42_tmap_1__14_premises__rerun.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t42_tmap_1__14_premises__rerun.E_input', dt_l1_pre_topc)).
fof(c_0_15, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>(r1_tmap_1(X2,X1,X3,X4)<=>![X5]:(m1_connsp_2(X5,X1,k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4))=>m1_connsp_2(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X5),X2,X4))))))))), inference(assume_negation,[status(cth)],[t42_tmap_1])).
fof(c_0_16, plain, (![X5]:![X6]:![X7]:![X8]:(~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|k8_relset_1(X5,X6,X7,X8)=k8_relat_1(X7,X8))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])])])).
fof(c_0_17, negated_conjecture, (![X11]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(((m1_connsp_2(esk5_0,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0))&(~m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0),esk2_0,esk4_0)|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)))&(r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)|(~m1_connsp_2(X11,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))|m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X11),esk2_0,esk4_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])])])])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|v1_relat_1(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:(~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|k7_relset_1(X5,X6,X7,X8)=k7_relat_1(X7,X8))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])])])).
fof(c_0_20, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X14]:((((((((m1_connsp_2(esk6_5(X7,X8,X9,X10,X11),X7,X10)|~m1_connsp_2(X11,X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10)))|~r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((r1_tarski(k7_relset_1(u1_struct_0(X7),u1_struct_0(X8),X9,esk6_5(X7,X8,X9,X10,X11)),X11)|~m1_connsp_2(X11,X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10)))|~r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))))&((((((m1_connsp_2(esk7_4(X7,X8,X9,X10),X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10))|r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((~m1_connsp_2(X14,X7,X10)|~r1_tarski(k7_relset_1(u1_struct_0(X7),u1_struct_0(X8),X9,X14),esk7_4(X7,X8,X9,X10)))|r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_tmap_1])])])])])])])])).
fof(c_0_21, plain, (![X3]:![X4]:((~v1_relat_1(X4)|~v1_funct_1(X4))|r1_tarski(k7_relat_1(X4,k8_relat_1(X4,X3)),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t75_funct_1])])).
cnf(c_0_22,plain,(k8_relset_1(X1,X2,X3,X4)=k8_relat_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,plain,(v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,plain,(k7_relset_1(X1,X2,X3,X4)=k7_relat_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tmap_1(X1,X2,X3,X4)|m1_connsp_2(esk7_4(X1,X2,X3,X4),X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,negated_conjecture,(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32,negated_conjecture,(v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_33,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_34,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tmap_1(X1,X2,X3,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X5),esk7_4(X1,X2,X3,X4))|~m1_connsp_2(X5,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36,plain,(r1_tarski(k7_relat_1(X1,k8_relat_1(X1,X2)),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37,negated_conjecture,(k8_relat_1(esk3_0,X1)=k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_38,negated_conjecture,(v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_39,negated_conjecture,(k7_relat_1(esk3_0,X1)=k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1)), inference(spm,[status(thm)],[c_0_25, c_0_23])).
cnf(c_0_40,negated_conjecture,(m1_connsp_2(esk7_4(esk2_0,esk1_0,esk3_0,X1),esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1))|r1_tmap_1(esk2_0,esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_23]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34])).
cnf(c_0_41,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_42,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,X1)|~r1_tarski(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X2),esk7_4(esk2_0,esk1_0,esk3_0,X1))|~m1_connsp_2(X2,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_27]), c_0_23]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34])).
cnf(c_0_43,negated_conjecture,(r1_tarski(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1)),X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_28])]), c_0_39])).
cnf(c_0_44,negated_conjecture,(m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1),esk2_0,esk4_0)|r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)|~m1_connsp_2(X1,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_45,negated_conjecture,(m1_connsp_2(esk7_4(esk2_0,esk1_0,esk3_0,esk4_0),esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))|r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,X1)|~m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk7_4(esk2_0,esk1_0,esk3_0,X1)),esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47,negated_conjecture,(m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk7_4(esk2_0,esk1_0,esk3_0,esk4_0)),esk2_0,esk4_0)|r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
fof(c_0_48, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X8,X5)))|m1_subset_1(k3_funct_2(X5,X6,X7,X8),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_funct_2])])])).
cnf(c_0_49,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_connsp_2(esk6_5(X1,X2,X3,X4,X5),X1,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_tmap_1(X1,X2,X3,X4)|~m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_50,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_41])])).
cnf(c_0_51,negated_conjecture,(m1_connsp_2(esk5_0,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_52, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(~m1_connsp_2(X6,X4,X5)|m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_connsp_2])])])])])])).
cnf(c_0_53,plain,(m1_subset_1(k3_funct_2(X1,X2,X3,X4),X2)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_54, plain, (![X5]:![X6]:![X7]:![X9]:((((((m1_subset_1(esk8_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(X5)))|~m1_connsp_2(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((v3_pre_topc(esk8_3(X5,X6,X7),X5)|~m1_connsp_2(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((r1_tarski(esk8_3(X5,X6,X7),X7)|~m1_connsp_2(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((r2_hidden(X6,esk8_3(X5,X6,X7))|~m1_connsp_2(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))))))&(((((~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X5)))|(~v3_pre_topc(X9,X5)|(~r1_tarski(X9,X7)|~r2_hidden(X6,X9))))|m1_connsp_2(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_connsp_2])])])])])])])])).
fof(c_0_55, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((~r1_tarski(k7_relset_1(X6,X7,X8,X9),X10)|r1_tarski(X9,k8_relset_1(X6,X7,X8,X10)))|~m1_subset_1(X10,k1_zfmisc_1(X7)))|~m1_subset_1(X9,k1_zfmisc_1(X6)))|(~v1_funct_1(X8)|(~v1_funct_2(X8,X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|v1_xboole_0(X7))|v1_xboole_0(X6))&((((((~r1_tarski(X9,k8_relset_1(X6,X7,X8,X10))|r1_tarski(k7_relset_1(X6,X7,X8,X9),X10))|~m1_subset_1(X10,k1_zfmisc_1(X7)))|~m1_subset_1(X9,k1_zfmisc_1(X6)))|(~v1_funct_1(X8)|(~v1_funct_2(X8,X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))|v1_xboole_0(X7))|v1_xboole_0(X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t95_funct_2])])])])])])])).
cnf(c_0_56,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk6_5(X1,X2,X3,X4,X5)),X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_tmap_1(X1,X2,X3,X4)|~m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_57,negated_conjecture,(m1_connsp_2(esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,X1),esk2_0,esk4_0)|~m1_connsp_2(X1,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_23]), c_0_41]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34])).
cnf(c_0_58,negated_conjecture,(m1_connsp_2(esk5_0,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_50])])).
cnf(c_0_59,plain,(m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m1_connsp_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_27]), c_0_23]), c_0_28])])).
fof(c_0_61, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_62,plain,(v2_struct_0(X1)|v3_pre_topc(esk8_3(X1,X2,X3),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_63,plain,(v2_struct_0(X1)|m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64,plain,(v1_xboole_0(X1)|v1_xboole_0(X2)|r1_tarski(X4,k8_relset_1(X1,X2,X3,X5))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X5,k1_zfmisc_1(X2))|~r1_tarski(k7_relset_1(X1,X2,X3,X4),X5)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_65,negated_conjecture,(r1_tarski(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,X1)),X1)|~m1_connsp_2(X1,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_50]), c_0_23]), c_0_41]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34])).
cnf(c_0_66,negated_conjecture,(m1_connsp_2(esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_67,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_58]), c_0_29]), c_0_31])]), c_0_33])).
cnf(c_0_68,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_60, c_0_41])).
cnf(c_0_69,plain,(v2_struct_0(X1)|m1_connsp_2(X3,X1,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,X4)|~r1_tarski(X4,X3)|~v3_pre_topc(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_70,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71,plain,(v3_pre_topc(esk8_3(X1,X2,X3),X1)|v2_struct_0(X1)|~m1_connsp_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_62, c_0_59])).
cnf(c_0_72,plain,(m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_connsp_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_63, c_0_59])).
cnf(c_0_73,plain,(v2_struct_0(X1)|r2_hidden(X2,esk8_3(X1,X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_74, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_75,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|r1_tarski(X1,k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X2))|~r1_tarski(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_27]), c_0_23]), c_0_28])])).
cnf(c_0_76,negated_conjecture,(r1_tarski(k7_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)),esk5_0)), inference(spm,[status(thm)],[c_0_65, c_0_58])).
cnf(c_0_77,negated_conjecture,(m1_subset_1(esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_66]), c_0_41]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_78,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_79,plain,(v2_struct_0(X1)|r1_tarski(esk8_3(X1,X2,X3),X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_80,plain,(m1_connsp_2(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X3,X4)|~v3_pre_topc(X4,X2)|~r1_tarski(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_81,negated_conjecture,(v3_pre_topc(esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_66]), c_0_41]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_82,negated_conjecture,(m1_subset_1(esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_66]), c_0_41]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_83,plain,(r2_hidden(X1,esk8_3(X2,X1,X3))|v2_struct_0(X2)|~m1_connsp_2(X3,X2,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_73, c_0_59])).
cnf(c_0_84,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_85,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))|r1_tarski(esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0),k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])]), c_0_78])).
cnf(c_0_86,plain,(r1_tarski(esk8_3(X1,X2,X3),X3)|v2_struct_0(X1)|~m1_connsp_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_79, c_0_59])).
fof(c_0_87, plain, (![X5]:![X6]:![X7]:![X8]:(~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|m1_subset_1(k8_relset_1(X5,X6,X7,X8),k1_zfmisc_1(X5)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relset_1])])])])).
cnf(c_0_88,negated_conjecture,(m1_connsp_2(X1,esk2_0,X2)|~r2_hidden(X2,esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)))|~r1_tarski(esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_89,negated_conjecture,(r2_hidden(esk4_0,esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_66]), c_0_41]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_90,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|r1_tarski(X1,k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0))|~r1_tarski(X1,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_91,negated_conjecture,(r1_tarski(esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)),esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_66]), c_0_41]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_92,plain,(m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_93,negated_conjecture,(~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)|~m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0),esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_94, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_95,negated_conjecture,(m1_connsp_2(X1,esk2_0,esk4_0)|~r1_tarski(esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_96,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))|r1_tarski(esk8_3(esk2_0,esk4_0,esk6_5(esk2_0,esk1_0,esk3_0,esk4_0,esk5_0)),k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_97,negated_conjecture,(m1_subset_1(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_92, c_0_23])).
cnf(c_0_98,negated_conjecture,(~m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0),esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_50])])).
cnf(c_0_99,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_100,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97])]), c_0_98])).
fof(c_0_101, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_102,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_33])).
cnf(c_0_103,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_104,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_29])])).
cnf(c_0_105,negated_conjecture,(~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_104]), c_0_34])).
cnf(c_0_106,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_103]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 107
# Proof object clause steps            : 76
# Proof object formula steps           : 31
# Proof object conjectures             : 53
# Proof object clause conjectures      : 50
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 107
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 156
# ...of these trivial                  : 3
# ...subsumed                          : 11
# ...remaining for further processing  : 142
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 25
# Generated clauses                    : 156
# ...of the previous two non-trivial   : 151
# Contextual simplify-reflections      : 7
# Paramodulations                      : 156
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 78
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 36
# ...number of literals in the above   : 124
# Current number of archived formulas  : 0
# Current number of archived clauses   : 64
# Clause-clause subsumption calls (NU) : 2287
# Rec. Clause-clause subsumption calls : 525
# Non-unit clause-clause subsumptions  : 22
# Unit Clause-clause subsumption calls : 39
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10668

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 45168 pages
