# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_r3_lattices, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3)))), file('out/t42_lattice3__4_premises__rerun.E_input', redefinition_r3_lattices)).
fof(cc1_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1)))))))))), file('out/t42_lattice3__4_premises__rerun.E_input', cc1_lattices)).
fof(d16_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4))))))), file('out/t42_lattice3__4_premises__rerun.E_input', d16_lattice3)).
fof(t34_lattice3, axiom, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2)))))))), file('out/t42_lattice3__4_premises__rerun.E_input', t34_lattice3)).
fof(t42_lattice3, conjecture, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r3_lattice3(X1,X2,X3))=>k16_lattice3(X1,X3)=X2)))), file('out/t42_lattice3__4_premises__rerun.E_input', t42_lattice3)).
fof(c_0_5, plain, (![X4]:![X5]:![X6]:(((~r3_lattices(X4,X5,X6)|r1_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_lattices(X4,X5,X6)|r3_lattices(X4,X5,X6))|((v2_struct_0(X4)|(~v6_lattices(X4)|(~v8_lattices(X4)|(~v9_lattices(X4)|~l3_lattices(X4)))))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
fof(c_0_6, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v4_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v5_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v6_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v7_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&(((v8_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))&((v9_lattices(X2)|(v2_struct_0(X2)|~v10_lattices(X2)))|~l3_lattices(X2))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:((((~r3_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|r1_lattices(X5,X6,X8))))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&((((r2_hidden(esk1_3(X5,X6,X7),X7)|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5)))&(((~r1_lattices(X5,X6,esk1_3(X5,X6,X7))|r3_lattice3(X5,X6,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|~l3_lattices(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])])).
fof(c_0_8, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:(((((r3_lattice3(X5,X6,X7)|X6!=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r3_lattice3(X5,X8,X7)|r3_lattices(X5,X8,X6)))|X6!=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))))&(((((m1_subset_1(esk2_3(X5,X6,X7),u1_struct_0(X5))|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&(((((r3_lattice3(X5,esk2_3(X5,X6,X7),X7)|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5)))))&((((~r3_lattices(X5,esk2_3(X5,X6,X7),X6)|~r3_lattice3(X5,X6,X7))|X6=k16_lattice3(X5,X7))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])])).
fof(c_0_9, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r3_lattice3(X1,X2,X3))=>k16_lattice3(X1,X3)=X2))))), inference(assume_negation,[status(cth)],[t42_lattice3])).
cnf(c_0_10,plain,(v2_struct_0(X2)|r3_lattices(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v9_lattices(X2)|~v8_lattices(X2)|~v6_lattices(X2)|~r1_lattices(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,plain,(v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,plain,(v2_struct_0(X1)|v6_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,plain,(v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,plain,(v2_struct_0(X1)|r1_lattices(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r3_lattice3(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,plain,(v2_struct_0(X1)|X2=k16_lattice3(X1,X3)|r3_lattice3(X1,esk2_3(X1,X2,X3),X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r3_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16,plain,(v2_struct_0(X1)|X2=k16_lattice3(X1,X3)|m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r3_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_17, negated_conjecture, (((~v2_struct_0(esk3_0)&(v10_lattices(esk3_0)&(v4_lattice3(esk3_0)&l3_lattices(esk3_0))))&(m1_subset_1(esk4_0,u1_struct_0(esk3_0))&((r2_hidden(esk4_0,esk5_0)&r3_lattice3(esk3_0,esk4_0,esk5_0))&k16_lattice3(esk3_0,esk5_0)!=esk4_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_18,plain,(v2_struct_0(X1)|X2=k16_lattice3(X1,X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r3_lattice3(X1,X2,X3)|~r3_lattices(X1,esk2_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19,plain,(r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~v10_lattices(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])).
cnf(c_0_20,plain,(X1=k16_lattice3(X2,X3)|r1_lattices(X2,esk2_3(X2,X1,X3),X4)|v2_struct_0(X2)|~v4_lattice3(X2)|~r2_hidden(X4,X3)|~r3_lattice3(X2,X1,X3)|~v10_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_21,negated_conjecture,(r2_hidden(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,plain,(X1=k16_lattice3(X2,X3)|v2_struct_0(X2)|~v4_lattice3(X2)|~r3_lattice3(X2,X1,X3)|~v10_lattices(X2)|~r1_lattices(X2,esk2_3(X2,X1,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_16])).
cnf(c_0_23,negated_conjecture,(X1=k16_lattice3(X2,esk5_0)|r1_lattices(X2,esk2_3(X2,X1,esk5_0),esk4_0)|v2_struct_0(X2)|~v4_lattice3(X2)|~r3_lattice3(X2,X1,esk5_0)|~v10_lattices(X2)|~m1_subset_1(esk4_0,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24,negated_conjecture,(k16_lattice3(esk3_0,esk5_0)!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,negated_conjecture,(k16_lattice3(X1,esk5_0)=esk4_0|v2_struct_0(X1)|~v4_lattice3(X1)|~r3_lattice3(X1,esk4_0,esk5_0)|~v10_lattices(X1)|~m1_subset_1(esk4_0,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26,negated_conjecture,(v4_lattice3(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,negated_conjecture,(r3_lattice3(esk3_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28,negated_conjecture,(v10_lattices(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30,negated_conjecture,(l3_lattices(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30])]), c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 22
# Proof object formula steps           : 11
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 25
# Processed clauses                    : 70
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 70
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 35
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 8
# Paramodulations                      : 33
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 42
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 35
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 697
# Rec. Clause-clause subsumption calls : 65
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3674

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45312 pages
