# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t29_ordinal1, axiom, (![X1]:(v3_ordinal1(X1)=>(~((~(v4_ordinal1(X1))&![X2]:(v3_ordinal1(X2)=>~(X1=k1_ordinal1(X2)))))&~((?[X2]:(v3_ordinal1(X2)&X1=k1_ordinal1(X2))&v4_ordinal1(X1)))))), file('out/t42_ordinal1__11_premises__rerun__rerun.E_input', t29_ordinal1)).
fof(d1_ordinal1, axiom, (![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), file('out/t42_ordinal1__11_premises__rerun__rerun.E_input', d1_ordinal1)).
fof(t42_ordinal1, conjecture, (![X1]:(v3_ordinal1(X1)=>(~((~(v4_ordinal1(X1))&![X2]:(v3_ordinal1(X2)=>~(X1=k1_ordinal1(X2)))))&~((?[X2]:(v3_ordinal1(X2)&X1=k1_ordinal1(X2))&v4_ordinal1(X1)))))), file('out/t42_ordinal1__11_premises__rerun__rerun.E_input', t42_ordinal1)).
fof(c_0_3, plain, (![X3]:![X5]:((((v3_ordinal1(esk1_1(X3))|v4_ordinal1(X3))|~v3_ordinal1(X3))&((X3=k1_ordinal1(esk1_1(X3))|v4_ordinal1(X3))|~v3_ordinal1(X3)))&(((~v3_ordinal1(X5)|X3!=k1_ordinal1(X5))|~v4_ordinal1(X3))|~v3_ordinal1(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_ordinal1])])])])])])])])).
fof(c_0_4, plain, (![X2]:k1_ordinal1(X2)=k2_xboole_0(X2,k1_tarski(X2))), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_5, negated_conjecture, (~(![X1]:(v3_ordinal1(X1)=>(~((~(v4_ordinal1(X1))&![X2]:(v3_ordinal1(X2)=>~(X1=k1_ordinal1(X2)))))&~((?[X2]:(v3_ordinal1(X2)&X1=k1_ordinal1(X2))&v4_ordinal1(X1))))))), inference(assume_negation,[status(cth)],[t42_ordinal1])).
cnf(c_0_6,plain,(~v3_ordinal1(X1)|~v4_ordinal1(X1)|X1!=k1_ordinal1(X2)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_7,plain,(k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_8, negated_conjecture, (![X4]:(v3_ordinal1(esk2_0)&((((v3_ordinal1(esk3_0)|~v4_ordinal1(esk2_0))&(esk2_0=k1_ordinal1(esk3_0)|~v4_ordinal1(esk2_0)))&(v4_ordinal1(esk2_0)|~v4_ordinal1(esk2_0)))&(((v3_ordinal1(esk3_0)|(~v3_ordinal1(X4)|esk2_0!=k1_ordinal1(X4)))&(esk2_0=k1_ordinal1(esk3_0)|(~v3_ordinal1(X4)|esk2_0!=k1_ordinal1(X4))))&(v4_ordinal1(esk2_0)|(~v3_ordinal1(X4)|esk2_0!=k1_ordinal1(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])])])).
cnf(c_0_9,plain,(X1!=k2_xboole_0(X2,k1_tarski(X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v4_ordinal1(X1)), inference(rw,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10,negated_conjecture,(esk2_0=k1_ordinal1(esk3_0)|~v4_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11,negated_conjecture,(v4_ordinal1(esk2_0)|esk2_0!=k1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(v4_ordinal1(X1)|X1=k1_ordinal1(esk1_1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_13,plain,(~v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_14,negated_conjecture,(esk2_0=k2_xboole_0(esk3_0,k1_tarski(esk3_0))|~v4_ordinal1(esk2_0)), inference(rw,[status(thm)],[c_0_10, c_0_7])).
cnf(c_0_15,negated_conjecture,(v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16,negated_conjecture,(v3_ordinal1(esk3_0)|~v4_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17,negated_conjecture,(v4_ordinal1(esk2_0)|esk2_0!=k2_xboole_0(X1,k1_tarski(X1))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_11, c_0_7])).
cnf(c_0_18,plain,(X1=k2_xboole_0(esk1_1(X1),k1_tarski(esk1_1(X1)))|v4_ordinal1(X1)|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_12, c_0_7])).
cnf(c_0_19,negated_conjecture,(~v4_ordinal1(esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_20,negated_conjecture,(~v3_ordinal1(esk1_1(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])]), c_0_15])]), c_0_19])).
cnf(c_0_21,plain,(v4_ordinal1(X1)|v3_ordinal1(esk1_1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_22,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_15])]), c_0_19]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 16
# Proof object formula steps           : 7
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 9
# Processed clauses                    : 22
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 21
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 13
# ...of the previous two non-trivial   : 6
# Contextual simplify-reflections      : 2
# Paramodulations                      : 9
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 11
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 8
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 4
# Current number of archived formulas  : 0
# Current number of archived clauses   : 10
# Clause-clause subsumption calls (NU) : 3
# Rec. Clause-clause subsumption calls : 3
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 837

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45236 pages
