# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t43_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0))))))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t43_zfmisc_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t7_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t71_enumset1)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', d1_tarski)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', d3_tarski)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t95_xboole_1)).
fof(t51_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t51_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', d1_xboole_0)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t91_xboole_1)).
fof(idempotence_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X1)=X1), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', idempotence_k2_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X1)=X1), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', idempotence_k3_xboole_0)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t12_xboole_1)).
fof(t92_xboole_1, axiom, (![X1]:k5_xboole_0(X1,X1)=k1_xboole_0), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t92_xboole_1)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t5_boole)).
fof(t7_xboole_0, axiom, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1))))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t7_xboole_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', d10_xboole_0)).
fof(l1_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', l1_zfmisc_1)).
fof(t75_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1)), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', t75_enumset1)).
fof(commutativity_k5_xboole_0, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), file('out/t43_zfmisc_1__26_premises__rerun__rerun.E_input', commutativity_k5_xboole_0)).
fof(c_0_22, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_23, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_24, negated_conjecture, (~(![X1]:![X2]:![X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))))), inference(assume_negation,[status(cth)],[t43_zfmisc_1])).
fof(c_0_25, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_26,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_28, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_29, negated_conjecture, ((k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)&((esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0))&((esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0))&(esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
cnf(c_0_30,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31,plain,(k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33,negated_conjecture,(k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
fof(c_0_35, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_36,plain,(r1_tarski(X1,k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_37,negated_conjecture,(k1_tarski(esk1_0)=k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_31]), c_0_32])).
fof(c_0_38, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_39,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41,negated_conjecture,(r1_tarski(esk2_0,k1_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_42, plain, (![X3]:![X4]:k2_xboole_0(k3_xboole_0(X3,X4),k4_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t51_xboole_1])).
fof(c_0_43, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_44,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_46,negated_conjecture,(r2_hidden(X1,k1_tarski(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_47, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk6_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_48,plain,(k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_31]), c_0_32])).
fof(c_0_51, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
fof(c_0_52, plain, (![X3]:k2_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_53, plain, (![X3]:k3_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_54, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_55,negated_conjecture,(esk1_0=X1|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56,plain,(v1_xboole_0(X1)|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57,plain,(k3_tarski(k2_enumset1(k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2))),k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2))),k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2))),k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2))))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_31]), c_0_49]), c_0_32]), c_0_50]), c_0_50]), c_0_50]), c_0_50])).
cnf(c_0_58,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59,plain,(k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_60, plain, (![X2]:k5_xboole_0(X2,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_61, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_62,plain,(k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_65,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_66,negated_conjecture,(esk6_1(esk2_0)=esk1_0|v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_67, plain, (![X3]:(X3=k1_xboole_0|r2_hidden(esk7_1(X3),X3))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])])])).
cnf(c_0_68,plain,(k3_tarski(k2_enumset1(k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)))),k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)))),k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)))),k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)))))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58]), c_0_58]), c_0_58]), c_0_58])).
cnf(c_0_69,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_31]), c_0_32])).
cnf(c_0_70,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72,plain,(k5_xboole_0(k5_xboole_0(X1,X1),k3_tarski(k2_enumset1(X1,X1,X1,X1)))=X1), inference(rw,[status(thm)],[c_0_62, c_0_50])).
fof(c_0_73, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_74,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_31]), c_0_32])).
cnf(c_0_75,plain,(r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_76,negated_conjecture,(v1_xboole_0(esk2_0)|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_56, c_0_66])).
cnf(c_0_77,plain,(r2_hidden(esk7_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_78,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_71]), c_0_70]), c_0_71]), c_0_70]), c_0_71]), c_0_70]), c_0_71]), c_0_70])).
cnf(c_0_79,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_70]), c_0_69])).
cnf(c_0_80,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
fof(c_0_81, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])])])).
cnf(c_0_82,negated_conjecture,(k1_tarski(esk1_0)=esk3_0|~r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_74])).
cnf(c_0_83,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_84,negated_conjecture,(esk7_1(esk2_0)=esk1_0|k1_xboole_0=esk2_0), inference(spm,[status(thm)],[c_0_55, c_0_77])).
cnf(c_0_85,plain,(k3_tarski(k2_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_78]), c_0_79]), c_0_70]), c_0_79]), c_0_70]), c_0_79]), c_0_70]), c_0_79]), c_0_70]), c_0_71])).
cnf(c_0_86,negated_conjecture,(k1_tarski(esk1_0)=esk2_0|~r1_tarski(k1_tarski(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_80, c_0_41])).
cnf(c_0_87,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_88,negated_conjecture,(esk3_0!=k1_tarski(esk1_0)|esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_89,negated_conjecture,(k1_tarski(esk1_0)=esk3_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_90,negated_conjecture,(k1_xboole_0=esk2_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_84])).
cnf(c_0_91,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_36, c_0_85])).
cnf(c_0_92,negated_conjecture,(k1_tarski(esk1_0)=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93,negated_conjecture,(r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90])).
cnf(c_0_94,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k1_tarski(X3))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_40, c_0_87])).
cnf(c_0_95,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_91])).
cnf(c_0_96,negated_conjecture,(k1_tarski(esk1_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_97,plain,(k1_tarski(X1)=k1_xboole_0|r2_hidden(esk7_1(k1_tarski(X1)),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_94, c_0_77])).
cnf(c_0_98,negated_conjecture,(esk3_0!=k1_tarski(esk1_0)|esk2_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_99,plain,(v1_xboole_0(k1_xboole_0)|r2_hidden(esk6_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_95, c_0_56])).
cnf(c_0_100,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_101,plain,(X1=k1_tarski(X2)|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_80, c_0_87])).
cnf(c_0_102,negated_conjecture,(r1_tarski(esk2_0,X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_87, c_0_96])).
cnf(c_0_103,negated_conjecture,(esk7_1(k1_tarski(X1))=esk1_0|k1_tarski(X1)=k1_xboole_0|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_97])).
cnf(c_0_104,negated_conjecture,(esk3_0!=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_98, c_0_92])).
cnf(c_0_105,plain,(X1=esk6_1(k1_xboole_0)|v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_45, c_0_99])).
fof(c_0_106, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X8,X6,X7,X5)), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_107,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_100])])).
cnf(c_0_108,negated_conjecture,(k1_tarski(X1)=esk2_0|~r2_hidden(esk1_0,k1_tarski(X1))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_109,negated_conjecture,(k1_tarski(X1)=k1_xboole_0|r2_hidden(esk1_0,k1_tarski(X1))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_103])).
cnf(c_0_110,negated_conjecture,(esk3_0!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_93])])).
cnf(c_0_111,plain,(X1=X2|v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_105, c_0_105])).
cnf(c_0_112,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_70]), c_0_79])).
cnf(c_0_113,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
fof(c_0_114, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k5_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_115,plain,(~v1_xboole_0(k1_tarski(X1))), inference(spm,[status(thm)],[c_0_64, c_0_107])).
cnf(c_0_116,negated_conjecture,(k1_tarski(X1)=k1_xboole_0|k1_tarski(X1)=esk2_0|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_117,negated_conjecture,(v1_xboole_0(k1_xboole_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_118,plain,(k3_tarski(k2_enumset1(k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X1))),k5_xboole_0(X2,k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X1)))),k5_xboole_0(X2,k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X1)))),k5_xboole_0(X2,k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X1))))))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_112]), c_0_113])).
cnf(c_0_119,negated_conjecture,(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))=esk2_0), inference(rw,[status(thm)],[c_0_37, c_0_96])).
cnf(c_0_120,plain,(k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_121,negated_conjecture,(k1_tarski(X1)=esk2_0|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117])])).
cnf(c_0_122,negated_conjecture,(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k5_xboole_0(esk2_0,esk3_0)))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120]), c_0_120]), c_0_112]), c_0_120]), c_0_112]), c_0_120]), c_0_112]), c_0_113])).
cnf(c_0_123,negated_conjecture,(esk3_0!=k1_xboole_0|esk2_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_124,negated_conjecture,(X1=esk2_0|~r2_hidden(X2,X1)|~r1_tarski(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_121]), c_0_40])).
cnf(c_0_125,negated_conjecture,(r1_tarski(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_122])).
cnf(c_0_126,negated_conjecture,(esk3_0!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_123, c_0_92])).
cnf(c_0_127,negated_conjecture,(~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_110])).
cnf(c_0_128,negated_conjecture,(esk3_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_93])])).
cnf(c_0_129,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_77]), c_0_128]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 130
# Proof object clause steps            : 85
# Proof object formula steps           : 45
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 71
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 29
# Processed clauses                    : 528
# ...of these trivial                  : 23
# ...subsumed                          : 301
# ...remaining for further processing  : 204
# Other redundant clauses eliminated   : 24
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 20
# Generated clauses                    : 1856
# ...of the previous two non-trivial   : 1524
# Contextual simplify-reflections      : 62
# Paramodulations                      : 1828
# Factorizations                       : 5
# Equation resolutions                 : 24
# Current number of processed clauses  : 136
#    Positive orientable unit clauses  : 31
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 86
# Current number of unprocessed clauses: 372
# ...number of literals in the above   : 867
# Current number of archived formulas  : 0
# Current number of archived clauses   : 69
# Clause-clause subsumption calls (NU) : 2794
# Rec. Clause-clause subsumption calls : 2353
# Non-unit clause-clause subsumptions  : 308
# Unit Clause-clause subsumption calls : 234
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 96
# BW rewrite match successes           : 61
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15961

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 45424 pages
