# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_tmap_1(X1,X2,X3,X4)<=>![X5]:(m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))=>?[X6]:(m1_connsp_2(X6,X1,X4)&r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X6),X5))))))))), file('out/t44_tmap_1__2_premises__rerun.E_input', d2_tmap_1)).
fof(d1_borsuk_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))=>?[X6]:(m1_connsp_2(X6,X1,X4)&r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X6),X5))))))))), file('out/t44_tmap_1__2_premises__rerun.E_input', d1_borsuk_1)).
fof(t44_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(v5_pre_topc(X3,X2,X1)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>r1_tmap_1(X2,X1,X3,X4))))))), file('out/t44_tmap_1__2_premises__rerun.E_input', t44_tmap_1)).
fof(c_0_3, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X14]:((((((((m1_connsp_2(esk1_5(X7,X8,X9,X10,X11),X7,X10)|~m1_connsp_2(X11,X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10)))|~r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((r1_tarski(k7_relset_1(u1_struct_0(X7),u1_struct_0(X8),X9,esk1_5(X7,X8,X9,X10,X11)),X11)|~m1_connsp_2(X11,X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10)))|~r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))))&((((((m1_connsp_2(esk2_4(X7,X8,X9,X10),X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10))|r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((~m1_connsp_2(X14,X7,X10)|~r1_tarski(k7_relset_1(u1_struct_0(X7),u1_struct_0(X8),X9,X14),esk2_4(X7,X8,X9,X10)))|r1_tmap_1(X7,X8,X9,X10))|~m1_subset_1(X10,u1_struct_0(X7)))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_tmap_1])])])])])])])])).
fof(c_0_4, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X15]:((((((((m1_connsp_2(esk3_5(X7,X8,X9,X10,X11),X7,X10)|~m1_connsp_2(X11,X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10)))|~m1_subset_1(X10,u1_struct_0(X7)))|~v5_pre_topc(X9,X7,X8))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((r1_tarski(k7_relset_1(u1_struct_0(X7),u1_struct_0(X8),X9,esk3_5(X7,X8,X9,X10,X11)),X11)|~m1_connsp_2(X11,X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,X10)))|~m1_subset_1(X10,u1_struct_0(X7)))|~v5_pre_topc(X9,X7,X8))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))))&(((((m1_subset_1(esk4_3(X7,X8,X9),u1_struct_0(X7))|v5_pre_topc(X9,X7,X8))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((m1_connsp_2(esk5_3(X7,X8,X9),X8,k3_funct_2(u1_struct_0(X7),u1_struct_0(X8),X9,esk4_3(X7,X8,X9)))|v5_pre_topc(X9,X7,X8))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((~m1_connsp_2(X15,X7,esk4_3(X7,X8,X9))|~r1_tarski(k7_relset_1(u1_struct_0(X7),u1_struct_0(X8),X9,X15),esk5_3(X7,X8,X9)))|v5_pre_topc(X9,X7,X8))|(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8)))))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_borsuk_1])])])])])])])])).
cnf(c_0_5,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tmap_1(X1,X2,X3,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X5),esk2_4(X1,X2,X3,X4))|~m1_connsp_2(X5,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk3_5(X1,X2,X3,X4,X5)),X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_connsp_2(esk3_5(X1,X2,X3,X4,X5),X1,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tmap_1(X1,X2,X3,X4)|m1_connsp_2(esk2_4(X1,X2,X3,X4),X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_9, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(v5_pre_topc(X3,X2,X1)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>r1_tmap_1(X2,X1,X3,X4)))))))), inference(assume_negation,[status(cth)],[t44_tmap_1])).
cnf(c_0_10,plain,(r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X1)|v2_struct_0(X2)|~v5_pre_topc(X3,X1,X2)|~m1_connsp_2(esk2_4(X1,X2,X3,X4),X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X5))|~m1_connsp_2(esk3_5(X1,X2,X3,X5,esk2_4(X1,X2,X3,X4)),X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_11,plain,(m1_connsp_2(esk3_5(X1,X2,X3,X4,esk2_4(X1,X2,X3,X4)),X1,X4)|r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X1)|v2_struct_0(X2)|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
fof(c_0_12, negated_conjecture, (![X9]:((~v2_struct_0(esk6_0)&(v2_pre_topc(esk6_0)&l1_pre_topc(esk6_0)))&((~v2_struct_0(esk7_0)&(v2_pre_topc(esk7_0)&l1_pre_topc(esk7_0)))&((v1_funct_1(esk8_0)&(v1_funct_2(esk8_0,u1_struct_0(esk7_0),u1_struct_0(esk6_0))&m1_subset_1(esk8_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk7_0),u1_struct_0(esk6_0))))))&(((m1_subset_1(esk9_0,u1_struct_0(esk7_0))|~v5_pre_topc(esk8_0,esk7_0,esk6_0))&(~r1_tmap_1(esk7_0,esk6_0,esk8_0,esk9_0)|~v5_pre_topc(esk8_0,esk7_0,esk6_0)))&(v5_pre_topc(esk8_0,esk7_0,esk6_0)|(~m1_subset_1(X9,u1_struct_0(esk7_0))|r1_tmap_1(esk7_0,esk6_0,esk8_0,X9)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])])])).
cnf(c_0_13,plain,(r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X2)|v2_struct_0(X1)|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_8]), c_0_11])).
cnf(c_0_14,negated_conjecture,(m1_subset_1(esk8_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk7_0),u1_struct_0(esk6_0))))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15,negated_conjecture,(v1_funct_2(esk8_0,u1_struct_0(esk7_0),u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(v1_funct_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(l1_pre_topc(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,negated_conjecture,(l1_pre_topc(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(v2_pre_topc(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,negated_conjecture,(v2_pre_topc(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,negated_conjecture,(~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,negated_conjecture,(~v2_struct_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(r1_tmap_1(esk7_0,esk6_0,esk8_0,X1)|v5_pre_topc(esk8_0,esk7_0,esk6_0)|~m1_subset_1(X1,u1_struct_0(esk7_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,X1,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),esk5_3(X1,X2,X3))|~m1_connsp_2(X4,X1,esk4_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_25,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk1_5(X1,X2,X3,X4,X5)),X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_tmap_1(X1,X2,X3,X4)|~m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_26,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_connsp_2(esk1_5(X1,X2,X3,X4,X5),X1,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_tmap_1(X1,X2,X3,X4)|~m1_connsp_2(X5,X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_27,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,X1,X2)|m1_connsp_2(esk5_3(X1,X2,X3),X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,esk4_3(X1,X2,X3)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_28,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,X1,X2)|m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_29,negated_conjecture,(~v5_pre_topc(esk8_0,esk7_0,esk6_0)|~r1_tmap_1(esk7_0,esk6_0,esk8_0,esk9_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30,negated_conjecture,(r1_tmap_1(esk7_0,esk6_0,esk8_0,X1)|~m1_subset_1(X1,u1_struct_0(esk7_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_31,negated_conjecture,(m1_subset_1(esk9_0,u1_struct_0(esk7_0))|~v5_pre_topc(esk8_0,esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32,plain,(v5_pre_topc(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_connsp_2(esk1_5(X2,X3,X1,X4,esk5_3(X2,X3,X1)),X2,esk4_3(X2,X3,X1))|~m1_connsp_2(esk5_3(X2,X3,X1),X3,k3_funct_2(u1_struct_0(X2),u1_struct_0(X3),X1,X4))|~r1_tmap_1(X2,X3,X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,plain,(v5_pre_topc(X1,X2,X3)|m1_connsp_2(esk1_5(X2,X3,X1,esk4_3(X2,X3,X1),esk5_3(X2,X3,X1)),X2,esk4_3(X2,X3,X1))|v2_struct_0(X2)|v2_struct_0(X3)|~r1_tmap_1(X2,X3,X1,esk4_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_34,negated_conjecture,(v5_pre_topc(esk8_0,esk7_0,esk6_0)|m1_subset_1(esk4_3(esk7_0,esk6_0,esk8_0),u1_struct_0(esk7_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21]), c_0_22])).
cnf(c_0_35,negated_conjecture,(~v5_pre_topc(esk8_0,esk7_0,esk6_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_36,plain,(v5_pre_topc(X1,X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tmap_1(X2,X3,X1,esk4_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_28]), c_0_27])).
cnf(c_0_37,negated_conjecture,(m1_subset_1(esk4_3(esk7_0,esk6_0,esk8_0),u1_struct_0(esk7_0))), inference(sr,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_30]), c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_37])]), c_0_35]), c_0_21]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 32
# Proof object formula steps           : 7
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 52
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 52
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 13
# Contextual simplify-reflections      : 6
# Paramodulations                      : 17
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 15
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 15
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 203
# Rec. Clause-clause subsumption calls : 12
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4619

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45352 pages
