# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_mcart_1, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_xtuple_0(k1_mcart_1(X1,X2,X3,X4),k2_mcart_1(X1,X2,X3,X4),k3_mcart_1(X1,X2,X3,X4))))))), file('out/t44_mcart_1__5_premises__rerun.E_input', t44_mcart_1)).
fof(redefinition_k1_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&(~(v1_xboole_0(X3))&m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3)))))=>k1_mcart_1(X1,X2,X3,X4)=k4_xtuple_0(X4))), file('out/t44_mcart_1__5_premises__rerun.E_input', redefinition_k1_mcart_1)).
fof(redefinition_k2_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&(~(v1_xboole_0(X3))&m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3)))))=>k2_mcart_1(X1,X2,X3,X4)=k5_xtuple_0(X4))), file('out/t44_mcart_1__5_premises__rerun.E_input', redefinition_k2_mcart_1)).
fof(redefinition_k3_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&(~(v1_xboole_0(X3))&m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3)))))=>k3_mcart_1(X1,X2,X3,X4)=k2_xtuple_0(X4))), file('out/t44_mcart_1__5_premises__rerun.E_input', redefinition_k3_mcart_1)).
fof(cc1_mcart_1, axiom, (![X1]:![X2]:![X3]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&~(v1_xboole_0(X3))))=>![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>v2_xtuple_0(X4)))), file('out/t44_mcart_1__5_premises__rerun.E_input', cc1_mcart_1)).
fof(rd7_xtuple_0, axiom, (![X1]:(v2_xtuple_0(X1)=>k3_xtuple_0(k4_xtuple_0(X1),k5_xtuple_0(X1),k2_xtuple_0(X1))=X1)), file('out/t44_mcart_1__5_premises__rerun.E_input', rd7_xtuple_0)).
fof(c_0_6, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_xtuple_0(k1_mcart_1(X1,X2,X3,X4),k2_mcart_1(X1,X2,X3,X4),k3_mcart_1(X1,X2,X3,X4)))))))), inference(assume_negation,[status(cth)],[t44_mcart_1])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|(v1_xboole_0(X6)|(v1_xboole_0(X7)|~m1_subset_1(X8,k3_zfmisc_1(X5,X6,X7)))))|k1_mcart_1(X5,X6,X7,X8)=k4_xtuple_0(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k1_mcart_1])])])).
fof(c_0_8, negated_conjecture, ((~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&esk4_0!=k3_xtuple_0(k1_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k2_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k3_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|(v1_xboole_0(X6)|(v1_xboole_0(X7)|~m1_subset_1(X8,k3_zfmisc_1(X5,X6,X7)))))|k2_mcart_1(X5,X6,X7,X8)=k5_xtuple_0(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k2_mcart_1])])])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|(v1_xboole_0(X6)|(v1_xboole_0(X7)|~m1_subset_1(X8,k3_zfmisc_1(X5,X6,X7)))))|k3_mcart_1(X5,X6,X7,X8)=k2_xtuple_0(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_mcart_1])])])).
cnf(c_0_11,plain,(k1_mcart_1(X1,X2,X3,X4)=k4_xtuple_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,negated_conjecture,(m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,negated_conjecture,(~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16,plain,(k2_mcart_1(X1,X2,X3,X4)=k5_xtuple_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,plain,(k3_mcart_1(X1,X2,X3,X4)=k2_xtuple_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_18, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|(v1_xboole_0(X6)|v1_xboole_0(X7)))|(~m1_subset_1(X8,k3_zfmisc_1(X5,X6,X7))|v2_xtuple_0(X8)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_mcart_1])])])])])])).
cnf(c_0_19,negated_conjecture,(esk4_0!=k3_xtuple_0(k1_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k2_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k3_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20,negated_conjecture,(k1_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k4_xtuple_0(esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14]), c_0_15])).
cnf(c_0_21,negated_conjecture,(k2_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k5_xtuple_0(esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_12]), c_0_13]), c_0_14]), c_0_15])).
cnf(c_0_22,negated_conjecture,(k3_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k2_xtuple_0(esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_12]), c_0_13]), c_0_14]), c_0_15])).
fof(c_0_23, plain, (![X2]:(~v2_xtuple_0(X2)|k3_xtuple_0(k4_xtuple_0(X2),k5_xtuple_0(X2),k2_xtuple_0(X2))=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rd7_xtuple_0])])).
cnf(c_0_24,plain,(v2_xtuple_0(X1)|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,negated_conjecture,(k3_xtuple_0(k4_xtuple_0(esk4_0),k5_xtuple_0(esk4_0),k2_xtuple_0(esk4_0))!=esk4_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])).
cnf(c_0_26,plain,(k3_xtuple_0(k4_xtuple_0(X1),k5_xtuple_0(X1),k2_xtuple_0(X1))=X1|~v2_xtuple_0(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27,negated_conjecture,(v2_xtuple_0(esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_12]), c_0_13]), c_0_14]), c_0_15])).
cnf(c_0_28,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 16
# Proof object formula steps           : 13
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 10
# Processed clauses                    : 15
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 15
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 5
# ...of the previous two non-trivial   : 5
# Contextual simplify-reflections      : 0
# Paramodulations                      : 5
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 14
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1060

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45256 pages
