# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t4_boole, axiom, (![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t4_boole)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t2_boole, axiom, (![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t2_boole)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t1_boole, axiom, (![X1]:k2_xboole_0(X1,k1_xboole_0)=X1), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t1_boole)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t71_enumset1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t98_xboole_1)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t48_xboole_1)).
fof(t112_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t112_xboole_1)).
fof(t16_xboole_1, axiom, (![X1]:![X2]:![X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t16_xboole_1)).
fof(t101_xboole_1, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2))), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t101_xboole_1)).
fof(t45_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1)))), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t45_xboole_1)).
fof(t45_zfmisc_1, conjecture, (![X1]:![X2]:(r1_tarski(k2_xboole_0(k1_tarski(X1),X2),X2)=>r2_hidden(X1,X2))), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t45_zfmisc_1)).
fof(t94_enumset1, axiom, (![X1]:k5_enumset1(X1,X1,X1,X1,X1,X1,X1)=k1_tarski(X1)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t94_enumset1)).
fof(t102_enumset1, axiom, (![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t102_enumset1)).
fof(t75_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1)), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', t75_enumset1)).
fof(l29_zfmisc_1, axiom, (![X1]:![X2]:(k3_xboole_0(X1,k1_tarski(X2))=k1_tarski(X2)=>r2_hidden(X2,X1))), file('out/t45_zfmisc_1__23_premises__rerun__rerun.E_input', l29_zfmisc_1)).
fof(c_0_19, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_20, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, plain, (![X2]:k4_xboole_0(k1_xboole_0,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_22, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_23, plain, (![X2]:k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_24, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_25, plain, (![X2]:k2_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_26,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_29, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
cnf(c_0_30,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_34, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_35,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36,plain,(k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_37,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39,plain,(k5_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40,plain,(k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_43,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X2))=k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_36]), c_0_31]), c_0_37])).
cnf(c_0_44,plain,(k5_xboole_0(k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45,plain,(k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_31]), c_0_31])).
cnf(c_0_46,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_40]), c_0_44])).
fof(c_0_47, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k3_xboole_0(X4,X5),k3_xboole_0(X6,X5))=k3_xboole_0(k5_xboole_0(X4,X6),X5)), inference(variable_rename,[status(thm)],[t112_xboole_1])).
cnf(c_0_48,plain,(k5_xboole_0(X1,k3_xboole_0(X1,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_32]), c_0_46])).
fof(c_0_49, plain, (![X4]:![X5]:![X6]:k3_xboole_0(k3_xboole_0(X4,X5),X6)=k3_xboole_0(X4,k3_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t16_xboole_1])).
fof(c_0_50, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k4_xboole_0(k2_xboole_0(X3,X4),k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t101_xboole_1])).
fof(c_0_51, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|X4=k2_xboole_0(X3,k4_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
cnf(c_0_52,plain,(k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53,plain,(k3_xboole_0(X1,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_48]), c_0_32]), c_0_46])).
cnf(c_0_54,plain,(k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_55,plain,(k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_56,plain,(X1=k2_xboole_0(X2,k4_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_57,plain,(k5_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_33])).
cnf(c_0_58,plain,(k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_33, c_0_54])).
fof(c_0_59, negated_conjecture, (~(![X1]:![X2]:(r1_tarski(k2_xboole_0(k1_tarski(X1),X2),X2)=>r2_hidden(X1,X2)))), inference(assume_negation,[status(cth)],[t45_zfmisc_1])).
cnf(c_0_60,plain,(k5_xboole_0(X1,X2)=k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_xboole_0(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_36]), c_0_31]), c_0_37]), c_0_37])).
cnf(c_0_61,plain,(X1=k3_tarski(k2_enumset1(X2,X2,X2,k5_xboole_0(X1,k3_xboole_0(X1,X2))))|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_36]), c_0_31]), c_0_37])).
cnf(c_0_62,plain,(k5_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_57, c_0_33])).
cnf(c_0_63,plain,(k5_xboole_0(k3_xboole_0(X1,X2),X2)=k3_xboole_0(X2,k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_33])).
cnf(c_0_64,plain,(k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_58, c_0_53])).
cnf(c_0_65,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_48, c_0_53])).
cnf(c_0_66,plain,(k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_54, c_0_53])).
fof(c_0_67, negated_conjecture, ((r1_tarski(k2_xboole_0(k1_tarski(esk1_0),esk2_0),esk2_0)&~r2_hidden(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_59])])])).
fof(c_0_68, plain, (![X2]:k5_enumset1(X2,X2,X2,X2,X2,X2,X2)=k1_tarski(X2)), inference(variable_rename,[status(thm)],[t94_enumset1])).
fof(c_0_69, plain, (![X4]:![X5]:![X6]:k1_enumset1(X4,X5,X6)=k1_enumset1(X6,X5,X4)), inference(variable_rename,[status(thm)],[t102_enumset1])).
cnf(c_0_70,plain,(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_xboole_0(X1,k3_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)))))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_33]), c_0_54])).
cnf(c_0_71,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k3_xboole_0(X2,k5_xboole_0(X2,X1))))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_72,plain,(k3_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X2,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_52]), c_0_65]), c_0_40]), c_0_57]), c_0_54]), c_0_66])).
cnf(c_0_73,negated_conjecture,(r1_tarski(k2_xboole_0(k1_tarski(esk1_0),esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74,plain,(k5_enumset1(X1,X1,X1,X1,X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_75, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X8,X6,X7,X5)), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_76,plain,(k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X2))=k5_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X2,X1)))), inference(rw,[status(thm)],[c_0_43, c_0_62])).
cnf(c_0_78,plain,(k5_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X2,X1)))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_54]), c_0_33]), c_0_66]), c_0_72]), c_0_46])).
cnf(c_0_79,negated_conjecture,(r1_tarski(k3_tarski(k2_enumset1(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),esk2_0)),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_36]), c_0_37]), c_0_74]), c_0_74]), c_0_74])).
cnf(c_0_80,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_81,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X3,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_37]), c_0_37])).
cnf(c_0_82,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_83,negated_conjecture,(r1_tarski(k3_tarski(k2_enumset1(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X1,X3,X2,X3)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85,plain,(k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X2,k3_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_33]), c_0_54])).
cnf(c_0_86,plain,(k3_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X1,X2)))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_62]), c_0_62]), c_0_62]), c_0_66])).
cnf(c_0_87,plain,(k3_xboole_0(X1,k5_xboole_0(X1,X2))=k5_xboole_0(X2,X1)|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_82]), c_0_53]), c_0_57])).
cnf(c_0_88,negated_conjecture,(r1_tarski(k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_89,plain,(k5_xboole_0(k3_xboole_0(X1,X2),X1)=k3_xboole_0(X1,k5_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_63, c_0_33])).
cnf(c_0_90,plain,(k3_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X1,X2)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_85, c_0_72])).
cnf(c_0_91,plain,(k3_xboole_0(X1,k5_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_78]), c_0_72])).
cnf(c_0_92,negated_conjecture,(k3_xboole_0(esk2_0,k5_xboole_0(esk2_0,k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))))=k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),esk2_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_93,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k3_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_89]), c_0_54]), c_0_90]), c_0_32]), c_0_46])).
cnf(c_0_94,negated_conjecture,(k3_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_91, c_0_88])).
cnf(c_0_95,negated_conjecture,(k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_92]), c_0_88])])).
cnf(c_0_96,plain,(k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X2,X1))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_72]), c_0_46])).
cnf(c_0_97,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k3_xboole_0(X2,k3_xboole_0(X1,X3))))=X1), inference(spm,[status(thm)],[c_0_93, c_0_58])).
cnf(c_0_98,negated_conjecture,(k3_xboole_0(esk2_0,k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))=k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_94]), c_0_46]), c_0_95]), c_0_33])).
cnf(c_0_99,plain,(k3_xboole_0(X1,k3_tarski(k2_enumset1(X1,X1,X1,X2)))=X1), inference(spm,[status(thm)],[c_0_96, c_0_77])).
cnf(c_0_100,negated_conjecture,(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k3_xboole_0(X1,k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))))=esk2_0), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_101,plain,(k3_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X1,X1)))=X1), inference(spm,[status(thm)],[c_0_99, c_0_81])).
cnf(c_0_102,negated_conjecture,(k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))=esk2_0), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_103,plain,(k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X2,X3))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_52, c_0_33])).
cnf(c_0_104,plain,(k5_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X2,X1)))=k3_tarski(k2_enumset1(X2,X2,X1,X1))), inference(spm,[status(thm)],[c_0_77, c_0_81])).
cnf(c_0_105,negated_conjecture,(k3_xboole_0(esk2_0,k5_xboole_0(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))=k5_xboole_0(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_102]), c_0_64]), c_0_57])).
fof(c_0_106, plain, (![X3]:![X4]:(k3_xboole_0(X3,k1_tarski(X4))!=k1_tarski(X4)|r2_hidden(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l29_zfmisc_1])])).
cnf(c_0_107,plain,(k3_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(k3_xboole_0(X1,X2),X3)))=k3_xboole_0(X1,k3_xboole_0(k5_xboole_0(X1,X3),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_54]), c_0_103]), c_0_57]), c_0_54])).
cnf(c_0_108,negated_conjecture,(k3_tarski(k2_enumset1(esk2_0,esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))=k5_xboole_0(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_109,plain,(r2_hidden(X1,X2)|k3_xboole_0(X2,k1_tarski(X1))!=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_110,plain,(k3_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(X1,k5_xboole_0(X2,X1))))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_89]), c_0_54]), c_0_54]), c_0_64]), c_0_107]), c_0_62]), c_0_54]), c_0_33]), c_0_66])).
cnf(c_0_111,negated_conjecture,(k3_xboole_0(esk2_0,k5_xboole_0(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))=k5_xboole_0(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_108]), c_0_108])).
cnf(c_0_112,negated_conjecture,(k3_xboole_0(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))=k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_96, c_0_105])).
cnf(c_0_113,plain,(r2_hidden(X1,X2)|k3_xboole_0(X2,k5_enumset1(X1,X1,X1,X1,X1,X1,X1))!=k5_enumset1(X1,X1,X1,X1,X1,X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_74]), c_0_74])).
cnf(c_0_114,negated_conjecture,(k3_xboole_0(esk2_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))=k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112])).
cnf(c_0_115,negated_conjecture,(~r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_116,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_115]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 117
# Proof object clause steps            : 78
# Proof object formula steps           : 39
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 87
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 15
# Processed clauses                    : 2708
# ...of these trivial                  : 290
# ...subsumed                          : 2001
# ...remaining for further processing  : 417
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 74
# Generated clauses                    : 25268
# ...of the previous two non-trivial   : 19719
# Contextual simplify-reflections      : 780
# Paramodulations                      : 25268
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 323
#    Positive orientable unit clauses  : 156
#    Positive unorientable unit clauses: 13
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 152
# Current number of unprocessed clauses: 13862
# ...number of literals in the above   : 20969
# Current number of archived formulas  : 0
# Current number of archived clauses   : 99
# Clause-clause subsumption calls (NU) : 13636
# Rec. Clause-clause subsumption calls : 13247
# Non-unit clause-clause subsumptions  : 2703
# Unit Clause-clause subsumption calls : 119
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1497
# BW rewrite match successes           : 188
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 484232

# -------------------------------------------------
# User time                : 0.340 s
# System time              : 0.008 s
# Total time               : 0.348 s
# Maximum resident set size: 45276 pages
