# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.004 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t47_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>(r1_tmap_1(X2,X1,X3,X4)<=>![X5]:(m1_connsp_2(X5,X1,k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4))=>m1_connsp_2(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X5),X2,X4)))))))), file('out/t47_tmap_1__148_premises__rerun__rerun.E_input', t47_tmap_1)).
fof(t42_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>(r1_tmap_1(X2,X1,X3,X4)<=>![X5]:(m1_connsp_2(X5,X1,k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4))=>m1_connsp_2(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X5),X2,X4)))))))), file('out/t47_tmap_1__148_premises__rerun__rerun.E_input', t42_tmap_1)).
fof(c_0_2, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>(r1_tmap_1(X2,X1,X3,X4)<=>![X5]:(m1_connsp_2(X5,X1,k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4))=>m1_connsp_2(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X5),X2,X4))))))))), inference(assume_negation,[status(cth)],[t47_tmap_1])).
fof(c_0_3, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((((((~r1_tmap_1(X7,X6,X8,X9)|(~m1_connsp_2(X10,X6,k3_funct_2(u1_struct_0(X7),u1_struct_0(X6),X8,X9))|m1_connsp_2(k8_relset_1(u1_struct_0(X7),u1_struct_0(X6),X8,X10),X7,X9)))|~m1_subset_1(X9,u1_struct_0(X7)))|(~v1_funct_1(X8)|(~v1_funct_2(X8,u1_struct_0(X7),u1_struct_0(X6))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X6)))))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((((((m1_connsp_2(esk6_4(X6,X7,X8,X9),X6,k3_funct_2(u1_struct_0(X7),u1_struct_0(X6),X8,X9))|r1_tmap_1(X7,X6,X8,X9))|~m1_subset_1(X9,u1_struct_0(X7)))|(~v1_funct_1(X8)|(~v1_funct_2(X8,u1_struct_0(X7),u1_struct_0(X6))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X6)))))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(((((~m1_connsp_2(k8_relset_1(u1_struct_0(X7),u1_struct_0(X6),X8,esk6_4(X6,X7,X8,X9)),X7,X9)|r1_tmap_1(X7,X6,X8,X9))|~m1_subset_1(X9,u1_struct_0(X7)))|(~v1_funct_1(X8)|(~v1_funct_2(X8,u1_struct_0(X7),u1_struct_0(X6))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X6)))))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_tmap_1])])])])])])])])).
fof(c_0_4, negated_conjecture, (![X11]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(((m1_connsp_2(esk5_0,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0))&(~m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0),esk2_0,esk4_0)|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)))&(r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)|(~m1_connsp_2(X11,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))|m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X11),esk2_0,esk4_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_2])])])])])])])])).
cnf(c_0_5,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tmap_1(X2,X1,X3,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_connsp_2(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,esk6_4(X1,X2,X3,X4)),X2,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1),esk2_0,esk4_0)|r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)|~m1_connsp_2(X1,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13,negated_conjecture,(v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)|~m1_connsp_2(esk6_4(esk1_0,esk2_0,esk3_0,esk4_0),esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_15]), c_0_16])).
cnf(c_0_18,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tmap_1(X2,X1,X3,X4)|m1_connsp_2(esk6_4(X1,X2,X3,X4),X1,k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_19,negated_conjecture,(m1_connsp_2(esk5_0,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_15]), c_0_16])).
cnf(c_0_21,negated_conjecture,(~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0)|~m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0),esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_connsp_2(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X5),X2,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_connsp_2(X5,X1,k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4))|~r1_tmap_1(X2,X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_23,negated_conjecture,(m1_connsp_2(esk5_0,esk1_0,k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20])])).
cnf(c_0_24,negated_conjecture,(~m1_connsp_2(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_0),esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_20])])).
cnf(c_0_25,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_20]), c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_24]), c_0_15]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 26
# Proof object clause steps            : 21
# Proof object formula steps           : 5
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 36
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 36
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 5
# ...of the previous two non-trivial   : 4
# Contextual simplify-reflections      : 0
# Paramodulations                      : 5
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 16
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 11
# Rec. Clause-clause subsumption calls : 3
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2005

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.004 s
# Total time               : 0.008 s
# Maximum resident set size: 43296 pages
