# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t47_yellow_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:![X3]:((r1_yellow_0(X1,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)))), file('out/t47_yellow_0__3_premises__rerun.E_input', t47_yellow_0)).
fof(d9_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))))), file('out/t47_yellow_0__3_premises__rerun.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, (![X1]:![X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1)))), file('out/t47_yellow_0__3_premises__rerun.E_input', dt_k1_yellow_0)).
fof(t46_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:![X3]:((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))&r1_yellow_0(X1,X2))=>r1_yellow_0(X1,X3)))), file('out/t47_yellow_0__3_premises__rerun.E_input', t46_yellow_0)).
fof(c_0_4, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:![X3]:((r1_yellow_0(X1,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3))))), inference(assume_negation,[status(cth)],[t47_yellow_0])).
fof(c_0_5, negated_conjecture, (![X8]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((r1_yellow_0(esk1_0,esk2_0)&(((~r2_lattice3(esk1_0,esk2_0,X8)|r2_lattice3(esk1_0,esk3_0,X8))|~m1_subset_1(X8,u1_struct_0(esk1_0)))&((~r2_lattice3(esk1_0,esk3_0,X8)|r2_lattice3(esk1_0,esk2_0,X8))|~m1_subset_1(X8,u1_struct_0(esk1_0)))))&k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,esk3_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])])])])])).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:((((((r2_lattice3(X5,X6,X7)|X7!=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X6,X8)|r1_orders_2(X5,X7,X8)))|X7!=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))&((((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((((r2_lattice3(X5,X6,esk4_3(X5,X6,X7))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~r1_orders_2(X5,X7,esk4_3(X5,X6,X7))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])])])).
fof(c_0_7, plain, (![X3]:![X4]:(~l1_orders_2(X3)|m1_subset_1(k1_yellow_0(X3,X4),u1_struct_0(X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])])])).
cnf(c_0_8,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(X2=k1_yellow_0(X1,X3)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,plain,(X2=k1_yellow_0(X1,X3)|r2_lattice3(X1,X3,esk4_3(X1,X3,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,plain,(m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,negated_conjecture,(X1=k1_yellow_0(esk1_0,X2)|r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_14,plain,(k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|r2_lattice3(X1,X3,esk4_3(X1,X3,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15,plain,(r1_orders_2(X1,X2,X4)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|X2!=k1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_10])])).
cnf(c_0_17,plain,(X1=k1_yellow_0(X2,X3)|r1_orders_2(X2,X4,esk4_3(X2,X3,X1))|X4!=k1_yellow_0(X2,X5)|~r2_lattice3(X2,X5,esk4_3(X2,X3,X1))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_yellow_0(X2,X5)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_15, c_0_9])).
cnf(c_0_18,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_12]), c_0_10])])).
cnf(c_0_19,negated_conjecture,(r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_20,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r1_orders_2(esk1_0,X2,esk4_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|X2!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_10])])).
cnf(c_0_21,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r1_orders_2(esk1_0,X2,esk4_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|X2!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_12]), c_0_10])])).
cnf(c_0_23,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_12]), c_0_10])])).
cnf(c_0_24,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|X2!=k1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r1_orders_2(esk1_0,X2,esk4_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|X2!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26,plain,(r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|k1_yellow_0(X1,X3)!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_27,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r1_orders_2(esk1_0,X2,esk4_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|X2!=k1_yellow_0(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_19]), c_0_10])])).
cnf(c_0_28,plain,(X2=k1_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_29,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),esk4_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|k1_yellow_0(esk1_0,X2)!=k1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_12]), c_0_10])])).
cnf(c_0_30,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_10])])).
cnf(c_0_31,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_12]), c_0_10])])).
fof(c_0_32, plain, (![X5]:![X6]:![X7]:((((m1_subset_1(esk5_3(X5,X6,X7),u1_struct_0(X5))|~r1_yellow_0(X5,X6))|r1_yellow_0(X5,X7))|(v2_struct_0(X5)|~l1_orders_2(X5)))&(((((~r2_lattice3(X5,X6,esk5_3(X5,X6,X7))|~r2_lattice3(X5,X7,esk5_3(X5,X6,X7)))|~r1_yellow_0(X5,X6))|r1_yellow_0(X5,X7))|(v2_struct_0(X5)|~l1_orders_2(X5)))&((((r2_lattice3(X5,X6,esk5_3(X5,X6,X7))|r2_lattice3(X5,X7,esk5_3(X5,X6,X7)))|~r1_yellow_0(X5,X6))|r1_yellow_0(X5,X7))|(v2_struct_0(X5)|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_yellow_0])])])])])])])])).
cnf(c_0_33,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_31, c_0_23])).
cnf(c_0_34,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk5_3(X1,X3,X2))|r2_lattice3(X1,X3,esk5_3(X1,X3,X2))|~l1_orders_2(X1)|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_35,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_36,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_26]), c_0_19]), c_0_10])])).
cnf(c_0_37,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk2_0,X1))|r2_lattice3(esk1_0,X1,esk5_3(esk1_0,esk2_0,X1))|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_19]), c_0_10])]), c_0_35])).
cnf(c_0_38,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|m1_subset_1(esk5_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,esk2_0,esk3_0))|r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk2_0,esk3_0))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40,negated_conjecture,(k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_41,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,X1,X2))|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,X1,X2))|~r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_38]), c_0_10])]), c_0_35])).
cnf(c_0_42,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk2_0,esk3_0))|r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_39]), c_0_40])).
cnf(c_0_43,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk2_0,esk3_0))|r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_19])])).
cnf(c_0_44,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk2_0,esk3_0))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_43])).
cnf(c_0_45,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X2,esk5_3(X1,X3,X2))|~r2_lattice3(X1,X3,esk5_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_44]), c_0_40])).
cnf(c_0_47,negated_conjecture,(r1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_19]), c_0_10])]), c_0_35])).
cnf(c_0_48,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,X1,X2))|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,X1,X2))|~r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_38]), c_0_10])]), c_0_35])).
cnf(c_0_49,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_36, c_0_47])).
cnf(c_0_50,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,esk2_0,X1))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk2_0,X1))), inference(spm,[status(thm)],[c_0_48, c_0_19])).
cnf(c_0_51,negated_conjecture,(~r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_49]), c_0_40])).
cnf(c_0_52,negated_conjecture,(k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_50]), c_0_46])]), c_0_51])).
cnf(c_0_53,negated_conjecture,($false), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_52]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 45
# Proof object formula steps           : 9
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 157
# ...of these trivial                  : 0
# ...subsumed                          : 20
# ...remaining for further processing  : 137
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 9
# Generated clauses                    : 403
# ...of the previous two non-trivial   : 346
# Contextual simplify-reflections      : 22
# Paramodulations                      : 357
# Factorizations                       : 34
# Equation resolutions                 : 12
# Current number of processed clauses  : 101
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 93
# Current number of unprocessed clauses: 167
# ...number of literals in the above   : 1375
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 1973
# Rec. Clause-clause subsumption calls : 349
# Non-unit clause-clause subsumptions  : 53
# Unit Clause-clause subsumption calls : 146
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 46
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 16188

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.000 s
# Total time               : 0.032 s
# Maximum resident set size: 45304 pages
