# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t47_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k3_xboole_0(k2_tarski(X1,X3),X2)=k2_tarski(X1,X3))), file('out/t47_zfmisc_1__14_premises__rerun.E_input', t47_zfmisc_1)).
fof(d1_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3))))))), file('out/t47_zfmisc_1__14_premises__rerun.E_input', d1_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t47_zfmisc_1__14_premises__rerun.E_input', t71_enumset1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t47_zfmisc_1__14_premises__rerun.E_input', t70_enumset1)).
fof(t102_enumset1, axiom, (![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), file('out/t47_zfmisc_1__14_premises__rerun.E_input', t102_enumset1)).
fof(t69_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), file('out/t47_zfmisc_1__14_premises__rerun.E_input', t69_enumset1)).
fof(t72_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X2,X4,X1)), file('out/t47_zfmisc_1__14_premises__rerun.E_input', t72_enumset1)).
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t47_zfmisc_1__14_premises__rerun.E_input', d4_xboole_0)).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k3_xboole_0(k2_tarski(X1,X3),X2)=k2_tarski(X1,X3)))), inference(assume_negation,[status(cth)],[t47_zfmisc_1])).
fof(c_0_9, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X10]:![X6]:![X7]:![X8]:![X9]:((((~r2_hidden(X10,X9)|(X10=X6|(X10=X7|X10=X8)))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X6|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X7|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&((X10!=X8|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8)))))&((((esk2_4(X6,X7,X8,X9)!=X6|~r2_hidden(esk2_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&(((esk2_4(X6,X7,X8,X9)!=X7|~r2_hidden(esk2_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&((esk2_4(X6,X7,X8,X9)!=X8|~r2_hidden(esk2_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))))&((r2_hidden(esk2_4(X6,X7,X8,X9),X9)|(esk2_4(X6,X7,X8,X9)=X6|(esk2_4(X6,X7,X8,X9)=X7|esk2_4(X6,X7,X8,X9)=X8)))|X9=k1_enumset1(X6,X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_11, negated_conjecture, (((r2_hidden(esk3_0,esk4_0)&r2_hidden(esk5_0,esk4_0))&k3_xboole_0(k2_tarski(esk3_0,esk5_0),esk4_0)!=k2_tarski(esk3_0,esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_12, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_13,plain,(X5=X4|X5=X3|X5=X2|X1!=k1_enumset1(X2,X3,X4)|~r2_hidden(X5,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:k1_enumset1(X4,X5,X6)=k1_enumset1(X6,X5,X4)), inference(variable_rename,[status(thm)],[t102_enumset1])).
cnf(c_0_16,negated_conjecture,(k3_xboole_0(k2_tarski(esk3_0,esk5_0),esk4_0)!=k2_tarski(esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(X1=k1_enumset1(X2,X3,X4)|esk2_4(X2,X3,X4,X1)=X4|esk2_4(X2,X3,X4,X1)=X3|esk2_4(X2,X3,X4,X1)=X2|r2_hidden(esk2_4(X2,X3,X4,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19,plain,(X4=X5|X3=X5|X2=X5|X1!=k2_enumset1(X2,X2,X3,X4)|~r2_hidden(X5,X1)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X6,X8,X5,X7)), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_21, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X7,X6,X8,X5)), inference(variable_rename,[status(thm)],[t72_enumset1])).
cnf(c_0_22,plain,(k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_23, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|~r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
cnf(c_0_24,negated_conjecture,(k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0)!=k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_17]), c_0_14]), c_0_14])).
cnf(c_0_25,plain,(X1=k2_enumset1(X2,X2,X3,X4)|esk2_4(X2,X3,X4,X1)=X4|esk2_4(X2,X3,X4,X1)=X3|esk2_4(X2,X3,X4,X1)=X2|r2_hidden(esk2_4(X2,X3,X4,X1),X1)), inference(rw,[status(thm)],[c_0_18, c_0_14])).
cnf(c_0_26,plain,(X1=X2|X3=X2|X4=X2|~r2_hidden(X2,k2_enumset1(X4,X4,X3,X1))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X2,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X3,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_14]), c_0_14])).
cnf(c_0_30,plain,(X1=k1_enumset1(X2,X3,X4)|~r2_hidden(esk2_4(X2,X3,X4,X1),X1)|esk2_4(X2,X3,X4,X1)!=X4), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31,plain,(r2_hidden(X4,X1)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(r2_hidden(X4,X2)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33,negated_conjecture,(esk2_4(X1,X2,X3,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=X1|esk2_4(X1,X2,X3,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=X2|esk2_4(X1,X2,X3,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=X3|r2_hidden(esk2_4(X1,X2,X3,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0)),k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))|k2_enumset1(X1,X1,X2,X3)!=k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_34,plain,(X1=X2|X3=X2|X4=X2|~r2_hidden(X2,k2_enumset1(X4,X1,X4,X3))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35,plain,(k2_enumset1(X1,X1,X2,X3)=k2_enumset1(X2,X3,X1,X3)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X4), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37,plain,(X1=k1_enumset1(X2,X3,X4)|~r2_hidden(esk2_4(X2,X3,X4,X1),X1)|esk2_4(X2,X3,X4,X1)!=X2), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_38,plain,(X1=k2_enumset1(X2,X2,X3,X4)|esk2_4(X2,X3,X4,X1)!=X4|~r2_hidden(esk2_4(X2,X3,X4,X1),X1)), inference(rw,[status(thm)],[c_0_30, c_0_14])).
cnf(c_0_39,plain,(r2_hidden(X1,k3_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_40,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_41,negated_conjecture,(esk2_4(esk3_0,esk3_0,esk5_0,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=esk5_0|esk2_4(esk3_0,esk3_0,esk5_0,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=esk3_0|r2_hidden(esk2_4(esk3_0,esk3_0,esk5_0,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0)),k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_42,plain,(X1=X2|X3=X2|~r2_hidden(X2,k2_enumset1(X1,X1,X1,X3))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43,plain,(r2_hidden(X5,X1)|X4!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_36, c_0_14])).
cnf(c_0_44,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X2), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_45,plain,(X1=k2_enumset1(X2,X2,X3,X4)|esk2_4(X2,X3,X4,X1)!=X2|~r2_hidden(esk2_4(X2,X3,X4,X1),X1)), inference(rw,[status(thm)],[c_0_37, c_0_14])).
cnf(c_0_46,plain,(k3_xboole_0(X1,X2)=k2_enumset1(X3,X3,X4,X5)|esk2_4(X3,X4,X5,k3_xboole_0(X1,X2))!=X5|~r2_hidden(esk2_4(X3,X4,X5,k3_xboole_0(X1,X2)),X2)|~r2_hidden(esk2_4(X3,X4,X5,k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_47,negated_conjecture,(esk2_4(esk3_0,esk3_0,esk5_0,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=esk3_0|esk2_4(esk3_0,esk3_0,esk5_0,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=esk5_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_48,negated_conjecture,(r2_hidden(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_49,plain,(r2_hidden(X1,k2_enumset1(X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])])).
cnf(c_0_50,plain,(r2_hidden(X5,X1)|X2!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_44, c_0_14])).
cnf(c_0_51,plain,(k3_xboole_0(X1,X2)=k2_enumset1(X3,X3,X4,X5)|esk2_4(X3,X4,X5,k3_xboole_0(X1,X2))!=X3|~r2_hidden(esk2_4(X3,X4,X5,k3_xboole_0(X1,X2)),X2)|~r2_hidden(esk2_4(X3,X4,X5,k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_45, c_0_39])).
cnf(c_0_52,negated_conjecture,(esk2_4(esk3_0,esk3_0,esk5_0,k3_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk5_0),esk4_0))=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49])]), c_0_24])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54,plain,(r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_50])])).
cnf(c_0_55,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_54])]), c_0_24]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 39
# Proof object formula steps           : 17
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 23
# Processed clauses                    : 1697
# ...of these trivial                  : 35
# ...subsumed                          : 1389
# ...remaining for further processing  : 273
# Other redundant clauses eliminated   : 235
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 26352
# ...of the previous two non-trivial   : 24753
# Contextual simplify-reflections      : 2
# Paramodulations                      : 26045
# Factorizations                       : 74
# Equation resolutions                 : 236
# Current number of processed clauses  : 240
#    Positive orientable unit clauses  : 35
#    Positive unorientable unit clauses: 15
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 189
# Current number of unprocessed clauses: 22941
# ...number of literals in the above   : 146844
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 19204
# Rec. Clause-clause subsumption calls : 10504
# Non-unit clause-clause subsumptions  : 1367
# Unit Clause-clause subsumption calls : 620
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2258
# BW rewrite match successes           : 1263
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 546810

# -------------------------------------------------
# User time                : 0.640 s
# System time              : 0.012 s
# Total time               : 0.652 s
# Maximum resident set size: 45288 pages
