# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v1_tops_1(X2,X1)&v2_tex_2(X2,X1))=>v3_tex_2(X2,X1))))), file('out/t48_tex_2__7_premises__rerun.E_input', t48_tex_2)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t48_tex_2__7_premises__rerun.E_input', d10_xboole_0)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t48_tex_2__7_premises__rerun.E_input', t3_subset)).
fof(d2_tops_3, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>k2_pre_topc(X1,X2)=u1_struct_0(X1))))), file('out/t48_tex_2__7_premises__rerun.E_input', d2_tops_3)).
fof(t49_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3))))))), file('out/t48_tex_2__7_premises__rerun.E_input', t49_pre_topc)).
fof(dt_k2_pre_topc, axiom, (![X1]:![X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t48_tex_2__7_premises__rerun.E_input', dt_k2_pre_topc)).
fof(d7_tex_2, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_tex_2(X2,X1)<=>(v2_tex_2(X2,X1)&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v2_tex_2(X3,X1)&r1_tarski(X2,X3))=>X2=X3))))))), file('out/t48_tex_2__7_premises__rerun.E_input', d7_tex_2)).
fof(t41_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X3,X2)=>k9_subset_1(u1_struct_0(X1),X2,k2_pre_topc(X1,X3))=X3)))))), file('out/t48_tex_2__7_premises__rerun.E_input', t41_tex_2)).
fof(c_0_8, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v1_tops_1(X2,X1)&v2_tex_2(X2,X1))=>v3_tex_2(X2,X1)))))), inference(assume_negation,[status(cth)],[t48_tex_2])).
fof(c_0_9, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
fof(c_0_10, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_11, plain, (![X3]:![X4]:((((~v1_tops_1(X4,X3)|k2_pre_topc(X3,X4)=u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3))&(((k2_pre_topc(X3,X4)!=u1_struct_0(X3)|v1_tops_1(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tops_3])])])])])])).
fof(c_0_12, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((v1_tops_1(esk2_0,esk1_0)&v2_tex_2(esk2_0,esk1_0))&~v3_tex_2(esk2_0,esk1_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
cnf(c_0_13,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(~l1_pre_topc(X4)|(~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(~r1_tarski(X5,X6)|r1_tarski(k2_pre_topc(X4,X5),k2_pre_topc(X4,X6))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_pre_topc])])])])])).
cnf(c_0_16,plain,(k2_pre_topc(X1,X2)=u1_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_21, plain, (![X3]:![X4]:((~l1_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|m1_subset_1(k2_pre_topc(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
cnf(c_0_22,plain,(r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3))|~r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,negated_conjecture,(k2_pre_topc(esk1_0,esk2_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_24,plain,(X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_20, c_0_14])).
cnf(c_0_25,plain,(m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27,negated_conjecture,(r1_tarski(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,X1))|~r1_tarski(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_18]), c_0_19])])).
cnf(c_0_28,plain,(k2_pre_topc(X1,X2)=u1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k2_pre_topc(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,negated_conjecture,(m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k2_pre_topc(esk1_0,X1)))|~r1_tarski(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30,negated_conjecture,(k2_pre_topc(esk1_0,X1)=u1_struct_0(esk1_0)|~r1_tarski(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_19])])).
fof(c_0_31, plain, (![X4]:![X5]:![X6]:(((((v2_tex_2(X5,X4)|~v3_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|~l1_pre_topc(X4))&((((~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|((~v2_tex_2(X6,X4)|~r1_tarski(X5,X6))|X5=X6))|~v3_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|~l1_pre_topc(X4)))&(((((m1_subset_1(esk4_2(X4,X5),k1_zfmisc_1(u1_struct_0(X4)))|~v2_tex_2(X5,X4))|v3_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|~l1_pre_topc(X4))&((((((v2_tex_2(esk4_2(X4,X5),X4)|~v2_tex_2(X5,X4))|v3_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|~l1_pre_topc(X4))&((((r1_tarski(X5,esk4_2(X4,X5))|~v2_tex_2(X5,X4))|v3_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|~l1_pre_topc(X4)))&((((X5!=esk4_2(X4,X5)|~v2_tex_2(X5,X4))|v3_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_tex_2])])])])])])])).
cnf(c_0_32,negated_conjecture,(r1_tarski(k2_pre_topc(esk1_0,X1),u1_struct_0(esk1_0))|~r1_tarski(esk2_0,X2)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_30]), c_0_19])])).
cnf(c_0_33,plain,(v3_tex_2(X2,X1)|m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_34,negated_conjecture,(r1_tarski(k2_pre_topc(esk1_0,X1),u1_struct_0(esk1_0))|v3_tex_2(X2,esk1_0)|~r1_tarski(esk2_0,esk4_2(esk1_0,X2))|~r1_tarski(X1,esk4_2(esk1_0,X2))|~v2_tex_2(X2,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_19])])).
cnf(c_0_35,plain,(v3_tex_2(X2,X1)|r1_tarski(X2,esk4_2(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36,negated_conjecture,(v2_tex_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37,negated_conjecture,(~v3_tex_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_38,plain,(r1_tarski(X2,X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39,negated_conjecture,(r1_tarski(k2_pre_topc(esk1_0,X1),u1_struct_0(esk1_0))|~r1_tarski(X1,esk4_2(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_18]), c_0_19])]), c_0_37])).
cnf(c_0_40,plain,(r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_41,plain,(k2_pre_topc(X1,X2)=k2_pre_topc(X1,X3)|~r1_tarski(X2,X3)|~m1_subset_1(k2_pre_topc(X1,X3),k1_zfmisc_1(k2_pre_topc(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_20, c_0_22])).
cnf(c_0_42,negated_conjecture,(r1_tarski(k2_pre_topc(esk1_0,esk4_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_43,negated_conjecture,(k2_pre_topc(esk1_0,X1)=u1_struct_0(esk1_0)|~r1_tarski(esk2_0,X2)|~r1_tarski(X2,X1)|~m1_subset_1(k2_pre_topc(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_30]), c_0_19])])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(k2_pre_topc(esk1_0,esk4_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_26, c_0_42])).
fof(c_0_45, plain, (![X4]:![X5]:![X6]:((((~v2_tex_2(X5,X4)|(~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(~r1_tarski(X6,X5)|k9_subset_1(u1_struct_0(X4),X5,k2_pre_topc(X4,X6))=X6)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((((m1_subset_1(esk3_2(X4,X5),k1_zfmisc_1(u1_struct_0(X4)))|v2_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((((r1_tarski(esk3_2(X4,X5),X5)|v2_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((k9_subset_1(u1_struct_0(X4),X5,k2_pre_topc(X4,esk3_2(X4,X5)))!=esk3_2(X4,X5)|v2_tex_2(X5,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t41_tex_2])])])])])])])])).
cnf(c_0_46,negated_conjecture,(k2_pre_topc(esk1_0,esk4_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)|~r1_tarski(X1,esk4_2(esk1_0,esk2_0))|~r1_tarski(esk2_0,X1)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47,plain,(v2_struct_0(X1)|k9_subset_1(u1_struct_0(X1),X2,k2_pre_topc(X1,X3))=X3|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_48,negated_conjecture,(k2_pre_topc(esk1_0,esk4_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_35]), c_0_40]), c_0_18]), c_0_36]), c_0_19])]), c_0_37])).
cnf(c_0_49,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51,negated_conjecture,(k9_subset_1(u1_struct_0(esk1_0),X1,u1_struct_0(esk1_0))=esk4_2(esk1_0,esk2_0)|~r1_tarski(esk4_2(esk1_0,esk2_0),X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_19]), c_0_49])]), c_0_50])).
cnf(c_0_52,negated_conjecture,(k9_subset_1(u1_struct_0(esk1_0),X1,u1_struct_0(esk1_0))=esk2_0|~r1_tarski(esk2_0,X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_23]), c_0_18]), c_0_19]), c_0_49])]), c_0_50])).
cnf(c_0_53,negated_conjecture,(k9_subset_1(u1_struct_0(esk1_0),X1,u1_struct_0(esk1_0))=esk4_2(esk1_0,esk2_0)|~r1_tarski(esk4_2(esk1_0,esk2_0),X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_33]), c_0_36]), c_0_18]), c_0_19])]), c_0_37])).
cnf(c_0_54,negated_conjecture,(esk4_2(esk1_0,esk2_0)=esk2_0|~r1_tarski(esk4_2(esk1_0,esk2_0),X1)|~r1_tarski(esk2_0,X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_55,negated_conjecture,(esk4_2(esk1_0,esk2_0)=esk2_0|~r1_tarski(esk2_0,esk4_2(esk1_0,esk2_0))|~v2_tex_2(esk4_2(esk1_0,esk2_0),esk1_0)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_54, c_0_40])).
cnf(c_0_56,negated_conjecture,(esk4_2(esk1_0,esk2_0)=esk2_0|~r1_tarski(esk2_0,esk4_2(esk1_0,esk2_0))|~v2_tex_2(esk4_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_33]), c_0_36]), c_0_18]), c_0_19])]), c_0_37])).
cnf(c_0_57,negated_conjecture,(esk4_2(esk1_0,esk2_0)=esk2_0|~v2_tex_2(esk4_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_35]), c_0_36]), c_0_18]), c_0_19])]), c_0_37])).
cnf(c_0_58,plain,(v3_tex_2(X2,X1)|v2_tex_2(esk4_2(X1,X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_59,plain,(v3_tex_2(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_tex_2(X2,X1)|X2!=esk4_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_60,negated_conjecture,(esk4_2(esk1_0,esk2_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_36]), c_0_18]), c_0_19])]), c_0_37])).
cnf(c_0_61,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_36]), c_0_18]), c_0_19])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 45
# Proof object formula steps           : 17
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 832
# ...of these trivial                  : 1
# ...subsumed                          : 524
# ...remaining for further processing  : 307
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 60
# Backward-rewritten                   : 45
# Generated clauses                    : 1581
# ...of the previous two non-trivial   : 1340
# Contextual simplify-reflections      : 991
# Paramodulations                      : 1579
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 175
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 162
# Current number of unprocessed clauses: 294
# ...number of literals in the above   : 2197
# Current number of archived formulas  : 0
# Current number of archived clauses   : 130
# Clause-clause subsumption calls (NU) : 26742
# Rec. Clause-clause subsumption calls : 8627
# Non-unit clause-clause subsumptions  : 1554
# Unit Clause-clause subsumption calls : 256
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 54977

# -------------------------------------------------
# User time                : 0.072 s
# System time              : 0.004 s
# Total time               : 0.076 s
# Maximum resident set size: 43100 pages
