# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3)))))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', t48_waybel_0)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', t4_subset)).
fof(d21_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_hidden(X3,X2)&r2_lattice3(X1,X2,X3))))))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', d21_waybel_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', d3_tarski)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', d9_lattice3)).
fof(dt_k5_waybel_0, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', dt_k5_waybel_0)).
fof(d19_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X4,X3))=>r2_hidden(X4,X2)))))))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', d19_waybel_0)).
fof(t17_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', t17_waybel_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', d10_xboole_0)).
fof(reflexivity_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', reflexivity_r3_orders_2)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t48_waybel_0__12_premises__rerun__rerun.E_input', redefinition_r3_orders_2)).
fof(c_0_11, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3))))))), inference(assume_negation,[status(cth)],[t48_waybel_0])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_13, negated_conjecture, (![X6]:((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&((~v1_xboole_0(esk2_0)&(v1_waybel_0(esk2_0,esk1_0)&(v12_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v14_waybel_0(esk2_0,esk1_0)|(~m1_subset_1(X6,u1_struct_0(esk1_0))|esk2_0!=k5_waybel_0(esk1_0,X6)))&((m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v14_waybel_0(esk2_0,esk1_0))&(esk2_0=k5_waybel_0(esk1_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X7]:(((((m1_subset_1(esk4_2(X4,X5),u1_struct_0(X4))|~v14_waybel_0(X5,X4))|(v1_xboole_0(X5)|(~v1_waybel_0(X5,X4)|(~v12_waybel_0(X5,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))))))|(v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|~l1_orders_2(X4)))))&((((r2_hidden(esk4_2(X4,X5),X5)|~v14_waybel_0(X5,X4))|(v1_xboole_0(X5)|(~v1_waybel_0(X5,X4)|(~v12_waybel_0(X5,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))))))|(v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|~l1_orders_2(X4)))))&(((r2_lattice3(X4,X5,esk4_2(X4,X5))|~v14_waybel_0(X5,X4))|(v1_xboole_0(X5)|(~v1_waybel_0(X5,X4)|(~v12_waybel_0(X5,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))))))|(v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|~l1_orders_2(X4)))))))&((((~m1_subset_1(X7,u1_struct_0(X4))|(~r2_hidden(X7,X5)|~r2_lattice3(X4,X5,X7)))|v14_waybel_0(X5,X4))|(v1_xboole_0(X5)|(~v1_waybel_0(X5,X4)|(~v12_waybel_0(X5,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))))))|(v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|~l1_orders_2(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_waybel_0])])])])])])])])).
cnf(c_0_15,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_18,plain,(v2_struct_0(X1)|v1_xboole_0(X2)|m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v14_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,negated_conjecture,(v12_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_26, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk6_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk6_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk6_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
cnf(c_0_27,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_28,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_29, plain, (![X3]:![X4]:(((v2_struct_0(X3)|~l1_orders_2(X3))|~m1_subset_1(X4,u1_struct_0(X3)))|m1_subset_1(k5_waybel_0(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~v14_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_16]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24]), c_0_25])).
cnf(c_0_31,negated_conjecture,(v14_waybel_0(esk2_0,esk1_0)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,negated_conjecture,(r1_tarski(esk2_0,X1)|m1_subset_1(esk5_2(esk2_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34,plain,(v2_struct_0(X1)|v1_xboole_0(X2)|r2_lattice3(X1,X2,esk4_2(X1,X2))|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v14_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_35, plain, (![X5]:![X6]:![X7]:![X8]:((((~v12_waybel_0(X6,X5)|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|((~r2_hidden(X7,X6)|~r1_orders_2(X5,X8,X7))|r2_hidden(X8,X6)))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_orders_2(X5))&((((m1_subset_1(esk7_2(X5,X6),u1_struct_0(X5))|v12_waybel_0(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_orders_2(X5))&((((m1_subset_1(esk8_2(X5,X6),u1_struct_0(X5))|v12_waybel_0(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_orders_2(X5))&(((((r2_hidden(esk7_2(X5,X6),X6)|v12_waybel_0(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_orders_2(X5))&(((r1_orders_2(X5,esk8_2(X5,X6),esk7_2(X5,X6))|v12_waybel_0(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_orders_2(X5)))&(((~r2_hidden(esk8_2(X5,X6),X6)|v12_waybel_0(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_waybel_0])])])])])])])).
cnf(c_0_36,plain,(m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,negated_conjecture,(m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:(((((~r2_hidden(X6,k5_waybel_0(X4,X5))|r1_orders_2(X4,X6,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((~r1_orders_2(X4,X6,X5)|r2_hidden(X6,k5_waybel_0(X4,X5)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])])])).
cnf(c_0_39,negated_conjecture,(r1_orders_2(esk1_0,esk5_2(esk2_0,X1),X2)|r1_tarski(esk2_0,X1)|~r2_lattice3(esk1_0,X3,X2)|~r2_hidden(esk5_2(esk2_0,X1),X3)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_21])])).
cnf(c_0_40,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))|~v14_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_16]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24]), c_0_25])).
cnf(c_0_41,plain,(r2_hidden(X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,X3,X4)|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v12_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42,negated_conjecture,(m1_subset_1(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_21])]), c_0_25])).
cnf(c_0_43,plain,(v2_struct_0(X1)|r2_hidden(X3,k5_waybel_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44,negated_conjecture,(r1_orders_2(esk1_0,esk5_2(esk2_0,X1),X2)|r1_tarski(esk2_0,X1)|~r2_lattice3(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_28])).
cnf(c_0_45,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_31])).
cnf(c_0_46,plain,(r2_hidden(X1,X2)|~r1_orders_2(X3,X1,X4)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v12_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_41, c_0_15])).
cnf(c_0_47,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))|m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_15, c_0_42])).
cnf(c_0_48,plain,(v2_struct_0(X1)|r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k5_waybel_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49,plain,(v2_struct_0(X1)|v1_xboole_0(X2)|r2_hidden(esk4_2(X1,X2),X2)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v14_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_50,negated_conjecture,(r1_tarski(esk2_0,X1)|r2_hidden(esk5_2(esk2_0,X1),k5_waybel_0(esk1_0,X2))|~r1_orders_2(esk1_0,esk5_2(esk2_0,X1),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_33]), c_0_21])]), c_0_25])).
cnf(c_0_51,negated_conjecture,(r1_orders_2(esk1_0,esk5_2(esk2_0,X1),esk4_2(esk1_0,esk2_0))|r1_tarski(esk2_0,X1)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_37])).
cnf(c_0_52,negated_conjecture,(esk2_0!=k5_waybel_0(esk1_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v14_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_53,negated_conjecture,(r2_hidden(X1,esk2_0)|~r1_orders_2(esk1_0,X1,X2)|~r2_hidden(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_16]), c_0_19]), c_0_21])])).
cnf(c_0_54,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|m1_subset_1(esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_47, c_0_28])).
cnf(c_0_55,plain,(r1_orders_2(X1,esk5_2(k5_waybel_0(X1,X2),X3),X2)|r1_tarski(k5_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(esk5_2(k5_waybel_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_48, c_0_28])).
cnf(c_0_56,negated_conjecture,(r2_hidden(esk4_2(esk1_0,esk2_0),esk2_0)|~v14_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_16]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24]), c_0_25])).
fof(c_0_57, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_58,plain,(r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_59,negated_conjecture,(r1_tarski(esk2_0,X1)|r2_hidden(esk5_2(esk2_0,X1),k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_37])).
cnf(c_0_60,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))|k5_waybel_0(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_52, c_0_31])).
cnf(c_0_61,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),X2)|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62,negated_conjecture,(r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk4_2(esk1_0,esk2_0))|r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_54]), c_0_21])]), c_0_25]), c_0_37])).
cnf(c_0_63,negated_conjecture,(r2_hidden(esk4_2(esk1_0,esk2_0),esk2_0)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_56, c_0_31])).
cnf(c_0_64,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65,negated_conjecture,(r1_tarski(esk2_0,k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_66,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))|k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0))!=esk2_0), inference(spm,[status(thm)],[c_0_60, c_0_37])).
fof(c_0_67, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|r3_orders_2(X4,X5,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])])])).
cnf(c_0_68,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_69,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))|~r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])).
fof(c_0_70, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_71,plain,(r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_72,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_68]), c_0_69])).
cnf(c_0_73,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_74,negated_conjecture,(r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_21]), c_0_23])]), c_0_25])).
cnf(c_0_75,negated_conjecture,(v14_waybel_0(esk2_0,esk1_0)|esk2_0=k5_waybel_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_76,plain,(r2_lattice3(X1,X3,X2)|r2_hidden(esk6_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_77,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|~r3_orders_2(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_72]), c_0_21]), c_0_23])]), c_0_25])).
cnf(c_0_78,negated_conjecture,(r3_orders_2(esk1_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_74, c_0_72])).
cnf(c_0_79,plain,(v2_struct_0(X1)|v1_xboole_0(X2)|v14_waybel_0(X2,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~r2_lattice3(X1,X2,X3)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_80,negated_conjecture,(r1_orders_2(esk1_0,X1,esk3_0)|v14_waybel_0(esk2_0,esk1_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_75]), c_0_21])]), c_0_25]), c_0_31]), c_0_27])).
cnf(c_0_81,negated_conjecture,(r2_lattice3(esk1_0,X1,esk3_0)|r2_hidden(esk6_3(esk1_0,X1,esk3_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_72]), c_0_21])])).
cnf(c_0_82,negated_conjecture,(r2_hidden(esk3_0,k5_waybel_0(esk1_0,X1))|~r1_orders_2(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_72]), c_0_21])]), c_0_25])).
cnf(c_0_83,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_72])])).
cnf(c_0_84,plain,(v14_waybel_0(X1,X2)|v1_xboole_0(X1)|v2_struct_0(X2)|~r2_lattice3(X2,X1,X3)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v12_waybel_0(X1,X2)|~v1_waybel_0(X1,X2)|~l1_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_79, c_0_15])).
cnf(c_0_85,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk6_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_86,negated_conjecture,(r1_orders_2(esk1_0,esk6_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_lattice3(esk1_0,esk2_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_87,negated_conjecture,(r2_hidden(esk3_0,k5_waybel_0(esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_72])])).
cnf(c_0_88,negated_conjecture,(v14_waybel_0(esk2_0,esk1_0)|~r2_lattice3(esk1_0,esk2_0,X1)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_16]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24]), c_0_25])).
cnf(c_0_89,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_72]), c_0_21])])).
cnf(c_0_90,negated_conjecture,(r2_hidden(esk3_0,esk2_0)|v14_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_87, c_0_75])).
cnf(c_0_91,negated_conjecture,(v14_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90])).
cnf(c_0_92,negated_conjecture,(m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_91])])).
cnf(c_0_93,negated_conjecture,(m1_subset_1(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_92]), c_0_21])]), c_0_25])).
cnf(c_0_94,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_91])])).
cnf(c_0_95,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_15, c_0_93])).
cnf(c_0_96,negated_conjecture,(r1_orders_2(esk1_0,esk5_2(esk2_0,X1),esk4_2(esk1_0,esk2_0))|r1_tarski(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_94]), c_0_92])])).
cnf(c_0_97,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|m1_subset_1(esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_95, c_0_28])).
cnf(c_0_98,negated_conjecture,(r1_tarski(esk2_0,X1)|r2_hidden(esk5_2(esk2_0,X1),k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_96]), c_0_92])])).
cnf(c_0_99,negated_conjecture,(k5_waybel_0(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_91])])).
cnf(c_0_100,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)|~r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),X2)|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_97])).
cnf(c_0_101,negated_conjecture,(r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk4_2(esk1_0,esk2_0))|r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_97]), c_0_92]), c_0_21])]), c_0_25])).
cnf(c_0_102,negated_conjecture,(r2_hidden(esk4_2(esk1_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_91])])).
cnf(c_0_103,negated_conjecture,(r1_tarski(esk2_0,k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_58, c_0_98])).
cnf(c_0_104,negated_conjecture,(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0))!=esk2_0), inference(spm,[status(thm)],[c_0_99, c_0_92])).
cnf(c_0_105,negated_conjecture,(r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102])])).
cnf(c_0_106,negated_conjecture,(~r1_tarski(k5_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_103]), c_0_104])).
cnf(c_0_107,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_105]), c_0_106]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 108
# Proof object clause steps            : 85
# Proof object formula steps           : 23
# Proof object conjectures             : 68
# Proof object clause conjectures      : 65
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 51
# Proof object simplifying inferences  : 103
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 716
# ...of these trivial                  : 2
# ...subsumed                          : 184
# ...remaining for further processing  : 530
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 64
# Backward-rewritten                   : 98
# Generated clauses                    : 1647
# ...of the previous two non-trivial   : 1606
# Contextual simplify-reflections      : 207
# Paramodulations                      : 1645
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 329
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 303
# Current number of unprocessed clauses: 651
# ...number of literals in the above   : 3090
# Current number of archived formulas  : 0
# Current number of archived clauses   : 199
# Clause-clause subsumption calls (NU) : 29115
# Rec. Clause-clause subsumption calls : 7515
# Non-unit clause-clause subsumptions  : 450
# Unit Clause-clause subsumption calls : 618
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 65930

# -------------------------------------------------
# User time                : 0.084 s
# System time              : 0.000 s
# Total time               : 0.084 s
# Maximum resident set size: 43224 pages
