# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_waybel_7, axiom, (![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v3_waybel_3(X1)&l1_orders_2(X1))))))))=>(v5_waybel_7(k1_waybel_4(X1),X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v4_waybel_7(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r1_waybel_3(X1,k12_lattice3(X1,X3,X4),X2)&(~(r3_orders_2(X1,X3,X2))&~(r3_orders_2(X1,X4,X2)))))))))))), file('out/t49_waybel_7__2_premises__rerun.E_input', t48_waybel_7)).
fof(l57_waybel_7, axiom, (![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v3_waybel_3(X1)&l1_orders_2(X1))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~((v5_waybel_7(k1_waybel_4(X1),X1)&(![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r1_waybel_3(X1,k12_lattice3(X1,X3,X4),X2)&(~(r3_orders_2(X1,X3,X2))&~(r3_orders_2(X1,X4,X2)))))))&~(v5_waybel_6(X2,X1)))))))), file('out/t49_waybel_7__2_premises__rerun.E_input', l57_waybel_7)).
fof(t49_waybel_7, conjecture, (![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v3_waybel_3(X1)&l1_orders_2(X1))))))))=>(v5_waybel_7(k1_waybel_4(X1),X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v4_waybel_7(X2,X1)=>v5_waybel_6(X2,X1)))))), file('out/t49_waybel_7__2_premises__rerun.E_input', t49_waybel_7)).
fof(c_0_3, plain, (![X5]:![X6]:![X7]:![X8]:(((((~v4_waybel_7(X6,X5)|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|(~r1_waybel_3(X5,k12_lattice3(X5,X7,X8),X6)|(r3_orders_2(X5,X7,X6)|r3_orders_2(X5,X8,X6))))))|~m1_subset_1(X6,u1_struct_0(X5)))|~v5_waybel_7(k1_waybel_4(X5),X5))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&(((((m1_subset_1(esk1_2(X5,X6),u1_struct_0(X5))|v4_waybel_7(X6,X5))|~m1_subset_1(X6,u1_struct_0(X5)))|~v5_waybel_7(k1_waybel_4(X5),X5))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&(((((m1_subset_1(esk2_2(X5,X6),u1_struct_0(X5))|v4_waybel_7(X6,X5))|~m1_subset_1(X6,u1_struct_0(X5)))|~v5_waybel_7(k1_waybel_4(X5),X5))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&(((((r1_waybel_3(X5,k12_lattice3(X5,esk1_2(X5,X6),esk2_2(X5,X6)),X6)|v4_waybel_7(X6,X5))|~m1_subset_1(X6,u1_struct_0(X5)))|~v5_waybel_7(k1_waybel_4(X5),X5))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&(((((~r3_orders_2(X5,esk1_2(X5,X6),X6)|v4_waybel_7(X6,X5))|~m1_subset_1(X6,u1_struct_0(X5)))|~v5_waybel_7(k1_waybel_4(X5),X5))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&((((~r3_orders_2(X5,esk2_2(X5,X6),X6)|v4_waybel_7(X6,X5))|~m1_subset_1(X6,u1_struct_0(X5)))|~v5_waybel_7(k1_waybel_4(X5),X5))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5))))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_waybel_7])])])])])])])])).
fof(c_0_4, plain, (![X5]:![X6]:(((((m1_subset_1(esk3_2(X5,X6),u1_struct_0(X5))|v5_waybel_6(X6,X5))|~v5_waybel_7(k1_waybel_4(X5),X5))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&(((((m1_subset_1(esk4_2(X5,X6),u1_struct_0(X5))|v5_waybel_6(X6,X5))|~v5_waybel_7(k1_waybel_4(X5),X5))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&(((((r1_waybel_3(X5,k12_lattice3(X5,esk3_2(X5,X6),esk4_2(X5,X6)),X6)|v5_waybel_6(X6,X5))|~v5_waybel_7(k1_waybel_4(X5),X5))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&(((((~r3_orders_2(X5,esk3_2(X5,X6),X6)|v5_waybel_6(X6,X5))|~v5_waybel_7(k1_waybel_4(X5),X5))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))&((((~r3_orders_2(X5,esk4_2(X5,X6),X6)|v5_waybel_6(X6,X5))|~v5_waybel_7(k1_waybel_4(X5),X5))|~m1_subset_1(X6,u1_struct_0(X5)))|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|(~v1_yellow_0(X5)|(~v1_lattice3(X5)|(~v2_lattice3(X5)|(~v3_waybel_3(X5)|~l1_orders_2(X5)))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l57_waybel_7])])])])])])])])).
fof(c_0_5, negated_conjecture, (~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v3_waybel_3(X1)&l1_orders_2(X1))))))))=>(v5_waybel_7(k1_waybel_4(X1),X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v4_waybel_7(X2,X1)=>v5_waybel_6(X2,X1))))))), inference(assume_negation,[status(cth)],[t49_waybel_7])).
cnf(c_0_6,plain,(r3_orders_2(X1,X3,X2)|r3_orders_2(X1,X4,X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v5_waybel_7(k1_waybel_4(X1),X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_waybel_3(X1,k12_lattice3(X1,X4,X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_waybel_7(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_7,plain,(v5_waybel_6(X2,X1)|r1_waybel_3(X1,k12_lattice3(X1,esk3_2(X1,X2),esk4_2(X1,X2)),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_waybel_7(k1_waybel_4(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(v5_waybel_6(X2,X1)|m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_waybel_7(k1_waybel_4(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,plain,(v5_waybel_6(X2,X1)|m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_waybel_7(k1_waybel_4(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,plain,(v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_waybel_7(k1_waybel_4(X1),X1)|~r3_orders_2(X1,esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,plain,(v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_waybel_7(k1_waybel_4(X1),X1)|~r3_orders_2(X1,esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_12, negated_conjecture, (((v3_orders_2(esk5_0)&(v4_orders_2(esk5_0)&(v5_orders_2(esk5_0)&(v1_yellow_0(esk5_0)&(v1_lattice3(esk5_0)&(v2_lattice3(esk5_0)&(v3_waybel_3(esk5_0)&l1_orders_2(esk5_0))))))))&(v5_waybel_7(k1_waybel_4(esk5_0),esk5_0)&(m1_subset_1(esk6_0,u1_struct_0(esk5_0))&(v4_waybel_7(esk6_0,esk5_0)&~v5_waybel_6(esk6_0,esk5_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
cnf(c_0_13,plain,(v5_waybel_6(X1,X2)|~v4_waybel_7(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_waybel_7(k1_waybel_4(X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v2_lattice3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11])).
cnf(c_0_14,negated_conjecture,(v5_waybel_7(k1_waybel_4(esk5_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15,negated_conjecture,(l1_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(v3_waybel_3(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(v2_lattice3(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,negated_conjecture,(v1_lattice3(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(v1_yellow_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,negated_conjecture,(v5_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,negated_conjecture,(v4_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,negated_conjecture,(v3_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(~v5_waybel_6(esk6_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,negated_conjecture,(v5_waybel_6(X1,esk5_0)|~v4_waybel_7(X1,esk5_0)|~m1_subset_1(X1,u1_struct_0(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])])).
cnf(c_0_25,negated_conjecture,(v4_waybel_7(esk6_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_26,negated_conjecture,(m1_subset_1(esk6_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 21
# Proof object formula steps           : 7
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 48
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 48
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 4
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 4
# Paramodulations                      : 4
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 197
# Rec. Clause-clause subsumption calls : 4
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3192

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 43248 pages
