# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d21_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v11_lattices(X1)&(v16_lattices(X1)&l3_lattices(X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k7_lattices(X1,X2)<=>r2_lattices(X1,X3,X2))))))), file('out/t49_lattices__4_premises__rerun.E_input', d21_lattices)).
fof(d18_lattices, axiom, (![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattices(X1,X2,X3)<=>(k1_lattices(X1,X2,X3)=k6_lattices(X1)&(k1_lattices(X1,X3,X2)=k6_lattices(X1)&(k2_lattices(X1,X2,X3)=k5_lattices(X1)&k2_lattices(X1,X3,X2)=k5_lattices(X1))))))))), file('out/t49_lattices__4_premises__rerun.E_input', d18_lattices)).
fof(dt_k7_lattices, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l3_lattices(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k7_lattices(X1,X2),u1_struct_0(X1)))), file('out/t49_lattices__4_premises__rerun.E_input', dt_k7_lattices)).
fof(t49_lattices, conjecture, (![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k7_lattices(X1,k7_lattices(X1,X2))=X2))), file('out/t49_lattices__4_premises__rerun.E_input', t49_lattices)).
fof(cc5_lattices, axiom, (![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v17_lattices(X1))=>(~(v2_struct_0(X1))&(v11_lattices(X1)&(v15_lattices(X1)&v16_lattices(X1))))))), file('out/t49_lattices__4_premises__rerun.E_input', cc5_lattices)).
fof(c_0_5, plain, (![X4]:![X5]:![X6]:((((((X6!=k7_lattices(X4,X5)|r2_lattices(X4,X6,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|(~v11_lattices(X4)|(~v16_lattices(X4)|~l3_lattices(X4))))))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l3_lattices(X4)))&(((((~r2_lattices(X4,X6,X5)|X6=k7_lattices(X4,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|(v2_struct_0(X4)|(~v10_lattices(X4)|(~v11_lattices(X4)|(~v16_lattices(X4)|~l3_lattices(X4))))))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l3_lattices(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattices])])])])])])])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:((((((k1_lattices(X4,X5,X6)=k6_lattices(X4)|~r2_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l3_lattices(X4)))&(((((k1_lattices(X4,X6,X5)=k6_lattices(X4)|~r2_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l3_lattices(X4)))&(((((k2_lattices(X4,X5,X6)=k5_lattices(X4)|~r2_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l3_lattices(X4)))&((((k2_lattices(X4,X6,X5)=k5_lattices(X4)|~r2_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l3_lattices(X4))))))&(((((k1_lattices(X4,X5,X6)!=k6_lattices(X4)|(k1_lattices(X4,X6,X5)!=k6_lattices(X4)|(k2_lattices(X4,X5,X6)!=k5_lattices(X4)|k2_lattices(X4,X6,X5)!=k5_lattices(X4))))|r2_lattices(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(v2_struct_0(X4)|~l3_lattices(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_lattices])])])])])])])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v2_struct_0(X1)|r2_lattices(X1,X3,X2)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v16_lattices(X1)|~v11_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|X3!=k7_lattices(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, (![X3]:![X4]:(((v2_struct_0(X3)|~l3_lattices(X3))|~m1_subset_1(X4,u1_struct_0(X3)))|m1_subset_1(k7_lattices(X3,X4),u1_struct_0(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_lattices])])])).
cnf(c_0_9,plain,(v2_struct_0(X1)|r2_lattices(X1,X2,X3)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|k2_lattices(X1,X3,X2)!=k5_lattices(X1)|k2_lattices(X1,X2,X3)!=k5_lattices(X1)|k1_lattices(X1,X3,X2)!=k6_lattices(X1)|k1_lattices(X1,X2,X3)!=k6_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(v2_struct_0(X1)|k2_lattices(X1,X2,X3)=k5_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(v2_struct_0(X1)|k2_lattices(X1,X3,X2)=k5_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,plain,(v2_struct_0(X1)|k1_lattices(X1,X3,X2)=k6_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,plain,(v2_struct_0(X1)|k1_lattices(X1,X2,X3)=k6_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,plain,(v2_struct_0(X1)|r2_lattices(X1,X3,X2)|X3!=k7_lattices(X1,X2)|~l3_lattices(X1)|~v10_lattices(X1)|~v11_lattices(X1)|~v16_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_7])).
cnf(c_0_15,plain,(m1_subset_1(k7_lattices(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_16, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k7_lattices(X1,k7_lattices(X1,X2))=X2)))), inference(assume_negation,[status(cth)],[t49_lattices])).
cnf(c_0_17,plain,(v2_struct_0(X1)|v2_struct_0(X1)|X3=k7_lattices(X1,X2)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v16_lattices(X1)|~v11_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattices(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18,plain,(r2_lattices(X1,X2,X3)|v2_struct_0(X1)|~r2_lattices(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13])).
cnf(c_0_19,plain,(r2_lattices(X1,k7_lattices(X1,X2),X2)|v2_struct_0(X1)|~v16_lattices(X1)|~v11_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15])).
fof(c_0_20, negated_conjecture, (((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v17_lattices(esk1_0)&l3_lattices(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&k7_lattices(esk1_0,k7_lattices(esk1_0,esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])])).
cnf(c_0_21,plain,(X3=k7_lattices(X1,X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v11_lattices(X1)|~v16_lattices(X1)|~r2_lattices(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_17])).
cnf(c_0_22,plain,(r2_lattices(X1,X2,k7_lattices(X1,X2))|v2_struct_0(X1)|~v16_lattices(X1)|~v11_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_15])).
cnf(c_0_23,negated_conjecture,(k7_lattices(esk1_0,k7_lattices(esk1_0,esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(k7_lattices(X1,k7_lattices(X1,X2))=X2|v2_struct_0(X1)|~v16_lattices(X1)|~v11_lattices(X1)|~v10_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_15])).
cnf(c_0_25,negated_conjecture,(v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,negated_conjecture,(l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, (![X2]:(((~v2_struct_0(X2)|(v2_struct_0(X2)|~v17_lattices(X2)))|~l3_lattices(X2))&(((v11_lattices(X2)|(v2_struct_0(X2)|~v17_lattices(X2)))|~l3_lattices(X2))&(((v15_lattices(X2)|(v2_struct_0(X2)|~v17_lattices(X2)))|~l3_lattices(X2))&((v16_lattices(X2)|(v2_struct_0(X2)|~v17_lattices(X2)))|~l3_lattices(X2)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_lattices])])])])).
cnf(c_0_30,negated_conjecture,(~v16_lattices(esk1_0)|~v11_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_31,plain,(v2_struct_0(X1)|v16_lattices(X1)|~l3_lattices(X1)|~v17_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_32,negated_conjecture,(v17_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33,negated_conjecture,(~v11_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_27])]), c_0_28])).
cnf(c_0_34,plain,(v2_struct_0(X1)|v11_lattices(X1)|~l3_lattices(X1)|~v17_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_32]), c_0_27])]), c_0_28]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 25
# Proof object formula steps           : 11
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 17
# Processed clauses                    : 25
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 25
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 15
# Contextual simplify-reflections      : 8
# Paramodulations                      : 17
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 56
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 161
# Rec. Clause-clause subsumption calls : 20
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2128

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45304 pages
