# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t49_pre_topc__7_premises__rerun.E_input', t3_subset)).
fof(dt_k2_pre_topc, axiom, (![X1]:![X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t49_pre_topc__7_premises__rerun.E_input', dt_k2_pre_topc)).
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t49_pre_topc__7_premises__rerun.E_input', t1_xboole_1)).
fof(t48_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_tarski(X2,k2_pre_topc(X1,X2))))), file('out/t49_pre_topc__7_premises__rerun.E_input', t48_pre_topc)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t49_pre_topc__7_premises__rerun.E_input', d3_tarski)).
fof(t49_pre_topc, conjecture, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3))))))), file('out/t49_pre_topc__7_premises__rerun.E_input', t49_pre_topc)).
fof(t45_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(r2_hidden(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X4,X1)&r1_tarski(X2,X4))=>r2_hidden(X3,X4)))))))), file('out/t49_pre_topc__7_premises__rerun.E_input', t45_pre_topc)).
fof(projectivity_k2_pre_topc, axiom, (![X1]:![X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>k2_pre_topc(X1,k2_pre_topc(X1,X2))=k2_pre_topc(X1,X2))), file('out/t49_pre_topc__7_premises__rerun.E_input', projectivity_k2_pre_topc)).
fof(c_0_8, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_9, plain, (![X3]:![X4]:((~l1_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|m1_subset_1(k2_pre_topc(X3,X4),k1_zfmisc_1(u1_struct_0(X3))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_11, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|r1_tarski(X4,k2_pre_topc(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_pre_topc])])])])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk1_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk1_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_13,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(r1_tarski(X1,k2_pre_topc(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, negated_conjecture, (~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3)))))))), inference(assume_negation,[status(cth)],[t49_pre_topc])).
fof(c_0_18, plain, (![X5]:![X6]:![X7]:![X8]:(((((~r2_hidden(X7,k2_pre_topc(X5,X6))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5)))|((~v4_pre_topc(X8,X5)|~r1_tarski(X6,X8))|r2_hidden(X7,X8))))|~r2_hidden(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))&(((((m1_subset_1(esk5_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(X5)))|r2_hidden(X7,k2_pre_topc(X5,X6)))|~r2_hidden(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))&((((((v4_pre_topc(esk5_3(X5,X6,X7),X5)|r2_hidden(X7,k2_pre_topc(X5,X6)))|~r2_hidden(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5))&((((r1_tarski(X6,esk5_3(X5,X6,X7))|r2_hidden(X7,k2_pre_topc(X5,X6)))|~r2_hidden(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5)))&((((~r2_hidden(X7,esk5_3(X5,X6,X7))|r2_hidden(X7,k2_pre_topc(X5,X6)))|~r2_hidden(X7,u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|~l1_pre_topc(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_pre_topc])])])])])])])).
cnf(c_0_19,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(r1_tarski(k2_pre_topc(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_21,plain,(r1_tarski(X1,k2_pre_topc(X2,X3))|~l1_pre_topc(X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_22, negated_conjecture, ((l1_pre_topc(esk2_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(r1_tarski(esk3_0,esk4_0)&~r1_tarski(k2_pre_topc(esk2_0,esk3_0),k2_pre_topc(esk2_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).
cnf(c_0_23,plain,(r2_hidden(X3,X4)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,u1_struct_0(X1))|~r1_tarski(X2,X4)|~v4_pre_topc(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,k2_pre_topc(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(r2_hidden(X3,k2_pre_topc(X1,X2))|m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,plain,(r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,k2_pre_topc(X2,X3))|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26,plain,(r2_hidden(X3,k2_pre_topc(X1,X2))|v4_pre_topc(esk5_3(X1,X2,X3),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(r2_hidden(X3,k2_pre_topc(X1,X2))|r1_tarski(X2,esk5_3(X1,X2,X3))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28,plain,(r1_tarski(X1,k2_pre_topc(X2,X3))|~l1_pre_topc(X2)|~r1_tarski(X1,X4)|~r1_tarski(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_15, c_0_21])).
cnf(c_0_29,negated_conjecture,(r1_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,plain,(r2_hidden(X1,esk5_3(X2,X3,X4))|r2_hidden(X4,k2_pre_topc(X2,X3))|~r2_hidden(X1,k2_pre_topc(X2,X5))|~r2_hidden(X4,u1_struct_0(X2))|~l1_pre_topc(X2)|~r1_tarski(X5,esk5_3(X2,X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_31,plain,(r2_hidden(X1,k2_pre_topc(X2,X3))|r1_tarski(X4,esk5_3(X2,X3,X1))|~r2_hidden(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~r1_tarski(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_15, c_0_27])).
cnf(c_0_32,negated_conjecture,(r1_tarski(esk3_0,k2_pre_topc(X1,X2))|~l1_pre_topc(X1)|~r1_tarski(esk4_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33,plain,(r2_hidden(X1,esk5_3(X2,X3,X4))|r2_hidden(X4,k2_pre_topc(X2,X3))|~r2_hidden(X1,k2_pre_topc(X2,X5))|~r2_hidden(X4,u1_struct_0(X2))|~l1_pre_topc(X2)|~r1_tarski(X5,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,negated_conjecture,(r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,esk3_0)|~l1_pre_topc(X2)|~r1_tarski(esk4_0,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_19, c_0_32])).
cnf(c_0_35,plain,(r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_36,plain,(r1_tarski(X1,X2)|r2_hidden(esk1_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_38,negated_conjecture,(r2_hidden(X1,esk5_3(X2,X3,X4))|r2_hidden(X4,k2_pre_topc(X2,X3))|~r2_hidden(X4,u1_struct_0(X2))|~r2_hidden(X1,esk3_0)|~l1_pre_topc(X2)|~r1_tarski(esk4_0,X5)|~r1_tarski(X5,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_40,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_41,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42,negated_conjecture,(r1_tarski(esk3_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_13, c_0_37])).
cnf(c_0_43,plain,(r2_hidden(X3,k2_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,u1_struct_0(X1))|~r2_hidden(X3,esk5_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_44,negated_conjecture,(r2_hidden(X1,esk5_3(esk2_0,X2,X3))|r2_hidden(X3,k2_pre_topc(esk2_0,X2))|~r2_hidden(X3,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)|~r1_tarski(esk4_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41])])).
cnf(c_0_45,negated_conjecture,(r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_19, c_0_42])).
cnf(c_0_46,negated_conjecture,(r2_hidden(X1,k2_pre_topc(esk2_0,X2))|~r2_hidden(X1,esk3_0)|~r1_tarski(esk4_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_40])]), c_0_45])).
cnf(c_0_47,negated_conjecture,(r1_tarski(X1,k2_pre_topc(esk2_0,X2))|~r2_hidden(esk1_2(X1,k2_pre_topc(esk2_0,X2)),esk3_0)|~r1_tarski(esk4_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_35, c_0_46])).
fof(c_0_48, plain, (![X3]:![X4]:((~l1_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|k2_pre_topc(X3,k2_pre_topc(X3,X4))=k2_pre_topc(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[projectivity_k2_pre_topc])])).
cnf(c_0_49,negated_conjecture,(r1_tarski(esk3_0,k2_pre_topc(esk2_0,X1))|~r1_tarski(esk4_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_47, c_0_36])).
cnf(c_0_50,plain,(k2_pre_topc(X1,k2_pre_topc(X1,X2))=k2_pre_topc(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_51,negated_conjecture,(r1_tarski(esk3_0,k2_pre_topc(esk2_0,X1))|~r1_tarski(esk4_0,k2_pre_topc(esk2_0,X1))|~m1_subset_1(k2_pre_topc(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_40])])).
cnf(c_0_52,negated_conjecture,(r1_tarski(esk3_0,k2_pre_topc(esk2_0,X1))|~r1_tarski(esk4_0,k2_pre_topc(esk2_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_14]), c_0_40])])).
cnf(c_0_53,plain,(r2_hidden(X1,esk5_3(X2,X3,X4))|r2_hidden(X4,k2_pre_topc(X2,X3))|~r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X4,u1_struct_0(X2))|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_30, c_0_27])).
cnf(c_0_54,negated_conjecture,(r1_tarski(esk3_0,k2_pre_topc(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_16]), c_0_39]), c_0_40])])).
cnf(c_0_55,plain,(r2_hidden(X1,k2_pre_topc(X2,X3))|r1_tarski(X4,esk5_3(X2,X3,X1))|~r2_hidden(esk1_2(X4,esk5_3(X2,X3,X1)),k2_pre_topc(X2,X3))|~r2_hidden(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_35, c_0_53])).
cnf(c_0_56,negated_conjecture,(r2_hidden(X1,k2_pre_topc(esk2_0,esk4_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_19, c_0_54])).
cnf(c_0_57,negated_conjecture,(r2_hidden(X1,k2_pre_topc(esk2_0,esk4_0))|r1_tarski(X2,esk5_3(esk2_0,esk4_0,X1))|~r2_hidden(esk1_2(X2,esk5_3(esk2_0,esk4_0,X1)),esk3_0)|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_40]), c_0_39])])).
cnf(c_0_58,negated_conjecture,(r2_hidden(X1,k2_pre_topc(esk2_0,esk4_0))|r1_tarski(esk3_0,esk5_3(esk2_0,esk4_0,X1))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_57, c_0_36])).
cnf(c_0_59,negated_conjecture,(r2_hidden(X1,esk5_3(esk2_0,esk4_0,X2))|r2_hidden(X2,k2_pre_topc(esk2_0,esk4_0))|~r2_hidden(X1,k2_pre_topc(esk2_0,esk3_0))|~r2_hidden(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_58]), c_0_40]), c_0_37]), c_0_39])])).
cnf(c_0_60,negated_conjecture,(r2_hidden(X1,k2_pre_topc(esk2_0,esk4_0))|~r2_hidden(X1,k2_pre_topc(esk2_0,esk3_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_59]), c_0_40]), c_0_39])])).
cnf(c_0_61,negated_conjecture,(r1_tarski(X1,k2_pre_topc(esk2_0,esk4_0))|~r2_hidden(esk1_2(X1,k2_pre_topc(esk2_0,esk4_0)),k2_pre_topc(esk2_0,esk3_0))|~r2_hidden(esk1_2(X1,k2_pre_topc(esk2_0,esk4_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_35, c_0_60])).
cnf(c_0_62,negated_conjecture,(~r1_tarski(k2_pre_topc(esk2_0,esk3_0),k2_pre_topc(esk2_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_63,negated_conjecture,(~r2_hidden(esk1_2(k2_pre_topc(esk2_0,esk3_0),k2_pre_topc(esk2_0,esk4_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_36]), c_0_62])).
cnf(c_0_64,plain,(r2_hidden(esk1_2(k2_pre_topc(X1,X2),X3),u1_struct_0(X1))|r1_tarski(k2_pre_topc(X1,X2),X3)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_25, c_0_36])).
cnf(c_0_65,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_40]), c_0_37])]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 49
# Proof object formula steps           : 17
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 12136
# ...of these trivial                  : 2
# ...subsumed                          : 9899
# ...remaining for further processing  : 2235
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 113
# Backward-rewritten                   : 1
# Generated clauses                    : 63984
# ...of the previous two non-trivial   : 62768
# Contextual simplify-reflections      : 21005
# Paramodulations                      : 63984
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 2121
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 2102
# Current number of unprocessed clauses: 49633
# ...number of literals in the above   : 591007
# Current number of archived formulas  : 0
# Current number of archived clauses   : 114
# Clause-clause subsumption calls (NU) : 2341867
# Rec. Clause-clause subsumption calls : 326034
# Non-unit clause-clause subsumptions  : 30167
# Unit Clause-clause subsumption calls : 2392
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2600308

# -------------------------------------------------
# User time                : 3.704 s
# System time              : 0.024 s
# Total time               : 3.728 s
# Maximum resident set size: 62488 pages
