# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_setfam_1, axiom, (![X1]:![X2]:((~(X1=k1_xboole_0)=>(X2=k1_setfam_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>![X4]:(r2_hidden(X4,X1)=>r2_hidden(X3,X4)))))&(X1=k1_xboole_0=>(X2=k1_setfam_1(X1)<=>X2=k1_xboole_0)))), file('out/t4_setfam_1__3_premises__rerun.E_input', d1_setfam_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t4_setfam_1__3_premises__rerun.E_input', d3_tarski)).
fof(t4_setfam_1, conjecture, (![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1))), file('out/t4_setfam_1__3_premises__rerun.E_input', t4_setfam_1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t4_setfam_1__3_premises__rerun.E_input', t2_xboole_1)).
fof(c_0_4, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:![X6]:![X12]:![X5]:![X6]:![X6]:((((((~r2_hidden(X7,X6)|(~r2_hidden(X8,X5)|r2_hidden(X7,X8)))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&((((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(X7,esk1_3(X5,X6,X7))|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((((r2_hidden(esk3_2(X5,X6),X5)|~r2_hidden(esk2_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(esk2_2(X5,X6),esk3_2(X5,X6))|~r2_hidden(esk2_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0))&(((r2_hidden(esk2_2(X5,X6),X6)|(~r2_hidden(X12,X5)|r2_hidden(esk2_2(X5,X6),X12)))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((X6!=k1_setfam_1(X5)|X6=k1_xboole_0)|X5!=k1_xboole_0)&((X6!=k1_xboole_0|X6=k1_setfam_1(X5))|X5!=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_setfam_1])])])])])])])])).
cnf(c_0_5,plain,(X1=k1_xboole_0|r2_hidden(X3,X4)|X2!=k1_setfam_1(X1)|~r2_hidden(X4,X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk4_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk4_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_7, negated_conjecture, (~(![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1)))), inference(assume_negation,[status(cth)],[t4_setfam_1])).
cnf(c_0_8,plain,(X1=k1_xboole_0|r2_hidden(X2,X3)|~r2_hidden(X2,k1_setfam_1(X1))|~r2_hidden(X3,X1)), inference(er,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(r1_tarski(X1,X2)|r2_hidden(esk4_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, ((r2_hidden(esk5_0,esk6_0)&~r1_tarski(k1_setfam_1(esk6_0),esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_11,plain,(X1=k1_xboole_0|r1_tarski(k1_setfam_1(X1),X2)|r2_hidden(esk4_2(k1_setfam_1(X1),X2),X3)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12,negated_conjecture,(r2_hidden(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13,plain,(X2=k1_xboole_0|X1!=k1_xboole_0|X2!=k1_setfam_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14,plain,(r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(esk6_0=k1_xboole_0|r1_tarski(k1_setfam_1(esk6_0),X1)|r2_hidden(esk4_2(k1_setfam_1(esk6_0),X1),esk5_0)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16,negated_conjecture,(~r1_tarski(k1_setfam_1(esk6_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,plain,(k1_setfam_1(X1)=k1_xboole_0|X1!=k1_xboole_0), inference(er,[status(thm)],[c_0_13])).
fof(c_0_18, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_19,negated_conjecture,(esk6_0=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_20,plain,(k1_setfam_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_19]), c_0_20]), c_0_21])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 14
# Proof object formula steps           : 9
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 14
# Processed clauses                    : 24
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 23
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 19
# ...of the previous two non-trivial   : 16
# Contextual simplify-reflections      : 0
# Paramodulations                      : 13
# Factorizations                       : 0
# Equation resolutions                 : 6
# Current number of processed clauses  : 20
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 12
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 30
# Rec. Clause-clause subsumption calls : 17
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1269

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 43064 pages
