# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_setfam_1, axiom, (![X1]:![X2]:((~(X1=k1_xboole_0)=>(X2=k1_setfam_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>![X4]:(r2_hidden(X4,X1)=>r2_hidden(X3,X4)))))&(X1=k1_xboole_0=>(X2=k1_setfam_1(X1)<=>X2=k1_xboole_0)))), file('out/t4_setfam_1__5_premises__rerun.E_input', d1_setfam_1)).
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t4_setfam_1__5_premises__rerun.E_input', t4_xboole_0)).
fof(t7_xboole_0, axiom, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1))))), file('out/t4_setfam_1__5_premises__rerun.E_input', t7_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t4_setfam_1__5_premises__rerun.E_input', d3_tarski)).
fof(t4_setfam_1, conjecture, (![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1))), file('out/t4_setfam_1__5_premises__rerun.E_input', t4_setfam_1)).
fof(t65_xboole_1, axiom, (![X1]:r1_xboole_0(X1,k1_xboole_0)), file('out/t4_setfam_1__5_premises__rerun.E_input', t65_xboole_1)).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:![X6]:![X12]:![X5]:![X6]:![X6]:((((((~r2_hidden(X7,X6)|(~r2_hidden(X8,X5)|r2_hidden(X7,X8)))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&((((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(X7,esk1_3(X5,X6,X7))|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((((r2_hidden(esk3_2(X5,X6),X5)|~r2_hidden(esk2_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(esk2_2(X5,X6),esk3_2(X5,X6))|~r2_hidden(esk2_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0))&(((r2_hidden(esk2_2(X5,X6),X6)|(~r2_hidden(X12,X5)|r2_hidden(esk2_2(X5,X6),X12)))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((X6!=k1_setfam_1(X5)|X6=k1_xboole_0)|X5!=k1_xboole_0)&((X6!=k1_xboole_0|X6=k1_setfam_1(X5))|X5!=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_setfam_1])])])])])])])])).
fof(c_0_7, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk4_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
fof(c_0_8, plain, (![X3]:(X3=k1_xboole_0|r2_hidden(esk5_1(X3),X3))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])])])).
cnf(c_0_9,plain,(X1=k1_xboole_0|r2_hidden(X3,X4)|X2!=k1_setfam_1(X1)|~r2_hidden(X4,X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk6_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk6_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_11,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(r2_hidden(esk5_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_13, negated_conjecture, (~(![X1]:![X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1)))), inference(assume_negation,[status(cth)],[t4_setfam_1])).
cnf(c_0_14,plain,(X1=k1_xboole_0|r2_hidden(X2,X3)|~r2_hidden(X2,k1_setfam_1(X1))|~r2_hidden(X3,X1)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(r1_tarski(X1,X2)|r2_hidden(esk6_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_17, negated_conjecture, ((r2_hidden(esk7_0,esk8_0)&~r1_tarski(k1_setfam_1(esk8_0),esk7_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_18,plain,(r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,plain,(X1=k1_xboole_0|r1_tarski(k1_setfam_1(X1),X2)|r2_hidden(esk6_2(k1_setfam_1(X1),X2),X3)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_11, c_0_16])).
fof(c_0_21, plain, (![X2]:r1_xboole_0(X2,k1_xboole_0)), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_22,negated_conjecture,(~r1_tarski(k1_setfam_1(esk8_0),esk7_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(X1=k1_xboole_0|r1_tarski(k1_setfam_1(X1),X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24,negated_conjecture,(r2_hidden(esk7_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(X2=k1_xboole_0|X1!=k1_xboole_0|X2!=k1_setfam_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26,plain,(r1_tarski(k1_xboole_0,X1)|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_20, c_0_15])).
cnf(c_0_27,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28,negated_conjecture,(esk8_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])])).
cnf(c_0_29,plain,(k1_setfam_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_25])])).
cnf(c_0_30,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_28]), c_0_29]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 32
# Proof object clause steps            : 19
# Proof object formula steps           : 13
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 63
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 57
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 7
# Generated clauses                    : 65
# ...of the previous two non-trivial   : 57
# Contextual simplify-reflections      : 7
# Paramodulations                      : 58
# Factorizations                       : 2
# Equation resolutions                 : 7
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 104
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 275
# Rec. Clause-clause subsumption calls : 139
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2265

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 43064 pages
