# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_yellow_0, conjecture, (![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3)))))))), file('out/t4_yellow_0__3_premises__rerun.E_input', t4_yellow_0)).
fof(d2_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>(v4_orders_2(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X4))=>r1_orders_2(X1,X2,X4)))))))), file('out/t4_yellow_0__3_premises__rerun.E_input', d2_yellow_0)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t4_yellow_0__3_premises__rerun.E_input', d9_lattice3)).
fof(d8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4))))))), file('out/t4_yellow_0__3_premises__rerun.E_input', d8_lattice3)).
fof(c_0_4, negated_conjecture, (~(![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3))))))))), inference(assume_negation,[status(cth)],[t4_yellow_0])).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:(((~v4_orders_2(X5)|(~m1_subset_1(X6,u1_struct_0(X5))|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|((~r1_orders_2(X5,X6,X7)|~r1_orders_2(X5,X7,X8))|r1_orders_2(X5,X6,X8))))))|~l1_orders_2(X5))&(((m1_subset_1(esk6_1(X5),u1_struct_0(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&(((m1_subset_1(esk7_1(X5),u1_struct_0(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&(((m1_subset_1(esk8_1(X5),u1_struct_0(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&((((r1_orders_2(X5,esk6_1(X5),esk7_1(X5))|v4_orders_2(X5))|~l1_orders_2(X5))&((r1_orders_2(X5,esk7_1(X5),esk8_1(X5))|v4_orders_2(X5))|~l1_orders_2(X5)))&((~r1_orders_2(X5,esk6_1(X5),esk8_1(X5))|v4_orders_2(X5))|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_yellow_0])])])])])])])).
fof(c_0_6, negated_conjecture, (((v4_orders_2(esk1_0)&l1_orders_2(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(r1_orders_2(esk1_0,esk2_0,esk3_0)&(((r2_lattice3(esk1_0,esk5_0,esk2_0)|r1_lattice3(esk1_0,esk4_0,esk3_0))&(~r2_lattice3(esk1_0,esk5_0,esk3_0)|r1_lattice3(esk1_0,esk4_0,esk3_0)))&((r2_lattice3(esk1_0,esk5_0,esk2_0)|~r1_lattice3(esk1_0,esk4_0,esk2_0))&(~r2_lattice3(esk1_0,esk5_0,esk3_0)|~r1_lattice3(esk1_0,esk4_0,esk2_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk9_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk9_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk9_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
cnf(c_0_8,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,negated_conjecture,(r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,negated_conjecture,(r1_lattice3(esk1_0,esk4_0,esk3_0)|r2_lattice3(esk1_0,esk5_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk9_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,negated_conjecture,(r1_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_18,negated_conjecture,(r1_lattice3(esk1_0,esk4_0,esk3_0)|r1_orders_2(esk1_0,X1,esk2_0)|~r2_hidden(X1,esk5_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_10]), c_0_12])])).
cnf(c_0_19,plain,(r2_lattice3(X1,X3,X2)|r2_hidden(esk9_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20,negated_conjecture,(r2_lattice3(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,esk9_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk9_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_11]), c_0_12])])).
cnf(c_0_21,negated_conjecture,(r2_lattice3(X1,esk5_0,X2)|r1_lattice3(esk1_0,esk4_0,esk3_0)|r1_orders_2(esk1_0,esk9_3(X1,esk5_0,X2),esk2_0)|~m1_subset_1(esk9_3(X1,esk5_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22,negated_conjecture,(r1_lattice3(esk1_0,esk4_0,esk3_0)|~r2_lattice3(esk1_0,esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_23, plain, (![X5]:![X6]:![X7]:![X8]:((((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk10_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk10_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,X7,esk10_3(X5,X6,X7))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])])])).
cnf(c_0_24,negated_conjecture,(r1_lattice3(esk1_0,esk4_0,esk3_0)|~m1_subset_1(esk9_3(esk1_0,esk5_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_11]), c_0_12])]), c_0_22])).
cnf(c_0_25,plain,(r2_lattice3(X1,X3,X2)|m1_subset_1(esk9_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_26,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27,negated_conjecture,(r1_lattice3(esk1_0,esk4_0,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_11]), c_0_12])]), c_0_22])).
cnf(c_0_28,negated_conjecture,(r2_lattice3(esk1_0,esk5_0,esk2_0)|~r1_lattice3(esk1_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_29,negated_conjecture,(r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_11]), c_0_12])])).
cnf(c_0_30,plain,(r1_lattice3(X1,X3,X2)|r2_hidden(esk10_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,negated_conjecture,(r1_orders_2(esk1_0,X1,esk2_0)|~r2_hidden(X1,esk5_0)|~r1_lattice3(esk1_0,esk4_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_28]), c_0_10]), c_0_12])])).
cnf(c_0_32,negated_conjecture,(r1_lattice3(X1,esk4_0,X2)|r1_orders_2(esk1_0,esk3_0,esk10_3(X1,esk4_0,X2))|~m1_subset_1(esk10_3(X1,esk4_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33,negated_conjecture,(r2_lattice3(X1,esk5_0,X2)|r1_orders_2(esk1_0,esk9_3(X1,esk5_0,X2),esk2_0)|~r1_lattice3(esk1_0,esk4_0,esk2_0)|~m1_subset_1(esk9_3(X1,esk5_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_31, c_0_19])).
cnf(c_0_34,negated_conjecture,(~r1_lattice3(esk1_0,esk4_0,esk2_0)|~r2_lattice3(esk1_0,esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_35,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk10_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36,negated_conjecture,(r1_lattice3(X1,esk4_0,X2)|r1_orders_2(esk1_0,X3,esk10_3(X1,esk4_0,X2))|~r1_orders_2(esk1_0,X3,esk3_0)|~m1_subset_1(esk10_3(X1,esk4_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_32]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_37,negated_conjecture,(~r1_lattice3(esk1_0,esk4_0,esk2_0)|~m1_subset_1(esk9_3(esk1_0,esk5_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_33]), c_0_11]), c_0_12])]), c_0_34])).
cnf(c_0_38,negated_conjecture,(r1_lattice3(esk1_0,esk4_0,X1)|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(esk10_3(esk1_0,esk4_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_12])])).
cnf(c_0_39,plain,(r1_lattice3(X1,X3,X2)|m1_subset_1(esk10_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40,negated_conjecture,(~r1_lattice3(esk1_0,esk4_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_25]), c_0_11]), c_0_12])]), c_0_34])).
cnf(c_0_41,negated_conjecture,(r1_lattice3(esk1_0,esk4_0,X1)|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_12])])).
cnf(c_0_42,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_9]), c_0_10])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 34
# Proof object formula steps           : 9
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 74
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 6
# Generated clauses                    : 39
# ...of the previous two non-trivial   : 35
# Contextual simplify-reflections      : 6
# Paramodulations                      : 39
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 37
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 669
# Rec. Clause-clause subsumption calls : 85
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 11
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3495

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45256 pages
