# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t51_subset_1, conjecture, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>~(r2_hidden(X4,X3))))=>X2=k3_subset_1(X1,X3))))), file('out/t51_subset_1__9_premises__rerun.E_input', t51_subset_1)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t51_subset_1__9_premises__rerun.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, (![X1]:~(v1_xboole_0(k1_zfmisc_1(X1)))), file('out/t51_subset_1__9_premises__rerun.E_input', fc1_subset_1)).
fof(d1_zfmisc_1, axiom, (![X1]:![X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1)))), file('out/t51_subset_1__9_premises__rerun.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t51_subset_1__9_premises__rerun.E_input', d3_tarski)).
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t51_subset_1__9_premises__rerun.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t51_subset_1__9_premises__rerun.E_input', d1_xboole_0)).
fof(dt_k3_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1)))), file('out/t51_subset_1__9_premises__rerun.E_input', dt_k3_subset_1)).
fof(d5_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2))), file('out/t51_subset_1__9_premises__rerun.E_input', d5_subset_1)).
fof(involutiveness_k3_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,k3_subset_1(X1,X2))=X2)), file('out/t51_subset_1__9_premises__rerun.E_input', involutiveness_k3_subset_1)).
fof(c_0_10, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>~(r2_hidden(X4,X3))))=>X2=k3_subset_1(X1,X3)))))), inference(assume_negation,[status(cth)],[t51_subset_1])).
fof(c_0_11, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
fof(c_0_12, negated_conjecture, (![X8]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((((~r2_hidden(X8,esk2_0)|~r2_hidden(X8,esk3_0))|~m1_subset_1(X8,esk1_0))&((r2_hidden(X8,esk3_0)|r2_hidden(X8,esk2_0))|~m1_subset_1(X8,esk1_0)))&esk2_0!=k3_subset_1(esk1_0,esk3_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])])])).
fof(c_0_13, plain, (![X2]:~v1_xboole_0(k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|r1_tarski(X6,X4))|X5!=k1_zfmisc_1(X4))&((~r1_tarski(X6,X4)|r2_hidden(X6,X5))|X5!=k1_zfmisc_1(X4)))&(((~r2_hidden(esk4_2(X4,X5),X5)|~r1_tarski(esk4_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))&((r2_hidden(esk4_2(X4,X5),X5)|r1_tarski(esk4_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])])).
cnf(c_0_15,plain,(v1_xboole_0(X1)|r2_hidden(X2,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(r1_tarski(X3,X2)|X1!=k1_zfmisc_1(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,negated_conjecture,(r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_21,negated_conjecture,(r1_tarski(esk3_0,X1)|k1_zfmisc_1(esk1_0)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_22, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk6_3(X5,X6,X7),X7)|(~r2_hidden(esk6_3(X5,X6,X7),X5)|r2_hidden(esk6_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk6_3(X5,X6,X7),X5)|r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk6_3(X5,X6,X7),X6)|r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
fof(c_0_23, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk7_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_24,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,negated_conjecture,(r1_tarski(esk3_0,esk1_0)), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_26,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk6_3(X2,X3,X1),X1)|r2_hidden(esk6_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30,plain,(k4_xboole_0(X1,X2)=X1|r2_hidden(esk6_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_26])).
cnf(c_0_31,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk6_3(X2,X3,X1),X3)|~r2_hidden(esk6_3(X2,X3,X1),X2)|~r2_hidden(esk6_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34,negated_conjecture,(k4_xboole_0(esk3_0,X1)=esk3_0|r2_hidden(esk6_3(esk3_0,X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,negated_conjecture,(k4_xboole_0(X1,X2)=esk3_0|r2_hidden(esk6_3(X1,X2,esk3_0),esk1_0)|r2_hidden(esk6_3(X1,X2,esk3_0),X1)), inference(spm,[status(thm)],[c_0_29, c_0_26])).
cnf(c_0_36,plain,(k4_xboole_0(X1,X2)=X1|X3!=k4_xboole_0(X4,X1)|~r2_hidden(esk6_3(X1,X2,X1),X3)), inference(spm,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_37,plain,(k4_xboole_0(X1,X2)=X1|r2_hidden(esk6_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_30]), c_0_30])).
cnf(c_0_38,negated_conjecture,(~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_39,negated_conjecture,(k4_xboole_0(esk3_0,X1)=esk3_0|m1_subset_1(esk6_3(esk3_0,X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40,negated_conjecture,(k4_xboole_0(esk1_0,X1)=esk3_0|r2_hidden(esk6_3(esk1_0,X1,esk3_0),esk1_0)), inference(ef,[status(thm)],[c_0_35])).
fof(c_0_41, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(X3))|m1_subset_1(k3_subset_1(X3,X4),k1_zfmisc_1(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_subset_1])])).
fof(c_0_42, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(X3))|k3_subset_1(X3,X4)=k4_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_43, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(X3))|k3_subset_1(X3,k3_subset_1(X3,X4))=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k3_subset_1])])).
cnf(c_0_44,plain,(k4_xboole_0(X1,X2)=X1|X2!=k4_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_45,negated_conjecture,(k4_xboole_0(esk3_0,X1)=esk3_0|~r2_hidden(esk6_3(esk3_0,X1,esk3_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_30])).
cnf(c_0_46,negated_conjecture,(r2_hidden(X1,esk2_0)|r2_hidden(X1,esk3_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_47,negated_conjecture,(k4_xboole_0(esk1_0,X1)=esk3_0|m1_subset_1(esk6_3(esk1_0,X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_40])).
cnf(c_0_48,plain,(m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50,plain,(k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51,plain,(k3_subset_1(X1,k3_subset_1(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_53,negated_conjecture,(k4_xboole_0(esk3_0,esk2_0)=esk3_0), inference(spm,[status(thm)],[c_0_45, c_0_37])).
cnf(c_0_54,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk6_3(X2,X3,X1),X1)|~r2_hidden(esk6_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_55,negated_conjecture,(k4_xboole_0(esk1_0,X1)=esk3_0|r2_hidden(esk6_3(esk1_0,X1,esk3_0),esk2_0)|r2_hidden(esk6_3(esk1_0,X1,esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_56,negated_conjecture,(m1_subset_1(k3_subset_1(esk1_0,esk2_0),k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57,negated_conjecture,(k3_subset_1(esk1_0,esk2_0)=k4_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_49])).
cnf(c_0_58,negated_conjecture,(k3_subset_1(esk1_0,k3_subset_1(esk1_0,esk2_0))=esk2_0), inference(spm,[status(thm)],[c_0_51, c_0_49])).
cnf(c_0_59,negated_conjecture,(k4_xboole_0(esk2_0,esk3_0)=esk2_0), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_60,negated_conjecture,(k4_xboole_0(esk1_0,esk2_0)=esk3_0|r2_hidden(esk6_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61,negated_conjecture,(m1_subset_1(k4_xboole_0(esk1_0,esk2_0),k1_zfmisc_1(esk1_0))), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62,negated_conjecture,(k3_subset_1(esk1_0,k4_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[c_0_58, c_0_57])).
cnf(c_0_63,negated_conjecture,(X1!=esk2_0|~r2_hidden(X2,esk3_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_59])).
cnf(c_0_64,negated_conjecture,(k4_xboole_0(esk1_0,esk2_0)=esk3_0|r2_hidden(esk6_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_60]), c_0_40])).
cnf(c_0_65,negated_conjecture,(esk2_0!=k3_subset_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_66,negated_conjecture,(k3_subset_1(esk1_0,esk3_0)=k4_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_50, c_0_16])).
cnf(c_0_67,negated_conjecture,(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_61]), c_0_62])).
cnf(c_0_68,negated_conjecture,(k4_xboole_0(esk1_0,esk2_0)=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_60])).
cnf(c_0_69,negated_conjecture,(k4_xboole_0(esk1_0,esk3_0)!=esk2_0), inference(rw,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_70,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 50
# Proof object formula steps           : 21
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 9453
# ...of these trivial                  : 71
# ...subsumed                          : 7513
# ...remaining for further processing  : 1869
# Other redundant clauses eliminated   : 174
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 151
# Backward-rewritten                   : 132
# Generated clauses                    : 32528
# ...of the previous two non-trivial   : 30788
# Contextual simplify-reflections      : 1522
# Paramodulations                      : 31945
# Factorizations                       : 80
# Equation resolutions                 : 503
# Current number of processed clauses  : 1586
#    Positive orientable unit clauses  : 68
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 30
#    Non-unit-clauses                  : 1486
# Current number of unprocessed clauses: 19168
# ...number of literals in the above   : 73803
# Current number of archived formulas  : 0
# Current number of archived clauses   : 283
# Clause-clause subsumption calls (NU) : 924772
# Rec. Clause-clause subsumption calls : 715298
# Non-unit clause-clause subsumptions  : 7853
# Unit Clause-clause subsumption calls : 4103
# Rewrite failures with RHS unbound    : 27
# BW rewrite match attempts            : 283
# BW rewrite match successes           : 33
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 436081

# -------------------------------------------------
# User time                : 0.692 s
# System time              : 0.012 s
# Total time               : 0.704 s
# Maximum resident set size: 45256 pages
