# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t51_xboole_1__6_premises__rerun.E_input', d5_xboole_0)).
fof(l36_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(X1,k3_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), file('out/t51_xboole_1__6_premises__rerun.E_input', l36_xboole_1)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t51_xboole_1__6_premises__rerun.E_input', t48_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), file('out/t51_xboole_1__6_premises__rerun.E_input', commutativity_k2_xboole_0)).
fof(t47_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), file('out/t51_xboole_1__6_premises__rerun.E_input', t47_xboole_1)).
fof(t51_xboole_1, conjecture, (![X1]:![X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), file('out/t51_xboole_1__6_premises__rerun.E_input', t51_xboole_1)).
fof(t39_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), file('out/t51_xboole_1__6_premises__rerun.E_input', t39_xboole_1)).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
fof(c_0_8, plain, (![X4]:![X5]:![X6]:k4_xboole_0(X4,k3_xboole_0(X5,X6))=k2_xboole_0(k4_xboole_0(X4,X5),k4_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[l36_xboole_1])).
fof(c_0_9, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_10,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(k4_xboole_0(X1,k3_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), inference(rw,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_15, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k2_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_16,plain,(~r2_hidden(X1,k2_xboole_0(k4_xboole_0(X2,X3),k4_xboole_0(X2,X4)))|~r2_hidden(X1,k4_xboole_0(X3,k4_xboole_0(X3,X4)))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17,plain,(k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, (![X3]:![X4]:k4_xboole_0(X3,k3_xboole_0(X3,X4))=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_19,plain,(~r2_hidden(X1,k2_xboole_0(k4_xboole_0(X2,X3),k4_xboole_0(X2,X4)))|~r2_hidden(X1,k4_xboole_0(X4,k4_xboole_0(X4,X3)))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,plain,(k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_21, negated_conjecture, (~(![X1]:![X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1)), inference(assume_negation,[status(cth)],[t51_xboole_1])).
cnf(c_0_22,plain,(~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X3,k4_xboole_0(X3,X4))))|~r2_hidden(X1,k4_xboole_0(X4,k4_xboole_0(X4,X3)))), inference(spm,[status(thm)],[c_0_19, c_0_14])).
cnf(c_0_23,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_12])).
cnf(c_0_24,plain,(r2_hidden(X4,X2)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_25, plain, (![X3]:![X4]:k2_xboole_0(X3,k4_xboole_0(X4,X3))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_26, negated_conjecture, (k2_xboole_0(k3_xboole_0(esk2_0,esk3_0),k4_xboole_0(esk2_0,esk3_0))!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
cnf(c_0_27,plain,(~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_29,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|~r2_hidden(esk1_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|r2_hidden(esk1_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31,plain,(k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,negated_conjecture,(k2_xboole_0(k3_xboole_0(esk2_0,esk3_0),k4_xboole_0(esk2_0,esk3_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,plain,(~r2_hidden(X1,k4_xboole_0(k4_xboole_0(X2,X3),X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_23]), c_0_28])).
cnf(c_0_34,plain,(X1=k4_xboole_0(X2,X2)|r2_hidden(esk1_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,plain,(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,k4_xboole_0(X1,X2))))=k2_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_14]), c_0_17])).
cnf(c_0_36,negated_conjecture,(k2_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)),k4_xboole_0(esk2_0,esk3_0))!=esk2_0), inference(rw,[status(thm)],[c_0_32, c_0_12])).
cnf(c_0_37,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X1)=k4_xboole_0(X3,X3)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38,plain,(k4_xboole_0(X1,k2_xboole_0(X1,X1))=k4_xboole_0(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_35]), c_0_31]), c_0_23])).
cnf(c_0_39,plain,(k2_xboole_0(X1,X1)=k4_xboole_0(X1,k4_xboole_0(X1,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_35]), c_0_31])).
cnf(c_0_40,negated_conjecture,(k2_xboole_0(k4_xboole_0(esk2_0,esk3_0),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))!=esk2_0), inference(rw,[status(thm)],[c_0_36, c_0_17])).
cnf(c_0_41,plain,(k2_xboole_0(X1,k4_xboole_0(k4_xboole_0(X2,X3),X2))=k2_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_31, c_0_37])).
cnf(c_0_42,plain,(~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_38]), c_0_13])).
cnf(c_0_43,plain,(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))=k2_xboole_0(k4_xboole_0(X1,X3),k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_17, c_0_14])).
cnf(c_0_44,plain,(k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X2))=k4_xboole_0(X1,k2_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_14, c_0_39])).
cnf(c_0_45,negated_conjecture,(k2_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_31]), c_0_17])).
cnf(c_0_46,plain,(k2_xboole_0(X1,k4_xboole_0(X1,X2))=k2_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_31, c_0_41])).
cnf(c_0_47,plain,(k4_xboole_0(X1,X1)=k4_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_42, c_0_34])).
cnf(c_0_48,plain,(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,k2_xboole_0(X1,X1))))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_38]), c_0_43]), c_0_23])).
cnf(c_0_49,plain,(k2_xboole_0(k2_xboole_0(X1,X1),k2_xboole_0(X1,X1))=k4_xboole_0(X1,k4_xboole_0(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_39]), c_0_44]), c_0_38])).
cnf(c_0_50,plain,(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,k4_xboole_0(X1,k4_xboole_0(X1,X3)))))=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_14, c_0_23])).
cnf(c_0_51,negated_conjecture,(k2_xboole_0(esk2_0,esk2_0)!=esk2_0), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52,plain,(k2_xboole_0(X1,X1)=k4_xboole_0(X1,k4_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_39, c_0_47])).
cnf(c_0_53,plain,(k4_xboole_0(X1,X2)=X1|r2_hidden(esk1_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_30])).
cnf(c_0_54,plain,(k4_xboole_0(X1,k2_xboole_0(X2,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_48]), c_0_44]), c_0_44]), c_0_44]), c_0_49]), c_0_50]), c_0_43]), c_0_23])).
cnf(c_0_55,negated_conjecture,(k4_xboole_0(esk2_0,k4_xboole_0(X1,X1))!=esk2_0), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56,plain,(k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X1),X3)=k4_xboole_0(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_33, c_0_53])).
cnf(c_0_57,plain,(~r2_hidden(X1,k4_xboole_0(k4_xboole_0(k2_xboole_0(X2,X2),X3),X2))), inference(spm,[status(thm)],[c_0_33, c_0_54])).
cnf(c_0_58,plain,(k4_xboole_0(k4_xboole_0(X1,X1),X2)=k4_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_42, c_0_53])).
cnf(c_0_59,negated_conjecture,(k4_xboole_0(esk2_0,k4_xboole_0(k4_xboole_0(X1,X2),X1))!=esk2_0), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60,plain,(k4_xboole_0(k4_xboole_0(k2_xboole_0(X1,X1),X2),X1)=k4_xboole_0(X3,X3)), inference(spm,[status(thm)],[c_0_57, c_0_34])).
cnf(c_0_61,plain,(k2_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k2_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_23]), c_0_17]), c_0_31]), c_0_17]), c_0_17])).
cnf(c_0_62,plain,(k2_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X2)))=k4_xboole_0(X1,k4_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_35, c_0_58])).
cnf(c_0_63,negated_conjecture,(k4_xboole_0(esk2_0,k4_xboole_0(k4_xboole_0(k4_xboole_0(k2_xboole_0(X1,X1),X2),X1),X3))!=esk2_0), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_64,plain,(k2_xboole_0(X1,k4_xboole_0(X2,X2))=k4_xboole_0(X1,k4_xboole_0(X2,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_47]), c_0_62])).
cnf(c_0_65,negated_conjecture,(k2_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(k4_xboole_0(k2_xboole_0(X1,X1),X2),X1)),k4_xboole_0(esk2_0,X3))!=esk2_0), inference(spm,[status(thm)],[c_0_63, c_0_50])).
cnf(c_0_66,plain,(k2_xboole_0(k4_xboole_0(X1,X1),X2)=k4_xboole_0(X2,k4_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_17, c_0_64])).
cnf(c_0_67,plain,(~r2_hidden(X1,k4_xboole_0(k2_xboole_0(X2,X2),X2))), inference(spm,[status(thm)],[c_0_33, c_0_39])).
cnf(c_0_68,negated_conjecture,(k2_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(X1,X1)),k4_xboole_0(esk2_0,X2))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_58]), c_0_58]), c_0_58])).
cnf(c_0_69,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))))=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_23, c_0_14])).
cnf(c_0_70,plain,(k4_xboole_0(k4_xboole_0(k2_xboole_0(X1,X1),X1),X2)=k4_xboole_0(k2_xboole_0(X1,X1),X1)), inference(spm,[status(thm)],[c_0_67, c_0_53])).
cnf(c_0_71,negated_conjecture,(k2_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(X1,X1)),k2_xboole_0(k4_xboole_0(esk2_0,X2),k4_xboole_0(esk2_0,X3)))!=esk2_0), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_72,plain,(k2_xboole_0(X1,X1)=k4_xboole_0(X1,k4_xboole_0(k2_xboole_0(X2,X2),X2))), inference(spm,[status(thm)],[c_0_52, c_0_70])).
cnf(c_0_73,negated_conjecture,(k2_xboole_0(k2_xboole_0(esk2_0,esk2_0),k2_xboole_0(k4_xboole_0(esk2_0,X1),k4_xboole_0(esk2_0,X2)))!=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_70]), c_0_72])).
cnf(c_0_74,negated_conjecture,(k2_xboole_0(k2_xboole_0(esk2_0,esk2_0),k2_xboole_0(k4_xboole_0(esk2_0,X1),k2_xboole_0(k4_xboole_0(esk2_0,X2),k4_xboole_0(esk2_0,X3))))!=esk2_0), inference(spm,[status(thm)],[c_0_73, c_0_69])).
cnf(c_0_75,plain,(k4_xboole_0(k2_xboole_0(X1,X1),X1)=k4_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_67, c_0_34])).
cnf(c_0_76,negated_conjecture,(k2_xboole_0(k2_xboole_0(esk2_0,esk2_0),k4_xboole_0(k2_xboole_0(k4_xboole_0(esk2_0,X1),k4_xboole_0(esk2_0,X2)),k4_xboole_0(esk2_0,esk2_0)))!=esk2_0), inference(spm,[status(thm)],[c_0_74, c_0_66])).
cnf(c_0_77,plain,(k4_xboole_0(k2_xboole_0(X1,X1),X1)=k4_xboole_0(k4_xboole_0(X2,X3),X2)), inference(spm,[status(thm)],[c_0_56, c_0_75])).
cnf(c_0_78,negated_conjecture,(k2_xboole_0(k2_xboole_0(esk2_0,esk2_0),k4_xboole_0(k4_xboole_0(X1,X2),X1))!=esk2_0), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_79,negated_conjecture,(k2_xboole_0(k2_xboole_0(esk2_0,esk2_0),k4_xboole_0(k4_xboole_0(k2_xboole_0(X1,X1),X2),X1))!=esk2_0), inference(spm,[status(thm)],[c_0_78, c_0_54])).
cnf(c_0_80,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X3)|~r2_hidden(esk1_3(X2,X3,X1),X2)|~r2_hidden(esk1_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_81,negated_conjecture,(k2_xboole_0(k2_xboole_0(esk2_0,esk2_0),k4_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk2_0)),X1))!=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_31]), c_0_49])).
cnf(c_0_82,plain,(k4_xboole_0(X1,X2)=X1|r2_hidden(esk1_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_53]), c_0_53])).
cnf(c_0_83,negated_conjecture,(k2_xboole_0(k2_xboole_0(esk2_0,esk2_0),k4_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(X1,X1)),X2))!=esk2_0), inference(spm,[status(thm)],[c_0_81, c_0_47])).
cnf(c_0_84,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X2))=X1), inference(spm,[status(thm)],[c_0_42, c_0_82])).
cnf(c_0_85,plain,(k2_xboole_0(k2_xboole_0(X1,X1),k4_xboole_0(X1,X2))=k4_xboole_0(X1,k4_xboole_0(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_39]), c_0_58]), c_0_58])).
cnf(c_0_86,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_84]), c_0_85]), c_0_84])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 87
# Proof object clause steps            : 72
# Proof object formula steps           : 15
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 52
# Proof object simplifying inferences  : 43
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 11
# Processed clauses                    : 7487
# ...of these trivial                  : 183
# ...subsumed                          : 6710
# ...remaining for further processing  : 594
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 18
# Backward-rewritten                   : 158
# Generated clauses                    : 380728
# ...of the previous two non-trivial   : 366995
# Contextual simplify-reflections      : 40
# Paramodulations                      : 380689
# Factorizations                       : 36
# Equation resolutions                 : 3
# Current number of processed clauses  : 404
#    Positive orientable unit clauses  : 80
#    Positive unorientable unit clauses: 130
#    Negative unit clauses             : 56
#    Non-unit-clauses                  : 138
# Current number of unprocessed clauses: 301500
# ...number of literals in the above   : 386964
# Current number of archived formulas  : 0
# Current number of archived clauses   : 188
# Clause-clause subsumption calls (NU) : 13227
# Rec. Clause-clause subsumption calls : 11732
# Non-unit clause-clause subsumptions  : 427
# Unit Clause-clause subsumption calls : 6215
# Rewrite failures with RHS unbound    : 536
# BW rewrite match attempts            : 13190
# BW rewrite match successes           : 979
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12305875

# -------------------------------------------------
# User time                : 7.224 s
# System time              : 0.244 s
# Total time               : 7.468 s
# Maximum resident set size: 299076 pages
