# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t52_subset_1, conjecture, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(~(r2_hidden(X4,X2))<=>r2_hidden(X4,X3)))=>X2=k3_subset_1(X1,X3))))), file('out/t52_subset_1__10_premises__rerun.E_input', t52_subset_1)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t52_subset_1__10_premises__rerun.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, (![X1]:~(v1_xboole_0(k1_zfmisc_1(X1)))), file('out/t52_subset_1__10_premises__rerun.E_input', fc1_subset_1)).
fof(d1_zfmisc_1, axiom, (![X1]:![X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1)))), file('out/t52_subset_1__10_premises__rerun.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t52_subset_1__10_premises__rerun.E_input', d3_tarski)).
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t52_subset_1__10_premises__rerun.E_input', d4_xboole_0)).
fof(d5_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2))), file('out/t52_subset_1__10_premises__rerun.E_input', d5_subset_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t52_subset_1__10_premises__rerun.E_input', t100_xboole_1)).
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t52_subset_1__10_premises__rerun.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t52_subset_1__10_premises__rerun.E_input', d1_xboole_0)).
fof(t2_tarski, axiom, (![X1]:![X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2)), file('out/t52_subset_1__10_premises__rerun.E_input', t2_tarski)).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(~(r2_hidden(X4,X2))<=>r2_hidden(X4,X3)))=>X2=k3_subset_1(X1,X3)))))), inference(assume_negation,[status(cth)],[t52_subset_1])).
fof(c_0_12, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
fof(c_0_13, negated_conjecture, (![X8]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((((r2_hidden(X8,esk2_0)|r2_hidden(X8,esk3_0))|~m1_subset_1(X8,esk1_0))&((~r2_hidden(X8,esk3_0)|~r2_hidden(X8,esk2_0))|~m1_subset_1(X8,esk1_0)))&esk2_0!=k3_subset_1(esk1_0,esk3_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])])])).
fof(c_0_14, plain, (![X2]:~v1_xboole_0(k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|r1_tarski(X6,X4))|X5!=k1_zfmisc_1(X4))&((~r1_tarski(X6,X4)|r2_hidden(X6,X5))|X5!=k1_zfmisc_1(X4)))&(((~r2_hidden(esk4_2(X4,X5),X5)|~r1_tarski(esk4_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))&((r2_hidden(esk4_2(X4,X5),X5)|r1_tarski(esk4_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])])).
cnf(c_0_16,plain,(v1_xboole_0(X1)|r2_hidden(X2,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(r1_tarski(X3,X2)|X1!=k1_zfmisc_1(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,negated_conjecture,(r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_22,negated_conjecture,(r1_tarski(esk3_0,X1)|k1_zfmisc_1(esk1_0)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_23, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk6_3(X5,X6,X7),X7)|(~r2_hidden(esk6_3(X5,X6,X7),X5)|~r2_hidden(esk6_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk6_3(X5,X6,X7),X5)|r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk6_3(X5,X6,X7),X6)|r2_hidden(esk6_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
cnf(c_0_24,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25,negated_conjecture,(r1_tarski(esk3_0,esk1_0)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_26,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_27, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(X3))|k3_subset_1(X3,X4)=k4_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_28, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_29,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk6_3(X2,X3,X1),X1)|r2_hidden(esk6_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk7_3(X5,X6,X7),X7)|(~r2_hidden(esk7_3(X5,X6,X7),X5)|r2_hidden(esk7_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk7_3(X5,X6,X7),X5)|r2_hidden(esk7_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk7_3(X5,X6,X7),X6)|r2_hidden(esk7_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
cnf(c_0_31,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk6_3(X2,X3,X1),X1)|r2_hidden(esk6_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,negated_conjecture,(r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_26]), c_0_18])).
cnf(c_0_34,plain,(k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36,plain,(k3_xboole_0(X1,X2)=X1|r2_hidden(esk6_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_29])).
fof(c_0_37, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk8_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_38,plain,(r2_hidden(X4,X2)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39,plain,(X1=k3_xboole_0(X2,X3)|~r2_hidden(esk6_3(X2,X3,X1),X3)|~r2_hidden(esk6_3(X2,X3,X1),X2)|~r2_hidden(esk6_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40,plain,(k3_xboole_0(X1,X2)=X2|r2_hidden(esk6_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_31])).
cnf(c_0_41,negated_conjecture,(k3_xboole_0(X1,X2)=esk3_0|r2_hidden(esk6_3(X1,X2,esk3_0),esk1_0)|r2_hidden(esk6_3(X1,X2,esk3_0),X1)), inference(spm,[status(thm)],[c_0_32, c_0_29])).
cnf(c_0_42,negated_conjecture,(r1_tarski(esk2_0,X1)|k1_zfmisc_1(esk1_0)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_33])).
cnf(c_0_43,plain,(k5_xboole_0(X1,k3_xboole_0(X1,X2))=k3_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44,negated_conjecture,(k3_xboole_0(esk3_0,X1)=esk3_0|r2_hidden(esk6_3(esk3_0,X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_36])).
cnf(c_0_45,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_46,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47,plain,(r2_hidden(X4,X1)|r2_hidden(X4,X3)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_48,plain,(r2_hidden(X4,X2)|X1!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[c_0_38, c_0_35])).
cnf(c_0_49,plain,(k3_xboole_0(X1,X2)=X2|~r2_hidden(esk6_3(X1,X2,X2),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_40])).
cnf(c_0_50,negated_conjecture,(k3_xboole_0(esk1_0,X1)=esk3_0|r2_hidden(esk6_3(esk1_0,X1,esk3_0),esk1_0)), inference(ef,[status(thm)],[c_0_41])).
cnf(c_0_51,negated_conjecture,(r1_tarski(esk2_0,esk1_0)), inference(er,[status(thm)],[c_0_42])).
fof(c_0_52, plain, (![X4]:![X5]:(((~r2_hidden(esk9_2(X4,X5),X4)|~r2_hidden(esk9_2(X4,X5),X5))|X4=X5)&((r2_hidden(esk9_2(X4,X5),X4)|r2_hidden(esk9_2(X4,X5),X5))|X4=X5))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])])])).
cnf(c_0_53,negated_conjecture,(esk2_0!=k3_subset_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,negated_conjecture,(k3_subset_1(esk1_0,esk3_0)=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_43, c_0_17])).
cnf(c_0_55,plain,(r2_hidden(X4,X3)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_56,negated_conjecture,(k3_xboole_0(esk3_0,esk1_0)=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_44]), c_0_36])).
cnf(c_0_57,negated_conjecture,(r2_hidden(X1,esk3_0)|r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_58,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_59,plain,(r2_hidden(X4,X3)|r2_hidden(X4,X1)|X1!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X4,X2)), inference(rw,[status(thm)],[c_0_47, c_0_35])).
cnf(c_0_60,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_61,negated_conjecture,(k3_xboole_0(esk1_0,esk3_0)=esk3_0), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_62,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_51])).
cnf(c_0_63,plain,(X1=X2|r2_hidden(esk9_2(X1,X2),X2)|r2_hidden(esk9_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64,negated_conjecture,(k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk3_0))!=esk2_0), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_65,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_66,negated_conjecture,(r2_hidden(X1,esk1_0)|X2!=esk3_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_67,negated_conjecture,(~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_68,plain,(X1=X2|~r2_hidden(esk9_2(X1,X2),X2)|~r2_hidden(esk9_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_69,negated_conjecture,(r2_hidden(X1,esk3_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_70,plain,(r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_71,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,k5_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_72,negated_conjecture,(X1=esk2_0|r2_hidden(esk9_2(X1,esk2_0),esk1_0)|r2_hidden(esk9_2(X1,esk2_0),X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_73,negated_conjecture,(k5_xboole_0(esk1_0,esk3_0)!=esk2_0), inference(rw,[status(thm)],[c_0_64, c_0_61])).
cnf(c_0_74,plain,(X1!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X4,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[c_0_65, c_0_35])).
cnf(c_0_75,negated_conjecture,(k3_xboole_0(X1,X2)=X2|r2_hidden(esk6_3(X1,X2,X2),esk1_0)|X2!=esk3_0), inference(spm,[status(thm)],[c_0_66, c_0_40])).
cnf(c_0_76,negated_conjecture,(~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_58]), c_0_62])).
cnf(c_0_77,negated_conjecture,(X1=esk2_0|r2_hidden(esk9_2(X1,esk2_0),esk3_0)|~r2_hidden(esk9_2(X1,esk2_0),esk1_0)|~r2_hidden(esk9_2(X1,esk2_0),X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_78,negated_conjecture,(r2_hidden(X1,k5_xboole_0(esk1_0,esk3_0))|r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_70, c_0_61])).
cnf(c_0_79,negated_conjecture,(r2_hidden(esk9_2(k5_xboole_0(esk1_0,esk3_0),esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
cnf(c_0_80,plain,(~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_81,negated_conjecture,(k3_xboole_0(esk1_0,X1)=X1|X1!=esk3_0), inference(spm,[status(thm)],[c_0_49, c_0_75])).
cnf(c_0_82,negated_conjecture,(X1=esk2_0|r2_hidden(esk9_2(X1,esk2_0),X1)|~r2_hidden(esk9_2(X1,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_63])).
cnf(c_0_83,negated_conjecture,(r2_hidden(esk9_2(k5_xboole_0(esk1_0,esk3_0),esk2_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])]), c_0_73])).
cnf(c_0_84,negated_conjecture,(X1!=esk3_0|~r2_hidden(X2,k5_xboole_0(esk1_0,X1))|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85,negated_conjecture,(r2_hidden(esk9_2(k5_xboole_0(esk1_0,esk3_0),esk2_0),k5_xboole_0(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_73])).
cnf(c_0_86,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_83])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 87
# Proof object clause steps            : 64
# Proof object formula steps           : 23
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 34
# Processed clauses                    : 9107
# ...of these trivial                  : 249
# ...subsumed                          : 6265
# ...remaining for further processing  : 2593
# Other redundant clauses eliminated   : 51
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 140
# Backward-rewritten                   : 525
# Generated clauses                    : 87131
# ...of the previous two non-trivial   : 81919
# Contextual simplify-reflections      : 2250
# Paramodulations                      : 86549
# Factorizations                       : 275
# Equation resolutions                 : 237
# Current number of processed clauses  : 1853
#    Positive orientable unit clauses  : 145
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 82
#    Non-unit-clauses                  : 1625
# Current number of unprocessed clauses: 55694
# ...number of literals in the above   : 238059
# Current number of archived formulas  : 0
# Current number of archived clauses   : 726
# Clause-clause subsumption calls (NU) : 586048
# Rec. Clause-clause subsumption calls : 404628
# Non-unit clause-clause subsumptions  : 5136
# Unit Clause-clause subsumption calls : 42121
# Rewrite failures with RHS unbound    : 41
# BW rewrite match attempts            : 625
# BW rewrite match successes           : 88
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1441599

# -------------------------------------------------
# User time                : 1.692 s
# System time              : 0.056 s
# Total time               : 1.748 s
# Maximum resident set size: 54516 pages
