# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', t3_subset)).
fof(t31_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', t31_zfmisc_1)).
fof(t52_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))))))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', t52_waybel_0)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', d9_lattice3)).
fof(reflexivity_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', reflexivity_r3_orders_2)).
fof(t7_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2))))))))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', t7_yellow_0)).
fof(d9_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', d9_yellow_0)).
fof(fc1_finset_1, axiom, (![X1]:v1_finset_1(k1_tarski(X1))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', fc1_finset_1)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', redefinition_r3_orders_2)).
fof(fc2_xboole_0, axiom, (![X1]:~(v1_xboole_0(k1_tarski(X1)))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', fc2_xboole_0)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', fc1_xboole_0)).
fof(t9_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4))))))), file('out/t52_waybel_0__12_premises__rerun__rerun.E_input', t9_yellow_0)).
fof(c_0_12, plain, (![X3]:![X2]:![X1]:(epred1_3(X1,X2,X3)<=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3))))))), introduced(definition)).
fof(c_0_13, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_14, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_zfmisc_1])])])])).
fof(c_0_15, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[t52_waybel_0]), c_0_12])).
fof(c_0_16, plain, (![X3]:![X2]:![X1]:(epred1_3(X1,X2,X3)=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3))))))), inference(split_equiv,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk5_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk5_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk5_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|r3_orders_2(X4,X5,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])])])).
fof(c_0_21, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(epred1_3(esk1_0,esk2_0,esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&((~r2_lattice3(esk1_0,esk2_0,esk4_0)|~r2_lattice3(esk1_0,esk3_0,esk4_0))&(r2_lattice3(esk1_0,esk2_0,esk4_0)|r2_lattice3(esk1_0,esk3_0,esk4_0))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:(((((~r1_lattice3(X4,k1_tarski(X6),X5)|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&(((((~r1_orders_2(X4,X5,X6)|r1_lattice3(X4,k1_tarski(X6),X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&(((((~r2_lattice3(X4,k1_tarski(X6),X5)|r1_orders_2(X4,X6,X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&((((~r1_orders_2(X4,X6,X5)|r2_lattice3(X4,k1_tarski(X6),X5))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])])])).
fof(c_0_23, plain, (![X5]:![X6]:![X7]:![X8]:((((((r2_lattice3(X5,X6,X7)|X7!=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X6,X8)|r1_orders_2(X5,X7,X8)))|X7!=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))&((((((m1_subset_1(esk6_3(X5,X6,X7),u1_struct_0(X5))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((((r2_lattice3(X5,X6,esk6_3(X5,X6,X7))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~r1_orders_2(X5,X7,esk6_3(X5,X6,X7))|~r2_lattice3(X5,X6,X7))|X7=k1_yellow_0(X5,X6))|~r1_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])])])).
fof(c_0_24, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X12]:((((~v1_finset_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(X7)))|(X9=k1_xboole_0|r1_yellow_0(X8,X9)))|~epred1_3(X8,X7,X6))&((((((v1_finset_1(esk7_4(X6,X7,X8,X10))|~r2_hidden(X10,X6))|~m1_subset_1(X10,u1_struct_0(X8)))|~epred1_3(X8,X7,X6))&(((m1_subset_1(esk7_4(X6,X7,X8,X10),k1_zfmisc_1(X7))|~r2_hidden(X10,X6))|~m1_subset_1(X10,u1_struct_0(X8)))|~epred1_3(X8,X7,X6)))&((((r1_yellow_0(X8,esk7_4(X6,X7,X8,X10))|~r2_hidden(X10,X6))|~m1_subset_1(X10,u1_struct_0(X8)))|~epred1_3(X8,X7,X6))&(((X10=k1_yellow_0(X8,esk7_4(X6,X7,X8,X10))|~r2_hidden(X10,X6))|~m1_subset_1(X10,u1_struct_0(X8)))|~epred1_3(X8,X7,X6))))&(((~v1_finset_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(X7)))|(X12=k1_xboole_0|r2_hidden(k1_yellow_0(X8,X12),X6)))|~epred1_3(X8,X7,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])])])])).
cnf(c_0_25,plain,(m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26,plain,(r2_lattice3(X1,X3,X2)|r2_hidden(esk5_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, plain, (![X2]:v1_finset_1(k1_tarski(X2))), inference(variable_rename,[status(thm)],[fc1_finset_1])).
fof(c_0_28, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_29,plain,(r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,k1_tarski(X3),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35,plain,(X2=k1_yellow_0(X1,X3)|r2_lattice3(X1,X3,esk6_3(X1,X3,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36,plain,(X2=k1_yellow_0(X1,X3)|m1_subset_1(esk6_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37,plain,(r1_yellow_0(X1,X4)|X4=k1_xboole_0|~epred1_3(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(X2))|~v1_finset_1(X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38,plain,(r2_lattice3(X1,X2,X3)|m1_subset_1(k1_tarski(esk5_3(X1,X2,X3)),k1_zfmisc_1(X2))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_39,plain,(v1_finset_1(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41,negated_conjecture,(r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])]), c_0_33])).
cnf(c_0_42,plain,(r2_hidden(k1_yellow_0(X1,X4),X3)|X4=k1_xboole_0|~epred1_3(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(X2))|~v1_finset_1(X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_43,plain,(X2=k1_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_44,plain,(X1=k1_yellow_0(X2,k1_tarski(X3))|r1_orders_2(X2,X3,esk6_3(X2,k1_tarski(X3),X1))|~r2_lattice3(X2,k1_tarski(X3),X1)|~r1_yellow_0(X2,k1_tarski(X3))|~l1_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_45,plain,(k1_tarski(esk5_3(X1,X2,X3))=k1_xboole_0|r2_lattice3(X1,X2,X3)|r1_yellow_0(X4,k1_tarski(esk5_3(X1,X2,X3)))|~epred1_3(X4,X2,X5)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
cnf(c_0_46,negated_conjecture,(epred1_3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_47,plain,(r2_lattice3(X1,k1_tarski(X3),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_48,negated_conjecture,(r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_31]), c_0_32])]), c_0_33])).
cnf(c_0_49,plain,(k1_tarski(esk5_3(X1,X2,X3))=k1_xboole_0|r2_lattice3(X1,X2,X3)|r2_hidden(k1_yellow_0(X4,k1_tarski(esk5_3(X1,X2,X3))),X5)|~epred1_3(X4,X2,X5)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_38]), c_0_39])])).
cnf(c_0_50,plain,(k1_yellow_0(X1,k1_tarski(X2))=X2|~r2_lattice3(X1,k1_tarski(X2),X2)|~r1_yellow_0(X1,k1_tarski(X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51,negated_conjecture,(k1_tarski(esk5_3(X1,esk2_0,X2))=k1_xboole_0|r2_lattice3(X1,esk2_0,X2)|r1_yellow_0(esk1_0,k1_tarski(esk5_3(X1,esk2_0,X2)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52,negated_conjecture,(r2_lattice3(esk1_0,k1_tarski(X1),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_31])])).
fof(c_0_53, plain, (![X2]:~v1_xboole_0(k1_tarski(X2))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
cnf(c_0_54,negated_conjecture,(k1_tarski(esk5_3(X1,esk2_0,X2))=k1_xboole_0|r2_lattice3(X1,esk2_0,X2)|r2_hidden(k1_yellow_0(esk1_0,k1_tarski(esk5_3(X1,esk2_0,X2))),esk3_0)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_49, c_0_46])).
cnf(c_0_55,negated_conjecture,(k1_yellow_0(esk1_0,k1_tarski(esk5_3(X1,esk2_0,X2)))=esk5_3(X1,esk2_0,X2)|k1_tarski(esk5_3(X1,esk2_0,X2))=k1_xboole_0|r2_lattice3(X1,esk2_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk5_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_31])]), c_0_52])).
cnf(c_0_56,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_57,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_0)|r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_58,plain,(~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_59,negated_conjecture,(k1_tarski(esk5_3(X1,esk2_0,X2))=k1_xboole_0|r2_lattice3(X1,esk2_0,X2)|r2_hidden(esk5_3(X1,esk2_0,X2),esk3_0)|~l1_orders_2(X1)|~m1_subset_1(esk5_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_61,negated_conjecture,(r1_orders_2(esk1_0,X1,esk4_0)|r2_lattice3(esk1_0,esk2_0,esk4_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_31]), c_0_30])])).
cnf(c_0_62,negated_conjecture,(r2_lattice3(X1,esk2_0,X2)|r2_hidden(esk5_3(X1,esk2_0,X2),esk3_0)|~l1_orders_2(X1)|~m1_subset_1(esk5_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60])])).
cnf(c_0_63,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk5_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_64,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(X1,esk2_0,X2),esk4_0)|r2_lattice3(esk1_0,esk2_0,esk4_0)|r2_lattice3(X1,esk2_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk5_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
fof(c_0_65, plain, (![X5]:![X6]:![X7]:![X8]:(((((~r1_lattice3(X5,X7,X8)|r1_lattice3(X5,X6,X8))|~m1_subset_1(X8,u1_struct_0(X5)))|~r1_tarski(X6,X7))|~l1_orders_2(X5))&((((~r2_lattice3(X5,X7,X8)|r2_lattice3(X5,X6,X8))|~m1_subset_1(X8,u1_struct_0(X5)))|~r1_tarski(X6,X7))|~l1_orders_2(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])])])).
cnf(c_0_66,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_0)|~m1_subset_1(esk5_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_31]), c_0_30])])).
cnf(c_0_67,plain,(r2_lattice3(X1,X3,X2)|m1_subset_1(esk5_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_68,plain,(r2_lattice3(X1,X2,X4)|~l1_orders_2(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_lattice3(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_69,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_31]), c_0_30])])).
cnf(c_0_70,negated_conjecture,(r2_lattice3(esk1_0,X1,esk4_0)|~r1_tarski(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_31]), c_0_30])])).
cnf(c_0_71,plain,(r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_72,negated_conjecture,(r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_73,plain,(m1_subset_1(esk7_4(X3,X2,X1,X4),k1_zfmisc_1(X2))|~epred1_3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_74,plain,(r1_orders_2(X1,X2,X4)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|X2!=k1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_75,plain,(r2_lattice3(esk1_0,esk7_4(X1,esk2_0,X2,X3),esk4_0)|~epred1_3(X2,esk2_0,X1)|~r2_hidden(X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_76,plain,(r1_orders_2(esk1_0,X1,esk4_0)|X1!=k1_yellow_0(esk1_0,esk7_4(X2,esk2_0,X3,X4))|~epred1_3(X3,esk2_0,X2)|~r1_yellow_0(esk1_0,esk7_4(X2,esk2_0,X3,X4))|~r2_hidden(X4,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_31]), c_0_30])])).
cnf(c_0_77,plain,(X4=k1_yellow_0(X1,esk7_4(X3,X2,X1,X4))|~epred1_3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_78,plain,(r1_yellow_0(X1,esk7_4(X3,X2,X1,X4))|~epred1_3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_79,plain,(r1_orders_2(esk1_0,X1,esk4_0)|~epred1_3(esk1_0,esk2_0,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77])]), c_0_78])).
cnf(c_0_80,negated_conjecture,(r1_orders_2(esk1_0,X1,esk4_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_79, c_0_46])).
cnf(c_0_81,negated_conjecture,(~r2_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_82,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(X1,esk3_0,X2),esk4_0)|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk5_3(X1,esk3_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_80, c_0_26])).
cnf(c_0_83,negated_conjecture,(~r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_69])])).
cnf(c_0_84,negated_conjecture,(~m1_subset_1(esk5_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_82]), c_0_31]), c_0_30])]), c_0_83])).
cnf(c_0_85,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_67]), c_0_31]), c_0_30])]), c_0_83]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 86
# Proof object clause steps            : 59
# Proof object formula steps           : 27
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 41
# Processed clauses                    : 549
# ...of these trivial                  : 0
# ...subsumed                          : 82
# ...remaining for further processing  : 467
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 83
# Generated clauses                    : 963
# ...of the previous two non-trivial   : 938
# Contextual simplify-reflections      : 77
# Paramodulations                      : 927
# Factorizations                       : 0
# Equation resolutions                 : 36
# Current number of processed clauses  : 342
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 328
# Current number of unprocessed clauses: 223
# ...number of literals in the above   : 2392
# Current number of archived formulas  : 0
# Current number of archived clauses   : 125
# Clause-clause subsumption calls (NU) : 85631
# Rec. Clause-clause subsumption calls : 3415
# Non-unit clause-clause subsumptions  : 160
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 56705

# -------------------------------------------------
# User time                : 0.132 s
# System time              : 0.008 s
# Total time               : 0.140 s
# Maximum resident set size: 43224 pages
