# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X3)=>r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&((r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r2_lattice3(X1,X2,X3))&((r1_lattice3(X1,X2,X3)=>r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r1_lattice3(X1,X2,X3)))))))), file('out/t52_yellow_0__6_premises__rerun__rerun.E_input', t5_yellow_0)).
fof(d8_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3)))))))), file('out/t52_yellow_0__6_premises__rerun__rerun.E_input', d8_yellow_0)).
fof(t52_yellow_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r2_yellow_0(X1,X2)|r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k2_yellow_0(X1,X2)=k2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))))), file('out/t52_yellow_0__6_premises__rerun__rerun.E_input', t52_yellow_0)).
fof(t50_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2))))))), file('out/t52_yellow_0__6_premises__rerun__rerun.E_input', t50_yellow_0)).
fof(dt_k2_yellow_0, axiom, (![X1]:![X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1)))), file('out/t52_yellow_0__6_premises__rerun__rerun.E_input', dt_k2_yellow_0)).
fof(d10_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))))), file('out/t52_yellow_0__6_premises__rerun__rerun.E_input', d10_yellow_0)).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:((((~r2_lattice3(X4,X5,X6)|r2_lattice3(X4,k3_xboole_0(X5,u1_struct_0(X4)),X6))|~m1_subset_1(X6,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((~r2_lattice3(X4,k3_xboole_0(X5,u1_struct_0(X4)),X6)|r2_lattice3(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&((((~r1_lattice3(X4,X5,X6)|r1_lattice3(X4,k3_xboole_0(X5,u1_struct_0(X4)),X6))|~m1_subset_1(X6,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4)))&(((~r1_lattice3(X4,k3_xboole_0(X5,u1_struct_0(X4)),X6)|r1_lattice3(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|(v2_struct_0(X4)|~l1_orders_2(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_yellow_0])])])])])])])).
fof(c_0_7, plain, (![X6]:![X7]:![X9]:![X10]:![X7]:![X12]:![X15]:((((m1_subset_1(esk1_2(X6,X7),u1_struct_0(X6))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((r1_lattice3(X6,X7,esk1_2(X6,X7))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((~m1_subset_1(X9,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X9)|r1_orders_2(X6,X9,esk1_2(X6,X7))))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((m1_subset_1(esk2_3(X6,X7,X10),u1_struct_0(X6))|~r1_lattice3(X6,X7,X10))|X10=esk1_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((r1_lattice3(X6,X7,esk2_3(X6,X7,X10))|~r1_lattice3(X6,X7,X10))|X10=esk1_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((~r1_orders_2(X6,esk2_3(X6,X7,X10),X10)|~r1_lattice3(X6,X7,X10))|X10=esk1_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r2_yellow_0(X6,X7))|~l1_orders_2(X6)))))))&(((((((m1_subset_1(esk4_3(X6,X7,X12),u1_struct_0(X6))|m1_subset_1(esk3_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r1_lattice3(X6,X7,esk4_3(X6,X7,X12))|m1_subset_1(esk3_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X15)|r1_orders_2(X6,X15,esk4_3(X6,X7,X12))))|m1_subset_1(esk3_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk4_3(X6,X7,X12)!=X12|m1_subset_1(esk3_3(X6,X7,X12),u1_struct_0(X6)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))))&(((((((m1_subset_1(esk4_3(X6,X7,X12),u1_struct_0(X6))|r1_lattice3(X6,X7,esk3_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r1_lattice3(X6,X7,esk4_3(X6,X7,X12))|r1_lattice3(X6,X7,esk3_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X15)|r1_orders_2(X6,X15,esk4_3(X6,X7,X12))))|r1_lattice3(X6,X7,esk3_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk4_3(X6,X7,X12)!=X12|r1_lattice3(X6,X7,esk3_3(X6,X7,X12)))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))))&((((((m1_subset_1(esk4_3(X6,X7,X12),u1_struct_0(X6))|~r1_orders_2(X6,esk3_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r1_lattice3(X6,X7,esk4_3(X6,X7,X12))|~r1_orders_2(X6,esk3_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r1_lattice3(X6,X7,X15)|r1_orders_2(X6,X15,esk4_3(X6,X7,X12))))|~r1_orders_2(X6,esk3_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk4_3(X6,X7,X12)!=X12|~r1_orders_2(X6,esk3_3(X6,X7,X12),X12))|~r1_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r2_yellow_0(X6,X7))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])])).
fof(c_0_8, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r2_yellow_0(X1,X2)|r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k2_yellow_0(X1,X2)=k2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t52_yellow_0])).
cnf(c_0_9,plain,(v2_struct_0(X1)|r1_lattice3(X1,k3_xboole_0(X3,u1_struct_0(X1)),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(m1_subset_1(esk1_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X3]:![X4]:![X4]:![X4]:![X4]:(((~r1_yellow_0(X3,X4)|r1_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3))))|(v2_struct_0(X3)|~l1_orders_2(X3)))&(((~r1_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3)))|r1_yellow_0(X3,X4))|(v2_struct_0(X3)|~l1_orders_2(X3)))&(((~r2_yellow_0(X3,X4)|r2_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3))))|(v2_struct_0(X3)|~l1_orders_2(X3)))&((~r2_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3)))|r2_yellow_0(X3,X4))|(v2_struct_0(X3)|~l1_orders_2(X3))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])])])).
fof(c_0_12, negated_conjecture, (((~v2_struct_0(esk5_0)&l1_orders_2(esk5_0))&((r2_yellow_0(esk5_0,esk6_0)|r2_yellow_0(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0))))&k2_yellow_0(esk5_0,esk6_0)!=k2_yellow_0(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_13, plain, (![X3]:![X4]:(~l1_orders_2(X3)|m1_subset_1(k2_yellow_0(X3,X4),u1_struct_0(X3)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])])])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X8]:((((((r1_lattice3(X5,X6,X7)|X7!=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X6,X8)|r1_orders_2(X5,X8,X7)))|X7!=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))&((((((m1_subset_1(esk7_3(X5,X6,X7),u1_struct_0(X5))|~r1_lattice3(X5,X6,X7))|X7=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((((r1_lattice3(X5,X6,esk7_3(X5,X6,X7))|~r1_lattice3(X5,X6,X7))|X7=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((((~r1_orders_2(X5,esk7_3(X5,X6,X7),X7)|~r1_lattice3(X5,X6,X7))|X7=k2_yellow_0(X5,X6))|~r2_yellow_0(X5,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])])])).
cnf(c_0_15,plain,(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),esk1_2(X1,X3))|v2_struct_0(X1)|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X2,esk1_2(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_16,plain,(r1_lattice3(X1,X2,esk1_2(X1,X2))|~l1_orders_2(X1)|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,plain,(v2_struct_0(X1)|r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(r2_yellow_0(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)))|r2_yellow_0(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(l1_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,negated_conjecture,(~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,plain,(m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|X2!=k2_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,plain,(v2_struct_0(X1)|r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k3_xboole_0(X3,u1_struct_0(X1)),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24,plain,(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),esk1_2(X1,X2))|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_25,negated_conjecture,(r2_yellow_0(esk5_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_26,plain,(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),k2_yellow_0(X1,X3))|v2_struct_0(X1)|~r1_lattice3(X1,X2,k2_yellow_0(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_9, c_0_21])).
cnf(c_0_27,plain,(r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_22]), c_0_21])).
cnf(c_0_28,plain,(X2=k2_yellow_0(X1,X3)|r1_lattice3(X1,X3,esk7_3(X1,X3,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_29,plain,(r1_lattice3(X1,X2,esk1_2(X1,X3))|v2_struct_0(X1)|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),esk1_2(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_23, c_0_10])).
cnf(c_0_30,negated_conjecture,(r1_lattice3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),esk1_2(esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_19])]), c_0_20])).
cnf(c_0_31,plain,(X2=k2_yellow_0(X1,X3)|m1_subset_1(esk7_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_32,plain,(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),k2_yellow_0(X1,X2))|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33,plain,(v2_struct_0(X1)|r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_34,plain,(X2=k2_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,esk7_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35,plain,(r1_orders_2(X1,X3,esk1_2(X1,X2))|~l1_orders_2(X1)|~r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_36,plain,(esk1_2(X1,X2)=k2_yellow_0(X1,X3)|r1_lattice3(X1,X3,esk7_3(X1,X3,esk1_2(X1,X2)))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~r1_lattice3(X1,X3,esk1_2(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_10])).
cnf(c_0_37,negated_conjecture,(r1_lattice3(esk5_0,esk6_0,esk1_2(esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_25]), c_0_19])]), c_0_20])).
cnf(c_0_38,plain,(k2_yellow_0(X1,X2)=k2_yellow_0(X1,X3)|m1_subset_1(esk7_3(X1,X3,k2_yellow_0(X1,X2)),u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,k2_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_31, c_0_21])).
cnf(c_0_39,negated_conjecture,(r1_lattice3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),k2_yellow_0(esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_25]), c_0_19])]), c_0_20])).
cnf(c_0_40,negated_conjecture,(r2_yellow_0(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_25]), c_0_19])]), c_0_20])).
cnf(c_0_41,negated_conjecture,(k2_yellow_0(esk5_0,esk6_0)!=k2_yellow_0(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42,plain,(esk1_2(X1,X2)=k2_yellow_0(X1,X3)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk7_3(X1,X3,esk1_2(X1,X2)))|~r1_lattice3(X1,X3,esk1_2(X1,X2))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_31]), c_0_10])).
cnf(c_0_43,negated_conjecture,(esk1_2(esk5_0,esk6_0)=k2_yellow_0(esk5_0,esk6_0)|r1_lattice3(esk5_0,esk6_0,esk7_3(esk5_0,esk6_0,esk1_2(esk5_0,esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_25]), c_0_19])])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk7_3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),k2_yellow_0(esk5_0,esk6_0)),u1_struct_0(esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_19])]), c_0_41])).
cnf(c_0_45,plain,(k2_yellow_0(X1,X2)=k2_yellow_0(X1,X3)|r1_lattice3(X1,X3,esk7_3(X1,X3,k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,k2_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_21])).
cnf(c_0_46,negated_conjecture,(esk1_2(esk5_0,esk6_0)=k2_yellow_0(esk5_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_25]), c_0_37]), c_0_19])])).
cnf(c_0_47,negated_conjecture,(r1_lattice3(esk5_0,X1,esk7_3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),k2_yellow_0(esk5_0,esk6_0)))|~r1_lattice3(esk5_0,k3_xboole_0(X1,u1_struct_0(esk5_0)),esk7_3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),k2_yellow_0(esk5_0,esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_44]), c_0_19])]), c_0_20])).
cnf(c_0_48,negated_conjecture,(r1_lattice3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),esk7_3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),k2_yellow_0(esk5_0,esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_39]), c_0_40]), c_0_19])]), c_0_41])).
cnf(c_0_49,negated_conjecture,(k2_yellow_0(esk5_0,esk6_0)=k2_yellow_0(esk5_0,X1)|~r2_yellow_0(esk5_0,X1)|~r1_lattice3(esk5_0,esk6_0,esk7_3(esk5_0,X1,k2_yellow_0(esk5_0,esk6_0)))|~r1_lattice3(esk5_0,X1,k2_yellow_0(esk5_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_46]), c_0_25]), c_0_19])])).
cnf(c_0_50,negated_conjecture,(r1_lattice3(esk5_0,esk6_0,esk7_3(esk5_0,k3_xboole_0(esk6_0,u1_struct_0(esk5_0)),k2_yellow_0(esk5_0,esk6_0)))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_51,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_40]), c_0_39])]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 39
# Proof object formula steps           : 13
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 277
# ...of these trivial                  : 3
# ...subsumed                          : 11
# ...remaining for further processing  : 263
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 17
# Generated clauses                    : 1256
# ...of the previous two non-trivial   : 815
# Contextual simplify-reflections      : 8
# Paramodulations                      : 1254
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 210
#    Positive orientable unit clauses  : 120
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 88
# Current number of unprocessed clauses: 549
# ...number of literals in the above   : 1725
# Current number of archived formulas  : 0
# Current number of archived clauses   : 53
# Clause-clause subsumption calls (NU) : 1633
# Rec. Clause-clause subsumption calls : 216
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 28
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2950
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 100093

# -------------------------------------------------
# User time                : 0.044 s
# System time              : 0.000 s
# Total time               : 0.044 s
# Maximum resident set size: 45308 pages
