# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t46_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>((u1_struct_0(X1)=u1_struct_0(X2)&r1_tarski(u1_pre_topc(X2),u1_pre_topc(X1)))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))))=>(r1_funct_2(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X2),u1_struct_0(X3),X4,X5)=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>![X7]:(m1_subset_1(X7,u1_struct_0(X2))=>((X6=X7&r1_tmap_1(X2,X3,X5,X7))=>r1_tmap_1(X1,X3,X4,X6)))))))))))), file('out/t52_tmap_1__5_premises__rerun.E_input', t46_tmap_1)).
fof(t52_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>((u1_struct_0(X1)=u1_struct_0(X2)&r1_tarski(u1_pre_topc(X2),u1_pre_topc(X1)))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))&(v5_pre_topc(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))))=>(v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))&(v5_pre_topc(X4,X1,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))))))))))), file('out/t52_tmap_1__5_premises__rerun.E_input', t52_tmap_1)).
fof(t44_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(v5_pre_topc(X3,X2,X1)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>r1_tmap_1(X2,X1,X3,X4))))))), file('out/t52_tmap_1__5_premises__rerun.E_input', t44_tmap_1)).
fof(reflexivity_r1_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:((~(v1_xboole_0(X2))&(~(v1_xboole_0(X4))&((v1_funct_1(X5)&(v1_funct_2(X5,X1,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X3,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))))=>r1_funct_2(X1,X2,X3,X4,X5,X5))), file('out/t52_tmap_1__5_premises__rerun.E_input', reflexivity_r1_funct_2)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t52_tmap_1__5_premises__rerun.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t52_tmap_1__5_premises__rerun.E_input', dt_l1_pre_topc)).
fof(c_0_6, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))|((v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))|((v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10)))|((u1_struct_0(X8)!=u1_struct_0(X9)|~r1_tarski(u1_pre_topc(X9),u1_pre_topc(X8)))|((~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X8),u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X10))))))|((~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X9),u1_struct_0(X10))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10))))))|(~r1_funct_2(u1_struct_0(X8),u1_struct_0(X10),u1_struct_0(X9),u1_struct_0(X10),X11,X12)|(~m1_subset_1(X13,u1_struct_0(X8))|(~m1_subset_1(X14,u1_struct_0(X9))|((X13!=X14|~r1_tmap_1(X9,X10,X12,X14))|r1_tmap_1(X8,X10,X11,X13)))))))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_tmap_1])])])])])])).
fof(c_0_7, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>((u1_struct_0(X1)=u1_struct_0(X2)&r1_tarski(u1_pre_topc(X2),u1_pre_topc(X1)))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))&(v5_pre_topc(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))))=>(v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))&(v5_pre_topc(X4,X1,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3)))))))))))))), inference(assume_negation,[status(cth)],[t52_tmap_1])).
cnf(c_0_8,plain,(r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X1)|~r1_tmap_1(X5,X2,X6,X7)|X4!=X7|~m1_subset_1(X7,u1_struct_0(X5))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_funct_2(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X5),u1_struct_0(X2),X3,X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~r1_tarski(u1_pre_topc(X5),u1_pre_topc(X1))|u1_struct_0(X1)!=u1_struct_0(X5)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&(v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0)))&((u1_struct_0(esk1_0)=u1_struct_0(esk2_0)&r1_tarski(u1_pre_topc(esk2_0),u1_pre_topc(esk1_0)))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))&(v5_pre_topc(esk4_0,esk2_0,esk3_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0)))))))&(~v1_funct_1(esk4_0)|(~v1_funct_2(esk4_0,u1_struct_0(esk1_0),u1_struct_0(esk3_0))|(~v5_pre_topc(esk4_0,esk1_0,esk3_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk3_0))))))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])])])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X8]:(((((~v5_pre_topc(X7,X6,X5)|(~m1_subset_1(X8,u1_struct_0(X6))|r1_tmap_1(X6,X5,X7,X8)))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X6),u1_struct_0(X5))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X5)))))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((m1_subset_1(esk5_3(X5,X6,X7),u1_struct_0(X6))|v5_pre_topc(X7,X6,X5))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X6),u1_struct_0(X5))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X5)))))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((~r1_tmap_1(X6,X5,X7,esk5_3(X5,X6,X7))|v5_pre_topc(X7,X6,X5))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X6),u1_struct_0(X5))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X5)))))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t44_tmap_1])])])])])])])])).
cnf(c_0_11,plain,(r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(X1)!=u1_struct_0(X5)|~r1_tmap_1(X5,X2,X6,X4)|~r1_funct_2(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X5),u1_struct_0(X2),X3,X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X5))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X6,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X3)|~r1_tarski(u1_pre_topc(X5),u1_pre_topc(X1))|~l1_pre_topc(X5)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X5)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_12,negated_conjecture,(u1_struct_0(esk1_0)=u1_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,X2,X1)|m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,negated_conjecture,(~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk3_0))))|~v5_pre_topc(esk4_0,esk1_0,esk3_0)|~v1_funct_2(esk4_0,u1_struct_0(esk1_0),u1_struct_0(esk3_0))|~v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19,negated_conjecture,(v1_funct_2(esk4_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0))))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21,negated_conjecture,(r1_tmap_1(esk1_0,X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X1)|u1_struct_0(esk2_0)!=u1_struct_0(X4)|~r1_tmap_1(X4,X1,X5,X3)|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),u1_struct_0(X4),u1_struct_0(X1),X2,X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X4))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X5)|~v1_funct_1(X2)|~r1_tarski(u1_pre_topc(X4),u1_pre_topc(esk1_0))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X4)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_22,negated_conjecture,(r1_tarski(u1_pre_topc(esk2_0),u1_pre_topc(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24,negated_conjecture,(v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26,negated_conjecture,(v5_pre_topc(X1,esk1_0,X2)|m1_subset_1(esk5_3(X2,esk1_0,X1),u1_struct_0(esk2_0))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_12]), c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_27,negated_conjecture,(l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28,negated_conjecture,(v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29,negated_conjecture,(~v5_pre_topc(esk4_0,esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_12]), c_0_19]), c_0_12]), c_0_20])])).
cnf(c_0_30,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~r1_tmap_1(X2,X1,X3,esk5_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_32, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((v1_xboole_0(X8)|(v1_xboole_0(X10)|((~v1_funct_1(X11)|(~v1_funct_2(X11,X7,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,X9,X10)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))))))|r1_funct_2(X7,X8,X9,X10,X11,X11))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_funct_2])])])])])).
cnf(c_0_33,negated_conjecture,(r1_tmap_1(esk1_0,X1,X2,X3)|v2_struct_0(X1)|~r1_tmap_1(esk2_0,X1,X4,X3)|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),u1_struct_0(esk2_0),u1_struct_0(X1),X2,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_21]), c_0_22]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_34,negated_conjecture,(m1_subset_1(esk5_3(esk3_0,esk1_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_19]), c_0_20]), c_0_18]), c_0_27]), c_0_28])]), c_0_29]), c_0_30])).
cnf(c_0_35,negated_conjecture,(v5_pre_topc(X1,esk1_0,X2)|v2_struct_0(X2)|~r1_tmap_1(esk1_0,X2,X1,esk5_3(X2,esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_12]), c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_36,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_tmap_1(X2,X1,X3,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v5_pre_topc(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_37,plain,(r1_funct_2(X1,X2,X3,X4,X5,X5)|v1_xboole_0(X4)|v1_xboole_0(X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_2(X6,X3,X4)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X5,X1,X2)|~v1_funct_1(X5)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38,negated_conjecture,(r1_tmap_1(esk1_0,X1,X2,esk5_3(esk3_0,esk1_0,esk4_0))|v2_struct_0(X1)|~r1_tmap_1(esk2_0,X1,X3,esk5_3(esk3_0,esk1_0,esk4_0))|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),u1_struct_0(esk2_0),u1_struct_0(X1),X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39,negated_conjecture,(~r1_tmap_1(esk1_0,esk3_0,esk4_0,esk5_3(esk3_0,esk1_0,esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_19]), c_0_20]), c_0_18]), c_0_27]), c_0_28])]), c_0_29]), c_0_30])).
cnf(c_0_40,negated_conjecture,(r1_tmap_1(esk2_0,X1,X2,esk5_3(esk3_0,esk1_0,esk4_0))|v2_struct_0(X1)|~v5_pre_topc(X2,esk2_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_34]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_41,negated_conjecture,(v5_pre_topc(esk4_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_42, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_43, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_44,negated_conjecture,(v1_xboole_0(u1_struct_0(esk3_0))|v1_xboole_0(X1)|r1_funct_2(X2,X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0),X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_20]), c_0_19]), c_0_18])])).
cnf(c_0_45,negated_conjecture,(~r1_tmap_1(esk2_0,esk3_0,X1,esk5_3(esk3_0,esk1_0,esk4_0))|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk3_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_19]), c_0_20]), c_0_18]), c_0_27]), c_0_28])]), c_0_39]), c_0_30])).
cnf(c_0_46,negated_conjecture,(r1_tmap_1(esk2_0,esk3_0,esk4_0,esk5_3(esk3_0,esk1_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_20]), c_0_19]), c_0_18]), c_0_27]), c_0_28])]), c_0_30])).
cnf(c_0_47,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49,negated_conjecture,(v1_xboole_0(u1_struct_0(esk3_0))|r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_20]), c_0_19]), c_0_18])])).
cnf(c_0_50,negated_conjecture,(~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk3_0),esk4_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_20]), c_0_19]), c_0_18])])).
cnf(c_0_51,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52,negated_conjecture,(v1_xboole_0(u1_struct_0(esk3_0))), inference(sr,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_53,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_27])]), c_0_30]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 41
# Proof object formula steps           : 13
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 70
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 75
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 35
# ...of the previous two non-trivial   : 32
# Contextual simplify-reflections      : 0
# Paramodulations                      : 30
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 47
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 41
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 1461
# Rec. Clause-clause subsumption calls : 6
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5349

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 45352 pages
