# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t53_yellow_0__7_premises__rerun.E_input', d4_xboole_0)).
fof(t53_yellow_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1)))), file('out/t53_yellow_0__7_premises__rerun.E_input', t53_yellow_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t53_yellow_0__7_premises__rerun.E_input', d3_tarski)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t53_yellow_0__7_premises__rerun.E_input', t3_subset)).
fof(d7_yellow_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3)))))))), file('out/t53_yellow_0__7_premises__rerun.E_input', d7_yellow_0)).
fof(d12_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>(v3_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('out/t53_yellow_0__7_premises__rerun.E_input', d12_lattice3)).
fof(t50_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2))))))), file('out/t53_yellow_0__7_premises__rerun.E_input', t50_yellow_0)).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|~r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
fof(c_0_8, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1))))), inference(assume_negation,[status(cth)],[t53_yellow_0])).
cnf(c_0_9,plain,(r2_hidden(X4,X3)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk3_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk3_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_11, negated_conjecture, (![X4]:((~v2_struct_0(esk2_0)&l1_orders_2(esk2_0))&((~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk2_0)))|r1_yellow_0(esk2_0,X4))&~v3_lattice3(esk2_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])])).
fof(c_0_12, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_13,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(r1_tarski(X1,X2)|r2_hidden(esk3_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, (![X6]:![X7]:![X9]:![X10]:![X7]:![X12]:![X15]:((((m1_subset_1(esk4_2(X6,X7),u1_struct_0(X6))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((r2_lattice3(X6,X7,esk4_2(X6,X7))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((~m1_subset_1(X9,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X9)|r1_orders_2(X6,esk4_2(X6,X7),X9)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((m1_subset_1(esk5_3(X6,X7,X10),u1_struct_0(X6))|~r2_lattice3(X6,X7,X10))|X10=esk4_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((r2_lattice3(X6,X7,esk5_3(X6,X7,X10))|~r2_lattice3(X6,X7,X10))|X10=esk4_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((~r1_orders_2(X6,X10,esk5_3(X6,X7,X10))|~r2_lattice3(X6,X7,X10))|X10=esk4_2(X6,X7))|~m1_subset_1(X10,u1_struct_0(X6)))|~r1_yellow_0(X6,X7))|~l1_orders_2(X6)))))))&(((((((m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6))|m1_subset_1(esk6_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r2_lattice3(X6,X7,esk7_3(X6,X7,X12))|m1_subset_1(esk6_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X15)|r1_orders_2(X6,esk7_3(X6,X7,X12),X15)))|m1_subset_1(esk6_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk7_3(X6,X7,X12)!=X12|m1_subset_1(esk6_3(X6,X7,X12),u1_struct_0(X6)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))))&(((((((m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6))|r2_lattice3(X6,X7,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r2_lattice3(X6,X7,esk7_3(X6,X7,X12))|r2_lattice3(X6,X7,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X15)|r1_orders_2(X6,esk7_3(X6,X7,X12),X15)))|r2_lattice3(X6,X7,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk7_3(X6,X7,X12)!=X12|r2_lattice3(X6,X7,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))))&((((((m1_subset_1(esk7_3(X6,X7,X12),u1_struct_0(X6))|~r1_orders_2(X6,X12,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&(((((((r2_lattice3(X6,X7,esk7_3(X6,X7,X12))|~r1_orders_2(X6,X12,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6))&((((((~m1_subset_1(X15,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X15)|r1_orders_2(X6,esk7_3(X6,X7,X12),X15)))|~r1_orders_2(X6,X12,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))&(((((esk7_3(X6,X7,X12)!=X12|~r1_orders_2(X6,X12,esk6_3(X6,X7,X12)))|~r2_lattice3(X6,X7,X12))|~m1_subset_1(X12,u1_struct_0(X6)))|r1_yellow_0(X6,X7))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])])).
fof(c_0_16, plain, (![X5]:![X6]:![X8]:![X10]:((((m1_subset_1(esk8_2(X5,X6),u1_struct_0(X5))|~v3_lattice3(X5))|~l1_orders_2(X5))&(((r2_lattice3(X5,X6,esk8_2(X5,X6))|~v3_lattice3(X5))|~l1_orders_2(X5))&(((~m1_subset_1(X8,u1_struct_0(X5))|(~r2_lattice3(X5,X6,X8)|r1_orders_2(X5,esk8_2(X5,X6),X8)))|~v3_lattice3(X5))|~l1_orders_2(X5))))&(((((m1_subset_1(esk10_2(X5,X10),u1_struct_0(X5))|~r2_lattice3(X5,esk9_1(X5),X10))|~m1_subset_1(X10,u1_struct_0(X5)))|v3_lattice3(X5))|~l1_orders_2(X5))&(((((r2_lattice3(X5,esk9_1(X5),esk10_2(X5,X10))|~r2_lattice3(X5,esk9_1(X5),X10))|~m1_subset_1(X10,u1_struct_0(X5)))|v3_lattice3(X5))|~l1_orders_2(X5))&((((~r1_orders_2(X5,X10,esk10_2(X5,X10))|~r2_lattice3(X5,esk9_1(X5),X10))|~m1_subset_1(X10,u1_struct_0(X5)))|v3_lattice3(X5))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_lattice3])])])])])])])).
cnf(c_0_17,negated_conjecture,(r1_yellow_0(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20,plain,(r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk3_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_21,plain,(r1_orders_2(X1,esk4_2(X1,X2),X3)|~l1_orders_2(X1)|~r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(v3_lattice3(X1)|m1_subset_1(esk10_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,esk9_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, (![X3]:![X4]:![X4]:![X4]:![X4]:(((~r1_yellow_0(X3,X4)|r1_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3))))|(v2_struct_0(X3)|~l1_orders_2(X3)))&(((~r1_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3)))|r1_yellow_0(X3,X4))|(v2_struct_0(X3)|~l1_orders_2(X3)))&(((~r2_yellow_0(X3,X4)|r2_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3))))|(v2_struct_0(X3)|~l1_orders_2(X3)))&((~r2_yellow_0(X3,k3_xboole_0(X4,u1_struct_0(X3)))|r2_yellow_0(X3,X4))|(v2_struct_0(X3)|~l1_orders_2(X3))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])])])).
cnf(c_0_24,negated_conjecture,(r1_yellow_0(esk2_0,X1)|~r1_tarski(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25,plain,(r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26,plain,(r1_orders_2(X1,esk4_2(X1,X2),esk10_2(X1,X3))|v3_lattice3(X1)|~r2_lattice3(X1,X2,esk10_2(X1,X3))|~r2_lattice3(X1,esk9_1(X1),X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27,plain,(m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~r1_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28,plain,(v3_lattice3(X1)|r2_lattice3(X1,esk9_1(X1),esk10_2(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,esk9_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,negated_conjecture,(r1_yellow_0(esk2_0,k3_xboole_0(X1,u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,negated_conjecture,(l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_32,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_33,plain,(v3_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,esk9_1(X1),X2)|~r1_orders_2(X1,X2,esk10_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_34,plain,(r1_orders_2(X1,esk4_2(X1,X2),esk10_2(X1,esk4_2(X1,X3)))|v3_lattice3(X1)|~r2_lattice3(X1,X2,esk10_2(X1,esk4_2(X1,X3)))|~r2_lattice3(X1,esk9_1(X1),esk4_2(X1,X3))|~r1_yellow_0(X1,X2)|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35,plain,(r2_lattice3(X1,esk9_1(X1),esk10_2(X1,esk4_2(X1,X2)))|v3_lattice3(X1)|~r2_lattice3(X1,esk9_1(X1),esk4_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_27])).
cnf(c_0_36,negated_conjecture,(r1_yellow_0(esk2_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_37,negated_conjecture,(~v3_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38,plain,(r2_lattice3(X1,X2,esk4_2(X1,X2))|~l1_orders_2(X1)|~r1_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_39,plain,(v3_lattice3(X1)|~r2_lattice3(X1,X2,esk10_2(X1,esk4_2(X1,X2)))|~r2_lattice3(X1,esk9_1(X1),esk4_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_27])).
cnf(c_0_40,negated_conjecture,(r2_lattice3(esk2_0,esk9_1(esk2_0),esk10_2(esk2_0,esk4_2(esk2_0,X1)))|~r2_lattice3(esk2_0,esk9_1(esk2_0),esk4_2(esk2_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_31])]), c_0_37])).
cnf(c_0_41,negated_conjecture,(r2_lattice3(esk2_0,X1,esk4_2(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_36]), c_0_31])])).
cnf(c_0_42,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_36]), c_0_31])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 28
# Proof object formula steps           : 15
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 43
# Processed clauses                    : 14526
# ...of these trivial                  : 3459
# ...subsumed                          : 8721
# ...remaining for further processing  : 2346
# Other redundant clauses eliminated   : 32
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 53
# Generated clauses                    : 628835
# ...of the previous two non-trivial   : 555928
# Contextual simplify-reflections      : 495
# Paramodulations                      : 627936
# Factorizations                       : 772
# Equation resolutions                 : 127
# Current number of processed clauses  : 2250
#    Positive orientable unit clauses  : 587
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 1659
# Current number of unprocessed clauses: 536964
# ...number of literals in the above   : 1545842
# Current number of archived formulas  : 0
# Current number of archived clauses   : 96
# Clause-clause subsumption calls (NU) : 1764737
# Rec. Clause-clause subsumption calls : 1074285
# Non-unit clause-clause subsumptions  : 9158
# Unit Clause-clause subsumption calls : 124573
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26310
# BW rewrite match successes           : 661
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14129497

# -------------------------------------------------
# User time                : 8.352 s
# System time              : 0.284 s
# Total time               : 8.636 s
# Maximum resident set size: 376872 pages
