# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t52_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))))))), file('out/t53_waybel_0__2_premises__rerun.E_input', t52_waybel_0)).
fof(t53_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))=>(r1_yellow_0(X1,X2)<=>r1_yellow_0(X1,X3))))))), file('out/t53_waybel_0__2_premises__rerun.E_input', t53_waybel_0)).
fof(t46_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:![X3]:((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))&r1_yellow_0(X1,X2))=>r1_yellow_0(X1,X3)))), file('out/t53_waybel_0__2_premises__rerun.E_input', t46_yellow_0)).
fof(c_0_3, plain, (![X3]:![X2]:![X1]:(epred1_3(X1,X2,X3)<=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3))))))), introduced(definition)).
fof(c_0_4, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))))))), inference(apply_def,[status(thm)],[t52_waybel_0, c_0_3])).
fof(c_0_5, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(epred1_3(X1,X2,X3)=>(r1_yellow_0(X1,X2)<=>r1_yellow_0(X1,X3)))))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[t53_waybel_0]), c_0_3])).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:(((((((~r2_lattice3(X5,X6,X8)|r2_lattice3(X5,X7,X8))|~m1_subset_1(X8,u1_struct_0(X5)))|~epred1_3(X5,X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|~l1_orders_2(X5)))))&((((((~r2_lattice3(X5,X7,X8)|r2_lattice3(X5,X6,X8))|~m1_subset_1(X8,u1_struct_0(X5)))|~epred1_3(X5,X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])])])])).
fof(c_0_7, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(epred1_3(esk1_0,esk2_0,esk3_0)&((~r1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,esk3_0))&(r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,esk3_0)))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])).
fof(c_0_8, plain, (![X5]:![X6]:![X7]:((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|~r1_yellow_0(X5,X6))|r1_yellow_0(X5,X7))|(v2_struct_0(X5)|~l1_orders_2(X5)))&(((((~r2_lattice3(X5,X6,esk4_3(X5,X6,X7))|~r2_lattice3(X5,X7,esk4_3(X5,X6,X7)))|~r1_yellow_0(X5,X6))|r1_yellow_0(X5,X7))|(v2_struct_0(X5)|~l1_orders_2(X5)))&((((r2_lattice3(X5,X6,esk4_3(X5,X6,X7))|r2_lattice3(X5,X7,esk4_3(X5,X6,X7)))|~r1_yellow_0(X5,X6))|r1_yellow_0(X5,X7))|(v2_struct_0(X5)|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_yellow_0])])])])])])])])).
cnf(c_0_9,plain,(v2_struct_0(X1)|r2_lattice3(X1,X2,X4)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~epred1_3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_lattice3(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(epred1_3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk4_3(X1,X3,X2))|r2_lattice3(X1,X3,esk4_3(X1,X3,X2))|~l1_orders_2(X1)|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18,negated_conjecture,(r1_yellow_0(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,X1)|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_20,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,esk3_0,X1))|r2_lattice3(esk1_0,X1,esk4_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_13])]), c_0_16])).
cnf(c_0_21,plain,(v2_struct_0(X1)|r2_lattice3(X1,X3,X4)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~epred1_3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_lattice3(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk3_0,X1))|r2_lattice3(esk1_0,X1,esk4_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk4_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_24,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X2,esk4_3(X1,X3,X2))|~r2_lattice3(X1,X3,esk4_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,X1)|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_26,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk3_0,X1))|r2_lattice3(esk1_0,X1,esk4_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_13])]), c_0_16]), c_0_18])).
cnf(c_0_27,negated_conjecture,(r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk3_0,X1))|~r2_lattice3(esk1_0,X1,esk4_3(esk1_0,esk3_0,X1))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk4_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_13])]), c_0_16])).
cnf(c_0_28,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk3_0,esk2_0))|r1_yellow_0(esk1_0,esk2_0)), inference(ef,[status(thm)],[c_0_26])).
cnf(c_0_29,negated_conjecture,(r1_yellow_0(esk1_0,esk2_0)|~m1_subset_1(esk4_3(esk1_0,esk3_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_18]), c_0_28])).
cnf(c_0_30,negated_conjecture,(r1_yellow_0(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_23]), c_0_13])]), c_0_16]), c_0_18])).
cnf(c_0_31,negated_conjecture,(~r1_yellow_0(esk1_0,esk3_0)|~r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,X1))|r2_lattice3(esk1_0,X1,esk4_3(esk1_0,esk2_0,X1))|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_30]), c_0_13])]), c_0_16])).
cnf(c_0_33,negated_conjecture,(~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_30])])).
cnf(c_0_34,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_32]), c_0_33])).
cnf(c_0_35,negated_conjecture,(~r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_34]), c_0_30]), c_0_13])]), c_0_33]), c_0_16])).
cnf(c_0_36,negated_conjecture,(~m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_25]), c_0_34])).
cnf(c_0_37,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_23]), c_0_30]), c_0_13])]), c_0_33]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 29
# Proof object formula steps           : 9
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 68
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 68
# Processed clauses                    : 136
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 133
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 6
# Generated clauses                    : 40
# ...of the previous two non-trivial   : 32
# Contextual simplify-reflections      : 7
# Paramodulations                      : 34
# Factorizations                       : 6
# Equation resolutions                 : 0
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 176
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 877
# Rec. Clause-clause subsumption calls : 63
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 33
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6749

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 43224 pages
