# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1))), file('out/t53_pre_topc__9_premises__rerun.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t53_pre_topc__9_premises__rerun.E_input', dt_l1_pre_topc)).
fof(t53_pre_topc, conjecture, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)=>k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))&((v2_pre_topc(X1)&k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=>v3_pre_topc(X2,X1)))))), file('out/t53_pre_topc__9_premises__rerun.E_input', t53_pre_topc)).
fof(d6_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)<=>v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1))))), file('out/t53_pre_topc__9_premises__rerun.E_input', d6_pre_topc)).
fof(t22_pre_topc, axiom, (![X1]:(l1_struct_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2))), file('out/t53_pre_topc__9_premises__rerun.E_input', t22_pre_topc)).
fof(dt_k2_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t53_pre_topc__9_premises__rerun.E_input', dt_k2_struct_0)).
fof(redefinition_k7_subset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3))), file('out/t53_pre_topc__9_premises__rerun.E_input', redefinition_k7_subset_1)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t53_pre_topc__9_premises__rerun.E_input', t3_subset)).
fof(t36_xboole_1, axiom, (![X1]:![X2]:r1_tarski(k4_xboole_0(X1,X2),X1)), file('out/t53_pre_topc__9_premises__rerun.E_input', t36_xboole_1)).
fof(t52_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1)))))), file('out/t53_pre_topc__9_premises__rerun.E_input', t52_pre_topc)).
fof(c_0_10, plain, (![X2]:(~l1_struct_0(X2)|k2_struct_0(X2)=u1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_11, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_12, negated_conjecture, (~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)=>k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))&((v2_pre_topc(X1)&k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=>v3_pre_topc(X2,X1))))))), inference(assume_negation,[status(cth)],[t53_pre_topc])).
cnf(c_0_13,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, negated_conjecture, ((l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((((v2_pre_topc(esk1_0)|v3_pre_topc(esk2_0,esk1_0))&(k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0)|v3_pre_topc(esk2_0,esk1_0)))&(~v3_pre_topc(esk2_0,esk1_0)|v3_pre_topc(esk2_0,esk1_0)))&(((v2_pre_topc(esk1_0)|k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))!=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))&(k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0)|k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))!=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0)))&(~v3_pre_topc(esk2_0,esk1_0)|k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))!=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])])).
fof(c_0_16, plain, (![X3]:![X4]:((((~v4_pre_topc(X4,X3)|v3_pre_topc(k7_subset_1(u1_struct_0(X3),k2_struct_0(X3),X4),X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3))&(((~v3_pre_topc(k7_subset_1(u1_struct_0(X3),k2_struct_0(X3),X4),X3)|v4_pre_topc(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_pre_topc])])])])])])).
fof(c_0_17, plain, (![X3]:![X4]:(~l1_struct_0(X3)|(~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|k7_subset_1(u1_struct_0(X3),k2_struct_0(X3),k7_subset_1(u1_struct_0(X3),k2_struct_0(X3),X4))=X4))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_pre_topc])])])])])).
fof(c_0_18, plain, (![X2]:(~l1_struct_0(X2)|m1_subset_1(k2_struct_0(X2),k1_zfmisc_1(u1_struct_0(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
cnf(c_0_19,plain,(k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X5,k1_zfmisc_1(X4))|k7_subset_1(X4,X5,X6)=k4_xboole_0(X5,X6))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])])])).
cnf(c_0_24,plain,(m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,negated_conjecture,(k2_struct_0(esk1_0)=u1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26,plain,(v3_pre_topc(X1,X2)|~v4_pre_topc(k7_subset_1(u1_struct_0(X2),k2_struct_0(X2),X1),X2)|~m1_subset_1(k7_subset_1(u1_struct_0(X2),k2_struct_0(X2),X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14])).
fof(c_0_27, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_28, plain, (![X3]:![X4]:r1_tarski(k4_xboole_0(X3,X4),X3)), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_29,plain,(k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_31, plain, (![X3]:![X4]:((((~v4_pre_topc(X4,X3)|k2_pre_topc(X3,X4)=X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3))&((((~v2_pre_topc(X3)|k2_pre_topc(X3,X4)!=X4)|v4_pre_topc(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])])])).
cnf(c_0_32,negated_conjecture,(v3_pre_topc(X1,esk1_0)|~v4_pre_topc(k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),esk1_0)|~m1_subset_1(k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_25]), c_0_20])])).
cnf(c_0_33,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35,negated_conjecture,(k4_xboole_0(u1_struct_0(esk1_0),X1)=k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1)|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36,plain,(v4_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|k2_pre_topc(X1,X2)!=X2|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37,plain,(v4_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38,negated_conjecture,(v3_pre_topc(X1,esk1_0)|~r1_tarski(k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),u1_struct_0(esk1_0))|~v4_pre_topc(k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39,negated_conjecture,(r1_tarski(k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40,plain,(v4_pre_topc(X1,X2)|k2_pre_topc(X2,X1)!=X1|~r1_tarski(X1,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_36, c_0_33])).
cnf(c_0_41,plain,(v4_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_22]), c_0_14])).
cnf(c_0_42,plain,(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2)=k4_xboole_0(k2_struct_0(X1),X2)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_29, c_0_24])).
cnf(c_0_43,negated_conjecture,(v3_pre_topc(X1,esk1_0)|~v4_pre_topc(k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44,negated_conjecture,(v4_pre_topc(k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),esk1_0)|k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1))!=k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1)|~v2_pre_topc(esk1_0)|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_39]), c_0_20])])).
cnf(c_0_45,negated_conjecture,(v3_pre_topc(esk2_0,esk1_0)|k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_46,plain,(k2_pre_topc(X1,X2)=X2|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_47,plain,(v4_pre_topc(k4_xboole_0(k2_struct_0(X1),X2),X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(k4_xboole_0(k2_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_14])).
cnf(c_0_48,negated_conjecture,(v3_pre_topc(X1,esk1_0)|k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1))!=k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1)|~v2_pre_topc(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49,negated_conjecture,(k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk2_0))=k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk2_0)|v3_pre_topc(esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_25]), c_0_25])).
cnf(c_0_50,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_51,negated_conjecture,(v3_pre_topc(esk2_0,esk1_0)|v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_52,plain,(k2_pre_topc(X1,X2)=X2|~r1_tarski(X2,u1_struct_0(X1))|~v4_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_46, c_0_33])).
cnf(c_0_53,plain,(v4_pre_topc(k4_xboole_0(k2_struct_0(X1),X2),X1)|~r1_tarski(k4_xboole_0(k2_struct_0(X1),X2),u1_struct_0(X1))|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_47, c_0_33])).
cnf(c_0_54,negated_conjecture,(k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))!=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0)|~v3_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_55,negated_conjecture,(v3_pre_topc(esk2_0,esk1_0)|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])]), c_0_51])).
cnf(c_0_56,plain,(k2_pre_topc(X1,k4_xboole_0(u1_struct_0(X1),X2))=k4_xboole_0(u1_struct_0(X1),X2)|~v4_pre_topc(k4_xboole_0(u1_struct_0(X1),X2),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_52, c_0_34])).
cnf(c_0_57,negated_conjecture,(v4_pre_topc(k4_xboole_0(u1_struct_0(esk1_0),X1),esk1_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_25]), c_0_34]), c_0_20])])).
cnf(c_0_58,negated_conjecture,(k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk2_0))!=k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk2_0)|~v3_pre_topc(esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_25]), c_0_25])).
cnf(c_0_59,negated_conjecture,(v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_14]), c_0_20])])).
cnf(c_0_60,negated_conjecture,(k2_pre_topc(esk1_0,k4_xboole_0(u1_struct_0(esk1_0),X1))=k4_xboole_0(u1_struct_0(esk1_0),X1)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_20])])).
cnf(c_0_61,negated_conjecture,(k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk2_0))!=k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_62,negated_conjecture,(k2_pre_topc(esk1_0,k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1))=k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_35])).
cnf(c_0_63,negated_conjecture,(~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_59]), c_0_50])])).
cnf(c_0_64,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_14]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 44
# Proof object formula steps           : 21
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 18
# Processed clauses                    : 235
# ...of these trivial                  : 0
# ...subsumed                          : 111
# ...remaining for further processing  : 124
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 8
# Generated clauses                    : 346
# ...of the previous two non-trivial   : 317
# Contextual simplify-reflections      : 70
# Paramodulations                      : 346
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 96
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 82
# Current number of unprocessed clauses: 117
# ...number of literals in the above   : 484
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 4033
# Rec. Clause-clause subsumption calls : 2414
# Non-unit clause-clause subsumptions  : 183
# Unit Clause-clause subsumption calls : 82
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8445

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.004 s
# Total time               : 0.024 s
# Maximum resident set size: 45312 pages
