# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t51_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(k9_xtuple_0(X2),X1)=>k3_relat_1(k4_relat_1(X1),X2)=X2))), file('out/t53_relat_1__7_premises__rerun.E_input', t51_relat_1)).
fof(t20_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>(k10_xtuple_0(X1)=k9_xtuple_0(k2_relat_1(X1))&k9_xtuple_0(X1)=k10_xtuple_0(k2_relat_1(X1))))), file('out/t53_relat_1__7_premises__rerun.E_input', t20_relat_1)).
fof(dt_k2_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>v1_relat_1(k2_relat_1(X1)))), file('out/t53_relat_1__7_premises__rerun.E_input', dt_k2_relat_1)).
fof(t53_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(k10_xtuple_0(X2),X1)=>k3_relat_1(X2,k4_relat_1(X1))=X2))), file('out/t53_relat_1__7_premises__rerun.E_input', t53_relat_1)).
fof(t35_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k2_relat_1(k3_relat_1(X1,X2))=k3_relat_1(k2_relat_1(X2),k2_relat_1(X1))))), file('out/t53_relat_1__7_premises__rerun.E_input', t35_relat_1)).
fof(involutiveness_k2_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>k2_relat_1(k2_relat_1(X1))=X1)), file('out/t53_relat_1__7_premises__rerun.E_input', involutiveness_k2_relat_1)).
fof(rd3_relat_1, axiom, (![X1]:k2_relat_1(k4_relat_1(X1))=k4_relat_1(X1)), file('out/t53_relat_1__7_premises__rerun.E_input', rd3_relat_1)).
fof(dt_k4_relat_1, axiom, (![X1]:v1_relat_1(k4_relat_1(X1))), file('out/t53_relat_1__7_premises__rerun.E_input', dt_k4_relat_1)).
fof(c_0_8, plain, (![X3]:![X4]:(~v1_relat_1(X4)|(~r1_tarski(k9_xtuple_0(X4),X3)|k3_relat_1(k4_relat_1(X3),X4)=X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t51_relat_1])])).
fof(c_0_9, plain, (![X2]:((k10_xtuple_0(X2)=k9_xtuple_0(k2_relat_1(X2))|~v1_relat_1(X2))&(k9_xtuple_0(X2)=k10_xtuple_0(k2_relat_1(X2))|~v1_relat_1(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
fof(c_0_10, plain, (![X2]:(~v1_relat_1(X2)|v1_relat_1(k2_relat_1(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_relat_1])])).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(k10_xtuple_0(X2),X1)=>k3_relat_1(X2,k4_relat_1(X1))=X2)))), inference(assume_negation,[status(cth)],[t53_relat_1])).
fof(c_0_12, plain, (![X3]:![X4]:(~v1_relat_1(X3)|(~v1_relat_1(X4)|k2_relat_1(k3_relat_1(X3,X4))=k3_relat_1(k2_relat_1(X4),k2_relat_1(X3))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_relat_1])])])])])).
fof(c_0_13, plain, (![X2]:(~v1_relat_1(X2)|k2_relat_1(k2_relat_1(X2))=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k2_relat_1])])).
cnf(c_0_14,plain,(k3_relat_1(k4_relat_1(X1),X2)=X2|~r1_tarski(k9_xtuple_0(X2),X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,plain,(k10_xtuple_0(X1)=k9_xtuple_0(k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,plain,(v1_relat_1(k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, negated_conjecture, ((v1_relat_1(esk2_0)&(r1_tarski(k10_xtuple_0(esk2_0),esk1_0)&k3_relat_1(esk2_0,k4_relat_1(esk1_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_18,plain,(k2_relat_1(k3_relat_1(X1,X2))=k3_relat_1(k2_relat_1(X2),k2_relat_1(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(k2_relat_1(k2_relat_1(X1))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k3_relat_1(k4_relat_1(X1),k2_relat_1(X2))=k2_relat_1(X2)|~r1_tarski(k10_xtuple_0(X2),X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_21,negated_conjecture,(r1_tarski(k10_xtuple_0(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, (![X2]:k2_relat_1(k4_relat_1(X2))=k4_relat_1(X2)), inference(variable_rename,[status(thm)],[rd3_relat_1])).
fof(c_0_24, plain, (![X2]:v1_relat_1(k4_relat_1(X2))), inference(variable_rename,[status(thm)],[dt_k4_relat_1])).
cnf(c_0_25,plain,(k2_relat_1(k3_relat_1(X1,k2_relat_1(X2)))=k3_relat_1(X2,k2_relat_1(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_16])).
cnf(c_0_26,negated_conjecture,(k3_relat_1(k4_relat_1(esk1_0),k2_relat_1(esk2_0))=k2_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
cnf(c_0_27,plain,(k2_relat_1(k4_relat_1(X1))=k4_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,plain,(v1_relat_1(k4_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29,negated_conjecture,(k3_relat_1(esk2_0,k4_relat_1(esk1_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30,negated_conjecture,(k3_relat_1(esk2_0,k4_relat_1(esk1_0))=k2_relat_1(k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28]), c_0_22])])).
cnf(c_0_31,negated_conjecture,(k2_relat_1(k2_relat_1(esk2_0))!=esk2_0), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_32,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_19]), c_0_22])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 16
# Proof object formula steps           : 17
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 11
# Processed clauses                    : 52
# ...of these trivial                  : 2
# ...subsumed                          : 3
# ...remaining for further processing  : 47
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 12
# Generated clauses                    : 104
# ...of the previous two non-trivial   : 105
# Contextual simplify-reflections      : 5
# Paramodulations                      : 104
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 15
# Current number of unprocessed clauses: 57
# ...number of literals in the above   : 204
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 116
# Rec. Clause-clause subsumption calls : 114
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2612

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 42948 pages
