# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t56_mcart_1, conjecture, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:(k4_zfmisc_1(X1,X2,X3,X4)=k4_zfmisc_1(X5,X6,X7,X8)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|(X1=X5&(X2=X6&(X3=X7&X4=X8))))))))), file('out/t56_mcart_1__12_premises__rerun__rerun.E_input', t56_mcart_1)).
fof(t53_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:(k4_zfmisc_1(X1,X2,X3,X4)=k4_zfmisc_1(X5,X6,X7,X8)=>(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0|(X1=X5&(X2=X6&(X3=X7&X4=X8)))))), file('out/t56_mcart_1__12_premises__rerun__rerun.E_input', t53_mcart_1)).
fof(t49_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), file('out/t56_mcart_1__12_premises__rerun__rerun.E_input', t49_mcart_1)).
fof(t55_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0))), file('out/t56_mcart_1__12_premises__rerun__rerun.E_input', t55_mcart_1)).
fof(c_0_4, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:(k4_zfmisc_1(X1,X2,X3,X4)=k4_zfmisc_1(X5,X6,X7,X8)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|(X1=X5&(X2=X6&(X3=X7&X4=X8)))))))))), inference(assume_negation,[status(cth)],[t56_mcart_1])).
fof(c_0_5, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:(((X9=X13|k4_zfmisc_1(X9,X10,X11,X12)=k1_xboole_0)|k4_zfmisc_1(X9,X10,X11,X12)!=k4_zfmisc_1(X13,X14,X15,X16))&(((X10=X14|k4_zfmisc_1(X9,X10,X11,X12)=k1_xboole_0)|k4_zfmisc_1(X9,X10,X11,X12)!=k4_zfmisc_1(X13,X14,X15,X16))&(((X11=X15|k4_zfmisc_1(X9,X10,X11,X12)=k1_xboole_0)|k4_zfmisc_1(X9,X10,X11,X12)!=k4_zfmisc_1(X13,X14,X15,X16))&((X12=X16|k4_zfmisc_1(X9,X10,X11,X12)=k1_xboole_0)|k4_zfmisc_1(X9,X10,X11,X12)!=k4_zfmisc_1(X13,X14,X15,X16)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t53_mcart_1])])])).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:k4_zfmisc_1(X5,X6,X7,X8)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),X7),X8)), inference(variable_rename,[status(thm)],[t49_mcart_1])).
fof(c_0_7, negated_conjecture, ((k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0)=k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(esk1_0!=esk5_0|(esk2_0!=esk6_0|(esk3_0!=esk7_0|esk4_0!=esk8_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
cnf(c_0_8,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0|X4=X8|k4_zfmisc_1(X1,X2,X3,X4)!=k4_zfmisc_1(X5,X6,X7,X8)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0)=k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X8]:![X5]:![X6]:![X7]:![X8]:(((X5=k1_xboole_0|(X6=k1_xboole_0|(X7=k1_xboole_0|X8=k1_xboole_0)))|k4_zfmisc_1(X5,X6,X7,X8)!=k1_xboole_0)&((X5!=k1_xboole_0|k4_zfmisc_1(X5,X6,X7,X8)=k1_xboole_0)&((X6!=k1_xboole_0|k4_zfmisc_1(X5,X6,X7,X8)=k1_xboole_0)&((X7!=k1_xboole_0|k4_zfmisc_1(X5,X6,X7,X8)=k1_xboole_0)&(X8!=k1_xboole_0|k4_zfmisc_1(X5,X6,X7,X8)=k1_xboole_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])])])).
cnf(c_0_12,plain,(X4=X8|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),X7),X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9]), c_0_9]), c_0_9])).
cnf(c_0_13,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_9]), c_0_9])).
cnf(c_0_14,plain,(X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)=k1_xboole_0|esk8_0=X1|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16,plain,(X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_14, c_0_9])).
cnf(c_0_17,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)=k1_xboole_0|esk8_0=esk4_0), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_18,negated_conjecture,(esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19,negated_conjecture,(esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20,negated_conjecture,(esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21,negated_conjecture,(esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0|X2=X6|k4_zfmisc_1(X1,X2,X3,X4)!=k4_zfmisc_1(X5,X6,X7,X8)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_23,negated_conjecture,(esk8_0=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21])).
cnf(c_0_24,plain,(X2=X6|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),X7),X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_9]), c_0_9]), c_0_9])).
cnf(c_0_25,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk4_0)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)), inference(rw,[status(thm)],[c_0_13, c_0_23])).
cnf(c_0_26,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)=k1_xboole_0|X2=esk6_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_27,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)=k1_xboole_0|esk6_0=esk2_0), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_28,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0|X1=X5|k4_zfmisc_1(X1,X2,X3,X4)!=k4_zfmisc_1(X5,X6,X7,X8)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29,negated_conjecture,(esk6_0=esk2_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_27]), c_0_18]), c_0_19]), c_0_20]), c_0_21])).
cnf(c_0_30,plain,(X1=X5|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),X7),X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_9]), c_0_9]), c_0_9])).
cnf(c_0_31,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk2_0),esk7_0),esk4_0)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)), inference(rw,[status(thm)],[c_0_25, c_0_29])).
cnf(c_0_32,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)=k1_xboole_0|X1=esk5_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_33,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)=k1_xboole_0|esk5_0=esk1_0), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_34,negated_conjecture,(esk4_0!=esk8_0|esk3_0!=esk7_0|esk2_0!=esk6_0|esk1_0!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_35,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0|X3=X7|k4_zfmisc_1(X1,X2,X3,X4)!=k4_zfmisc_1(X5,X6,X7,X8)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_36,negated_conjecture,(esk5_0=esk1_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_33]), c_0_18]), c_0_19]), c_0_20]), c_0_21])).
cnf(c_0_37,negated_conjecture,(esk5_0!=esk1_0|esk6_0!=esk2_0|esk7_0!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_23])])).
cnf(c_0_38,plain,(X3=X7|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),X7),X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_9]), c_0_9]), c_0_9])).
cnf(c_0_39,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk7_0),esk4_0)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)), inference(rw,[status(thm)],[c_0_31, c_0_36])).
cnf(c_0_40,negated_conjecture,(esk5_0!=esk1_0|esk7_0!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_29])])).
cnf(c_0_41,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)=k1_xboole_0|esk7_0=X1|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)!=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X1),X4)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42,negated_conjecture,(esk7_0!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_36])])).
cnf(c_0_43,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_41]), c_0_42])).
cnf(c_0_44,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_43]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 36
# Proof object formula steps           : 9
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 15
# Processed clauses                    : 93
# ...of these trivial                  : 7
# ...subsumed                          : 18
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 18
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 30
# Generated clauses                    : 743
# ...of the previous two non-trivial   : 297
# Contextual simplify-reflections      : 0
# Paramodulations                      : 675
# Factorizations                       : 42
# Equation resolutions                 : 26
# Current number of processed clauses  : 16
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 13
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 140
# Rec. Clause-clause subsumption calls : 119
# Non-unit clause-clause subsumptions  : 20
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9600

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45224 pages
