# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(X2=k6_relat_1(X1,X3)<=>(![X4]:(r2_hidden(X4,k9_xtuple_0(X2))<=>(r2_hidden(X4,k9_xtuple_0(X3))&r2_hidden(k1_funct_1(X3,X4),X1)))&![X4]:(r2_hidden(X4,k9_xtuple_0(X2))=>k1_funct_1(X2,X4)=k1_funct_1(X3,X4))))))), file('out/t56_funct_1__44_premises__rerun__rerun.E_input', t53_funct_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t56_funct_1__44_premises__rerun__rerun.E_input', d3_tarski)).
fof(dt_k6_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>v1_relat_1(k6_relat_1(X1,X2)))), file('out/t56_funct_1__44_premises__rerun__rerun.E_input', dt_k6_relat_1)).
fof(fc9_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(v1_relat_1(k6_relat_1(X1,X2))&v1_funct_1(k6_relat_1(X1,X2))))), file('out/t56_funct_1__44_premises__rerun__rerun.E_input', fc9_funct_1)).
fof(t56_funct_1, conjecture, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>r1_tarski(k9_xtuple_0(k6_relat_1(X1,X2)),k9_xtuple_0(X2)))), file('out/t56_funct_1__44_premises__rerun__rerun.E_input', t56_funct_1)).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X9]:((((((((r2_hidden(X8,k9_xtuple_0(X7))|~r2_hidden(X8,k9_xtuple_0(X6)))|X6!=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((r2_hidden(k1_funct_1(X7,X8),X5)|~r2_hidden(X8,k9_xtuple_0(X6)))|X6!=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&(((((~r2_hidden(X8,k9_xtuple_0(X7))|~r2_hidden(k1_funct_1(X7,X8),X5))|r2_hidden(X8,k9_xtuple_0(X6)))|X6!=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&((((~r2_hidden(X9,k9_xtuple_0(X6))|k1_funct_1(X6,X9)=k1_funct_1(X7,X9))|X6!=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&((((((r2_hidden(esk2_3(X5,X6,X7),k9_xtuple_0(X6))|(~r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X6))|(~r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X7))|~r2_hidden(k1_funct_1(X7,esk1_3(X5,X6,X7)),X5))))|X6=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((k1_funct_1(X6,esk2_3(X5,X6,X7))!=k1_funct_1(X7,esk2_3(X5,X6,X7))|(~r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X6))|(~r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X7))|~r2_hidden(k1_funct_1(X7,esk1_3(X5,X6,X7)),X5))))|X6=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&((((((r2_hidden(esk2_3(X5,X6,X7),k9_xtuple_0(X6))|(r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X7))|r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X6))))|X6=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((k1_funct_1(X6,esk2_3(X5,X6,X7))!=k1_funct_1(X7,esk2_3(X5,X6,X7))|(r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X7))|r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X6))))|X6=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6))))&(((((r2_hidden(esk2_3(X5,X6,X7),k9_xtuple_0(X6))|(r2_hidden(k1_funct_1(X7,esk1_3(X5,X6,X7)),X5)|r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X6))))|X6=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6)))&((((k1_funct_1(X6,esk2_3(X5,X6,X7))!=k1_funct_1(X7,esk2_3(X5,X6,X7))|(r2_hidden(k1_funct_1(X7,esk1_3(X5,X6,X7)),X5)|r2_hidden(esk1_3(X5,X6,X7),k9_xtuple_0(X6))))|X6=k6_relat_1(X5,X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))|(~v1_relat_1(X6)|~v1_funct_1(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t53_funct_1])])])])])])])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk3_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk3_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_7,plain,(r2_hidden(X4,k9_xtuple_0(X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|X1!=k6_relat_1(X3,X2)|~r2_hidden(X4,k9_xtuple_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(r1_tarski(X1,X2)|r2_hidden(esk3_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, (![X3]:![X4]:(~v1_relat_1(X4)|v1_relat_1(k6_relat_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k6_relat_1])])).
fof(c_0_10, plain, (![X3]:![X4]:((v1_relat_1(k6_relat_1(X3,X4))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&(v1_funct_1(k6_relat_1(X3,X4))|(~v1_relat_1(X4)|~v1_funct_1(X4))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc9_funct_1])])])).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>r1_tarski(k9_xtuple_0(k6_relat_1(X1,X2)),k9_xtuple_0(X2))))), inference(assume_negation,[status(cth)],[t56_funct_1])).
cnf(c_0_12,plain,(r1_tarski(k9_xtuple_0(X1),X2)|r2_hidden(esk3_2(k9_xtuple_0(X1),X2),k9_xtuple_0(X3))|X1!=k6_relat_1(X4,X3)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_13,plain,(v1_relat_1(k6_relat_1(X1,X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(v1_funct_1(k6_relat_1(X2,X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, negated_conjecture, (((v1_relat_1(esk5_0)&v1_funct_1(esk5_0))&~r1_tarski(k9_xtuple_0(k6_relat_1(esk4_0,esk5_0)),k9_xtuple_0(esk5_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_16,plain,(r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17,plain,(r1_tarski(k9_xtuple_0(k6_relat_1(X1,X2)),X3)|r2_hidden(esk3_2(k9_xtuple_0(k6_relat_1(X1,X2)),X3),k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_12]), c_0_13]), c_0_14])).
cnf(c_0_18,negated_conjecture,(~r1_tarski(k9_xtuple_0(k6_relat_1(esk4_0,esk5_0)),k9_xtuple_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19,plain,(r1_tarski(k9_xtuple_0(k6_relat_1(X1,X2)),k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,negated_conjecture,(v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,negated_conjecture,(v1_relat_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 12
# Proof object formula steps           : 11
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 87
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 74
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 219
# ...of the previous two non-trivial   : 208
# Contextual simplify-reflections      : 21
# Paramodulations                      : 201
# Factorizations                       : 8
# Equation resolutions                 : 10
# Current number of processed clauses  : 56
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 158
# ...number of literals in the above   : 934
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 1328
# Rec. Clause-clause subsumption calls : 250
# Non-unit clause-clause subsumptions  : 34
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7747

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 46896 pages
