# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t56_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(k10_xtuple_0(X2),X1)&(k3_relat_1(X2,X3)=k4_relat_1(k9_xtuple_0(X4))&k3_relat_1(X3,X4)=k4_relat_1(X1)))=>X4=X2))))), file('out/t56_relat_1__3_premises__rerun.E_input', t56_relat_1)).
fof(t52_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>k3_relat_1(k4_relat_1(k9_xtuple_0(X1)),X1)=X1)), file('out/t56_relat_1__3_premises__rerun.E_input', t52_relat_1)).
fof(t53_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(k10_xtuple_0(X2),X1)=>k3_relat_1(X2,k4_relat_1(X1))=X2))), file('out/t56_relat_1__3_premises__rerun.E_input', t53_relat_1)).
fof(t36_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>k3_relat_1(k3_relat_1(X1,X2),X3)=k3_relat_1(X1,k3_relat_1(X2,X3)))))), file('out/t56_relat_1__3_premises__rerun.E_input', t36_relat_1)).
fof(c_0_4, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(k10_xtuple_0(X2),X1)&(k3_relat_1(X2,X3)=k4_relat_1(k9_xtuple_0(X4))&k3_relat_1(X3,X4)=k4_relat_1(X1)))=>X4=X2)))))), inference(assume_negation,[status(cth)],[t56_relat_1])).
fof(c_0_5, plain, (![X2]:(~v1_relat_1(X2)|k3_relat_1(k4_relat_1(k9_xtuple_0(X2)),X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_relat_1])])).
fof(c_0_6, negated_conjecture, ((v1_relat_1(esk2_0)&(v1_relat_1(esk3_0)&(v1_relat_1(esk4_0)&((r1_tarski(k10_xtuple_0(esk2_0),esk1_0)&(k3_relat_1(esk2_0,esk3_0)=k4_relat_1(k9_xtuple_0(esk4_0))&k3_relat_1(esk3_0,esk4_0)=k4_relat_1(esk1_0)))&esk4_0!=esk2_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
fof(c_0_7, plain, (![X3]:![X4]:(~v1_relat_1(X4)|(~r1_tarski(k10_xtuple_0(X4),X3)|k3_relat_1(X4,k4_relat_1(X3))=X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t53_relat_1])])).
fof(c_0_8, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X4)|(~v1_relat_1(X5)|(~v1_relat_1(X6)|k3_relat_1(k3_relat_1(X4,X5),X6)=k3_relat_1(X4,k3_relat_1(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t36_relat_1])])])])])).
cnf(c_0_9,plain,(k3_relat_1(k4_relat_1(k9_xtuple_0(X1)),X1)=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,negated_conjecture,(k3_relat_1(esk2_0,esk3_0)=k4_relat_1(k9_xtuple_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,plain,(k3_relat_1(X1,k4_relat_1(X2))=X1|~r1_tarski(k10_xtuple_0(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,negated_conjecture,(r1_tarski(k10_xtuple_0(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,negated_conjecture,(k3_relat_1(esk3_0,esk4_0)=k4_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,plain,(k3_relat_1(k3_relat_1(X1,X2),X3)=k3_relat_1(X1,k3_relat_1(X2,X3))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17,negated_conjecture,(k3_relat_1(k3_relat_1(esk2_0,esk3_0),esk4_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])])).
cnf(c_0_18,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19,negated_conjecture,(k3_relat_1(esk2_0,k3_relat_1(esk3_0,esk4_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])])).
cnf(c_0_20,negated_conjecture,(esk4_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_11]), c_0_18]), c_0_15])]), c_0_19]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 22
# Proof object clause steps            : 13
# Proof object formula steps           : 9
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 10
# Processed clauses                    : 23
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 22
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 7
# ...of the previous two non-trivial   : 7
# Contextual simplify-reflections      : 0
# Paramodulations                      : 7
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 12
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 16
# Current number of archived formulas  : 0
# Current number of archived clauses   : 10
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 688

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 42988 pages
