# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_zfmisc_1, axiom, (![X1]:![X2]:![X3]:![X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), file('out/t59_mcart_1__12_premises__rerun.E_input', d4_zfmisc_1)).
fof(d3_zfmisc_1, axiom, (![X1]:![X2]:![X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), file('out/t59_mcart_1__12_premises__rerun.E_input', d3_zfmisc_1)).
fof(d4_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), file('out/t59_mcart_1__12_premises__rerun.E_input', d4_mcart_1)).
fof(d3_mcart_1, axiom, (![X1]:![X2]:![X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), file('out/t59_mcart_1__12_premises__rerun.E_input', d3_mcart_1)).
fof(d8_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>![X6]:(m1_subset_1(X6,X1)=>(X6=k8_mcart_1(X1,X2,X3,X4,X5)<=>![X7]:![X8]:![X9]:![X10]:(X5=k4_mcart_1(X7,X8,X9,X10)=>X6=X7))))))))))), file('out/t59_mcart_1__12_premises__rerun.E_input', d8_mcart_1)).
fof(dt_k8_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k8_mcart_1(X1,X2,X3,X4,X5),X1))), file('out/t59_mcart_1__12_premises__rerun.E_input', dt_k8_mcart_1)).
fof(t59_mcart_1, conjecture, (![X1]:![X2]:![X3]:![X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))&?[X6]:?[X7]:?[X8]:?[X9]:(X5=k4_mcart_1(X6,X7,X8,X9)&~((k8_mcart_1(X1,X2,X3,X4,X5)=X6&(k9_mcart_1(X1,X2,X3,X4,X5)=X7&(k10_mcart_1(X1,X2,X3,X4,X5)=X8&k11_mcart_1(X1,X2,X3,X4,X5)=X9)))))))))))), file('out/t59_mcart_1__12_premises__rerun.E_input', t59_mcart_1)).
fof(d9_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>![X6]:(m1_subset_1(X6,X2)=>(X6=k9_mcart_1(X1,X2,X3,X4,X5)<=>![X7]:![X8]:![X9]:![X10]:(X5=k4_mcart_1(X7,X8,X9,X10)=>X6=X8))))))))))), file('out/t59_mcart_1__12_premises__rerun.E_input', d9_mcart_1)).
fof(dt_k9_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k9_mcart_1(X1,X2,X3,X4,X5),X2))), file('out/t59_mcart_1__12_premises__rerun.E_input', dt_k9_mcart_1)).
fof(d10_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>![X6]:(m1_subset_1(X6,X3)=>(X6=k10_mcart_1(X1,X2,X3,X4,X5)<=>![X7]:![X8]:![X9]:![X10]:(X5=k4_mcart_1(X7,X8,X9,X10)=>X6=X9))))))))))), file('out/t59_mcart_1__12_premises__rerun.E_input', d10_mcart_1)).
fof(dt_k10_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k10_mcart_1(X1,X2,X3,X4,X5),X3))), file('out/t59_mcart_1__12_premises__rerun.E_input', dt_k10_mcart_1)).
fof(d11_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>![X6]:(m1_subset_1(X6,X4)=>(X6=k11_mcart_1(X1,X2,X3,X4,X5)<=>![X7]:![X8]:![X9]:![X10]:(X5=k4_mcart_1(X7,X8,X9,X10)=>X6=X10))))))))))), file('out/t59_mcart_1__12_premises__rerun.E_input', d11_mcart_1)).
fof(dt_k11_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k11_mcart_1(X1,X2,X3,X4,X5),X4))), file('out/t59_mcart_1__12_premises__rerun.E_input', dt_k11_mcart_1)).
fof(c_0_13, plain, (![X5]:![X6]:![X7]:![X8]:k4_zfmisc_1(X5,X6,X7,X8)=k2_zfmisc_1(k3_zfmisc_1(X5,X6,X7),X8)), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:k3_zfmisc_1(X4,X5,X6)=k2_zfmisc_1(k2_zfmisc_1(X4,X5),X6)), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:![X8]:k4_mcart_1(X5,X6,X7,X8)=k4_tarski(k3_mcart_1(X5,X6,X7),X8)), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:k3_mcart_1(X4,X5,X6)=k4_tarski(k4_tarski(X4,X5),X6)), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_17, plain, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:((((((((X16!=k8_mcart_1(X11,X12,X13,X14,X15)|(X15!=k4_mcart_1(X17,X18,X19,X20)|X16=X17))|~m1_subset_1(X16,X11))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&((((((((X15=k4_mcart_1(esk1_6(X11,X12,X13,X14,X15,X16),esk2_6(X11,X12,X13,X14,X15,X16),esk3_6(X11,X12,X13,X14,X15,X16),esk4_6(X11,X12,X13,X14,X15,X16))|X16=k8_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X11))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&(((((((X16!=esk1_6(X11,X12,X13,X14,X15,X16)|X16=k8_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X11))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_mcart_1])])])])])])])])).
cnf(c_0_18,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(~m1_subset_1(X10,k4_zfmisc_1(X6,X7,X8,X9))|m1_subset_1(k8_mcart_1(X6,X7,X8,X9,X10),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_mcart_1])])).
fof(c_0_23, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))&?[X6]:?[X7]:?[X8]:?[X9]:(X5=k4_mcart_1(X6,X7,X8,X9)&~((k8_mcart_1(X1,X2,X3,X4,X5)=X6&(k9_mcart_1(X1,X2,X3,X4,X5)=X7&(k10_mcart_1(X1,X2,X3,X4,X5)=X8&k11_mcart_1(X1,X2,X3,X4,X5)=X9))))))))))))), inference(assume_negation,[status(cth)],[t59_mcart_1])).
cnf(c_0_24,plain,(X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X6=X7|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))|~m1_subset_1(X6,X1)|X5!=k4_mcart_1(X7,X8,X9,X10)|X6!=k8_mcart_1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26,plain,(k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,plain,(m1_subset_1(k8_mcart_1(X1,X2,X3,X4,X5),X1)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, negated_conjecture, ((esk5_0!=k1_xboole_0&(esk6_0!=k1_xboole_0&(esk7_0!=k1_xboole_0&(esk8_0!=k1_xboole_0&(m1_subset_1(esk9_0,k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))&(esk9_0=k4_mcart_1(esk10_0,esk11_0,esk12_0,esk13_0)&(k8_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk10_0|(k9_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk11_0|(k10_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk12_0|k11_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk13_0)))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])])])).
fof(c_0_29, plain, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:((((((((X16!=k9_mcart_1(X11,X12,X13,X14,X15)|(X15!=k4_mcart_1(X17,X18,X19,X20)|X16=X18))|~m1_subset_1(X16,X12))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&((((((((X15=k4_mcart_1(esk14_6(X11,X12,X13,X14,X15,X16),esk15_6(X11,X12,X13,X14,X15,X16),esk16_6(X11,X12,X13,X14,X15,X16),esk17_6(X11,X12,X13,X14,X15,X16))|X16=k9_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X12))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&(((((((X16!=esk15_6(X11,X12,X13,X14,X15,X16)|X16=k9_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X12))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_mcart_1])])])])])])])])).
fof(c_0_30, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(~m1_subset_1(X10,k4_zfmisc_1(X6,X7,X8,X9))|m1_subset_1(k9_mcart_1(X6,X7,X8,X9,X10),X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k9_mcart_1])])).
cnf(c_0_31,plain,(X6=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X6!=k8_mcart_1(X1,X2,X3,X4,X5)|X5!=k4_tarski(k4_tarski(k4_tarski(X7,X8),X9),X10)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_32,plain,(m1_subset_1(k8_mcart_1(X1,X2,X3,X4,X5),X1)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_27, c_0_25])).
cnf(c_0_33,negated_conjecture,(esk9_0=k4_mcart_1(esk10_0,esk11_0,esk12_0,esk13_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34,plain,(X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X6=X7|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))|~m1_subset_1(X6,X2)|X5!=k4_mcart_1(X8,X7,X9,X10)|X6!=k9_mcart_1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35,plain,(m1_subset_1(k9_mcart_1(X1,X2,X3,X4,X5),X2)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_36, plain, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:((((((((X16!=k10_mcart_1(X11,X12,X13,X14,X15)|(X15!=k4_mcart_1(X17,X18,X19,X20)|X16=X19))|~m1_subset_1(X16,X13))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&((((((((X15=k4_mcart_1(esk18_6(X11,X12,X13,X14,X15,X16),esk19_6(X11,X12,X13,X14,X15,X16),esk20_6(X11,X12,X13,X14,X15,X16),esk21_6(X11,X12,X13,X14,X15,X16))|X16=k10_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X13))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&(((((((X16!=esk20_6(X11,X12,X13,X14,X15,X16)|X16=k10_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X13))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_mcart_1])])])])])])])])).
fof(c_0_37, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(~m1_subset_1(X10,k4_zfmisc_1(X6,X7,X8,X9))|m1_subset_1(k10_mcart_1(X6,X7,X8,X9,X10),X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_mcart_1])])).
cnf(c_0_38,plain,(k8_mcart_1(X1,X2,X3,X4,k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8))=X5|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])]), c_0_32])).
cnf(c_0_39,negated_conjecture,(esk9_0=k4_tarski(k4_tarski(k4_tarski(esk10_0,esk11_0),esk12_0),esk13_0)), inference(rw,[status(thm)],[c_0_33, c_0_26])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk9_0,k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41,plain,(X6=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X6!=k9_mcart_1(X1,X2,X3,X4,X5)|X5!=k4_tarski(k4_tarski(k4_tarski(X8,X7),X9),X10)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_25]), c_0_26])).
cnf(c_0_42,plain,(m1_subset_1(k9_mcart_1(X1,X2,X3,X4,X5),X2)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_35, c_0_25])).
cnf(c_0_43,plain,(X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X6=X7|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))|~m1_subset_1(X6,X3)|X5!=k4_mcart_1(X8,X9,X7,X10)|X6!=k10_mcart_1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44,plain,(m1_subset_1(k10_mcart_1(X1,X2,X3,X4,X5),X3)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_45, plain, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:((((((((X16!=k11_mcart_1(X11,X12,X13,X14,X15)|(X15!=k4_mcart_1(X17,X18,X19,X20)|X16=X20))|~m1_subset_1(X16,X14))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&((((((((X15=k4_mcart_1(esk22_6(X11,X12,X13,X14,X15,X16),esk23_6(X11,X12,X13,X14,X15,X16),esk24_6(X11,X12,X13,X14,X15,X16),esk25_6(X11,X12,X13,X14,X15,X16))|X16=k11_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X14))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)&(((((((X16!=esk25_6(X11,X12,X13,X14,X15,X16)|X16=k11_mcart_1(X11,X12,X13,X14,X15))|~m1_subset_1(X16,X14))|~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14)))|X14=k1_xboole_0)|X13=k1_xboole_0)|X12=k1_xboole_0)|X11=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_mcart_1])])])])])])])])).
fof(c_0_46, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(~m1_subset_1(X10,k4_zfmisc_1(X6,X7,X8,X9))|m1_subset_1(k11_mcart_1(X6,X7,X8,X9,X10),X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k11_mcart_1])])).
cnf(c_0_47,negated_conjecture,(k8_mcart_1(X1,X2,X3,X4,esk9_0)=esk10_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk9_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48,negated_conjecture,(m1_subset_1(esk9_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))), inference(rw,[status(thm)],[c_0_40, c_0_25])).
cnf(c_0_49,negated_conjecture,(esk8_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_50,negated_conjecture,(esk7_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_51,negated_conjecture,(esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_52,negated_conjecture,(esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_53,plain,(k9_mcart_1(X1,X2,X3,X4,k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8))=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_41])]), c_0_42])).
cnf(c_0_54,plain,(X6=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X6!=k10_mcart_1(X1,X2,X3,X4,X5)|X5!=k4_tarski(k4_tarski(k4_tarski(X8,X9),X7),X10)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_25]), c_0_26])).
cnf(c_0_55,plain,(m1_subset_1(k10_mcart_1(X1,X2,X3,X4,X5),X3)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_44, c_0_25])).
cnf(c_0_56,plain,(X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X6=X7|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))|~m1_subset_1(X6,X4)|X5!=k4_mcart_1(X8,X9,X10,X7)|X6!=k11_mcart_1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57,plain,(m1_subset_1(k11_mcart_1(X1,X2,X3,X4,X5),X4)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58,negated_conjecture,(k11_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk13_0|k10_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk12_0|k9_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk11_0|k8_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk10_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_59,negated_conjecture,(k8_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)=esk10_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_60,negated_conjecture,(k9_mcart_1(X1,X2,X3,X4,esk9_0)=esk11_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk9_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_53, c_0_39])).
cnf(c_0_61,plain,(k10_mcart_1(X1,X2,X3,X4,k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8))=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_54])]), c_0_55])).
cnf(c_0_62,plain,(X6=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X6!=k11_mcart_1(X1,X2,X3,X4,X5)|X5!=k4_tarski(k4_tarski(k4_tarski(X8,X9),X10),X7)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_25]), c_0_26])).
cnf(c_0_63,plain,(m1_subset_1(k11_mcart_1(X1,X2,X3,X4,X5),X4)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_57, c_0_25])).
cnf(c_0_64,negated_conjecture,(k9_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk11_0|k10_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk12_0|k11_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk13_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_65,negated_conjecture,(k9_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)=esk11_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_48]), c_0_49]), c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_66,negated_conjecture,(k10_mcart_1(X1,X2,X3,X4,esk9_0)=esk12_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk9_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_61, c_0_39])).
cnf(c_0_67,plain,(k11_mcart_1(X1,X2,X3,X4,k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8))=X8|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_62])]), c_0_63])).
cnf(c_0_68,negated_conjecture,(k10_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk12_0|k11_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk13_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65])])).
cnf(c_0_69,negated_conjecture,(k10_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)=esk12_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_48]), c_0_49]), c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_70,negated_conjecture,(k11_mcart_1(X1,X2,X3,X4,esk9_0)=esk13_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk9_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_67, c_0_39])).
cnf(c_0_71,negated_conjecture,(k11_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=esk13_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69])])).
cnf(c_0_72,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_48]), c_0_71]), c_0_49]), c_0_50]), c_0_51]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 46
# Proof object formula steps           : 27
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 23
# Processed clauses                    : 73
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 91
# ...of the previous two non-trivial   : 96
# Contextual simplify-reflections      : 4
# Paramodulations                      : 87
# Factorizations                       : 0
# Equation resolutions                 : 8
# Current number of processed clauses  : 40
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 27
# Current number of unprocessed clauses: 69
# ...number of literals in the above   : 782
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 145
# Rec. Clause-clause subsumption calls : 19
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6020

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45404 pages
