# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l62_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6))), file('out/t59_enumset1__7_premises__rerun.E_input', l62_enumset1)).
fof(t98_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), file('out/t59_enumset1__7_premises__rerun.E_input', t98_xboole_1)).
fof(t59_enumset1, conjecture, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_enumset1(X5,X6,X7))), file('out/t59_enumset1__7_premises__rerun.E_input', t59_enumset1)).
fof(t44_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), file('out/t59_enumset1__7_premises__rerun.E_input', t44_enumset1)).
fof(t56_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_tarski(X1),k4_enumset1(X2,X3,X4,X5,X6,X7))), file('out/t59_enumset1__7_premises__rerun.E_input', t56_enumset1)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t59_enumset1__7_premises__rerun.E_input', t91_xboole_1)).
fof(t41_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t59_enumset1__7_premises__rerun.E_input', t41_xboole_1)).
fof(t4_xboole_1, axiom, (![X1]:![X2]:![X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t59_enumset1__7_premises__rerun.E_input', t4_xboole_1)).
fof(c_0_8, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:k4_enumset1(X7,X8,X9,X10,X11,X12)=k2_xboole_0(k1_enumset1(X7,X8,X9),k1_enumset1(X10,X11,X12))), inference(variable_rename,[status(thm)],[l62_enumset1])).
fof(c_0_9, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k5_xboole_0(X3,k4_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_10, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_enumset1(X5,X6,X7)))), inference(assume_negation,[status(cth)],[t59_enumset1])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X5),k1_enumset1(X6,X7,X8))), inference(variable_rename,[status(thm)],[t44_enumset1])).
fof(c_0_12, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:k5_enumset1(X8,X9,X10,X11,X12,X13,X14)=k2_xboole_0(k1_tarski(X8),k4_enumset1(X9,X10,X11,X12,X13,X14))), inference(variable_rename,[status(thm)],[t56_enumset1])).
cnf(c_0_13,plain,(k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_15, negated_conjecture, (k5_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)!=k2_xboole_0(k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0),k1_enumset1(esk5_0,esk6_0,esk7_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_16,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_tarski(X1),k4_enumset1(X2,X3,X4,X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(k4_enumset1(X1,X2,X3,X4,X5,X6)=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19,negated_conjecture,(k5_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)!=k2_xboole_0(k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0),k1_enumset1(esk5_0,esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,plain,(k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k1_tarski(X1),k4_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)))), inference(rw,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_21,plain,(k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k5_xboole_0(k1_tarski(X1),k4_xboole_0(k5_xboole_0(k1_enumset1(X2,X3,X4),k4_xboole_0(k1_enumset1(X5,X6,X7),k1_enumset1(X2,X3,X4))),k1_tarski(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_14]), c_0_18])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:k4_xboole_0(k4_xboole_0(X4,X5),X6)=k4_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_24, plain, (![X4]:![X5]:![X6]:k2_xboole_0(k2_xboole_0(X4,X5),X6)=k2_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t4_xboole_1])).
cnf(c_0_25,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k4_xboole_0(k5_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k4_xboole_0(k1_enumset1(esk5_0,esk6_0,esk7_0),k1_enumset1(esk2_0,esk3_0,esk4_0))),k1_tarski(esk1_0)))!=k5_xboole_0(k5_xboole_0(k1_tarski(esk1_0),k4_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k1_tarski(esk1_0))),k4_xboole_0(k1_enumset1(esk5_0,esk6_0,esk7_0),k5_xboole_0(k1_tarski(esk1_0),k4_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k1_tarski(esk1_0)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_14]), c_0_20]), c_0_20]), c_0_21])).
cnf(c_0_26,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,plain,(k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k4_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k1_tarski(esk1_0)),k4_xboole_0(k1_enumset1(esk5_0,esk6_0,esk7_0),k5_xboole_0(k1_tarski(esk1_0),k4_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k1_tarski(esk1_0))))))!=k5_xboole_0(k1_tarski(esk1_0),k4_xboole_0(k5_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k4_xboole_0(k1_enumset1(esk5_0,esk6_0,esk7_0),k1_enumset1(esk2_0,esk3_0,esk4_0))),k1_tarski(esk1_0)))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_27, c_0_14])).
cnf(c_0_31,plain,(k5_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k4_xboole_0(X3,k5_xboole_0(X1,k4_xboole_0(X2,X1))))=k5_xboole_0(X1,k4_xboole_0(k5_xboole_0(X2,k4_xboole_0(X3,X2)),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_14]), c_0_14]), c_0_14]), c_0_14])).
cnf(c_0_32,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k4_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k1_tarski(esk1_0)),k4_xboole_0(k4_xboole_0(k1_enumset1(esk5_0,esk6_0,esk7_0),k1_tarski(esk1_0)),k1_enumset1(esk2_0,esk3_0,esk4_0))))!=k5_xboole_0(k1_tarski(esk1_0),k4_xboole_0(k5_xboole_0(k1_enumset1(esk2_0,esk3_0,esk4_0),k4_xboole_0(k1_enumset1(esk5_0,esk6_0,esk7_0),k1_enumset1(esk2_0,esk3_0,esk4_0))),k1_tarski(esk1_0)))), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33,plain,(k5_xboole_0(X1,k5_xboole_0(k4_xboole_0(X2,X1),k4_xboole_0(k4_xboole_0(X3,X1),X2)))=k5_xboole_0(X1,k4_xboole_0(k5_xboole_0(X2,k4_xboole_0(X3,X2)),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_30]), c_0_26])).
cnf(c_0_34,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 18
# Proof object formula steps           : 17
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 0
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 8
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 4
# Processed clauses                    : 5
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 5
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 0
# ...of the previous two non-trivial   : 1
# Contextual simplify-reflections      : 0
# Paramodulations                      : 0
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 605

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45236 pages
