# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t59_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:~((k3_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)&(r2_hidden(X2,X3)&~(X1=X2))))), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', t59_zfmisc_1)).
fof(t90_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k4_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', t90_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t47_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', t47_xboole_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', t70_enumset1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t55_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((r1_xboole_0(k2_tarski(X1,X2),X3)&r2_hidden(X1,X3)))), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', t55_zfmisc_1)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(symmetry_r1_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1))), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', symmetry_r1_xboole_0)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', d1_tarski)).
fof(t1_xboole_0, axiom, (![X1]:![X2]:![X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3))))), file('out/t59_zfmisc_1__11_premises__rerun__rerun.E_input', t1_xboole_0)).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:![X3]:~((k3_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)&(r2_hidden(X2,X3)&~(X1=X2)))))), inference(assume_negation,[status(cth)],[t59_zfmisc_1])).
fof(c_0_12, plain, (![X3]:![X4]:r1_xboole_0(k4_xboole_0(X3,k3_xboole_0(X3,X4)),X4)), inference(variable_rename,[status(thm)],[t90_xboole_1])).
fof(c_0_13, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_14, plain, (![X3]:![X4]:k4_xboole_0(X3,k3_xboole_0(X3,X4))=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_15, negated_conjecture, ((k3_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_tarski(esk1_0)&(r2_hidden(esk2_0,esk3_0)&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
fof(c_0_16, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_17,plain,(r1_xboole_0(k4_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,negated_conjecture,(k3_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:(~r1_xboole_0(k2_tarski(X4,X5),X6)|~r2_hidden(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_zfmisc_1])])).
fof(c_0_24, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_25,plain,(r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2))),X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26,plain,(k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_18]), c_0_18])).
cnf(c_0_27,negated_conjecture,(k3_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)=k1_tarski(esk1_0)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(~r2_hidden(X1,X2)|~r1_xboole_0(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|r1_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_32,plain,(r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33,negated_conjecture,(k3_xboole_0(esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0))=k1_tarski(esk1_0)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34,plain,(~r2_hidden(X1,X2)|~r1_xboole_0(k1_enumset1(X1,X1,X3),X2)), inference(rw,[status(thm)],[c_0_29, c_0_21])).
cnf(c_0_35,plain,(k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_21]), c_0_21])).
cnf(c_0_36,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37,negated_conjecture,(r1_xboole_0(k5_xboole_0(esk3_0,k1_tarski(esk1_0)),k1_enumset1(esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_38, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk4_2(X4,X5),X5)|esk4_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_39,plain,(~r1_xboole_0(k1_enumset1(X1,X1,X2),X3)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40,negated_conjecture,(r1_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),k5_xboole_0(esk3_0,k1_tarski(esk1_0)))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_41, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:((((~r2_hidden(X4,X5)|~r2_hidden(X4,X6))|~r2_hidden(X4,k5_xboole_0(X5,X6)))&((r2_hidden(X4,X5)|r2_hidden(X4,X6))|~r2_hidden(X4,k5_xboole_0(X5,X6))))&(((~r2_hidden(X4,X5)|r2_hidden(X4,X6))|r2_hidden(X4,k5_xboole_0(X5,X6)))&((~r2_hidden(X4,X6)|r2_hidden(X4,X5))|r2_hidden(X4,k5_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])])])).
cnf(c_0_42,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43,negated_conjecture,(~r2_hidden(esk2_0,k5_xboole_0(esk3_0,k1_tarski(esk1_0)))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44,plain,(r2_hidden(X1,k5_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_45,negated_conjecture,(r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_46,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_47,negated_conjecture,(r2_hidden(esk2_0,k1_tarski(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_48,negated_conjecture,(esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_49,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 27
# Proof object formula steps           : 23
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 17
# Processed clauses                    : 56
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 51
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 64
# ...of the previous two non-trivial   : 57
# Contextual simplify-reflections      : 0
# Paramodulations                      : 62
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 11
# Current number of unprocessed clauses: 35
# ...number of literals in the above   : 70
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 25
# Rec. Clause-clause subsumption calls : 23
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1709

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45252 pages
