# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t59_xboole_1__4_premises__rerun.E_input', t1_xboole_1)).
fof(d8_xboole_0, axiom, (![X1]:![X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2)))), file('out/t59_xboole_1__4_premises__rerun.E_input', d8_xboole_0)).
fof(t59_xboole_1, conjecture, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r2_xboole_0(X2,X3))=>r2_xboole_0(X1,X3))), file('out/t59_xboole_1__4_premises__rerun.E_input', t59_xboole_1)).
fof(antisymmetry_r2_xboole_0, axiom, (![X1]:![X2]:(r2_xboole_0(X1,X2)=>~(r2_xboole_0(X2,X1)))), file('out/t59_xboole_1__4_premises__rerun.E_input', antisymmetry_r2_xboole_0)).
fof(irreflexivity_r2_xboole_0, axiom, (![X1]:![X2]:~(r2_xboole_0(X1,X1))), file('out/t59_xboole_1__4_premises__rerun.E_input', irreflexivity_r2_xboole_0)).
fof(c_0_5, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_6, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|~r2_xboole_0(X3,X4))&(X3!=X4|~r2_xboole_0(X3,X4)))&((~r1_tarski(X3,X4)|X3=X4)|r2_xboole_0(X3,X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])])])).
fof(c_0_7, negated_conjecture, (~(![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r2_xboole_0(X2,X3))=>r2_xboole_0(X1,X3)))), inference(assume_negation,[status(cth)],[t59_xboole_1])).
cnf(c_0_8,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(r1_tarski(X1,X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, (((r1_tarski(esk1_0,esk2_0)&r2_xboole_0(esk2_0,esk3_0))&~r2_xboole_0(esk1_0,esk3_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_11,plain,(r1_tarski(X1,X2)|~r2_xboole_0(X3,X2)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12,negated_conjecture,(r2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, (![X3]:![X4]:(~r2_xboole_0(X3,X4)|~r2_xboole_0(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_xboole_0])])])).
cnf(c_0_14,plain,(r2_xboole_0(X1,X2)|X1=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,negated_conjecture,(r1_tarski(X1,esk3_0)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17,plain,(~r2_xboole_0(X1,X2)|~r2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,(esk1_0=esk2_0|r2_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19,negated_conjecture,(X1=esk3_0|r2_xboole_0(X1,esk3_0)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_14, c_0_16])).
cnf(c_0_20,negated_conjecture,(~r2_xboole_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21,negated_conjecture,(esk1_0=esk2_0|~r2_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,negated_conjecture,(esk1_0=esk3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_15]), c_0_20])).
fof(c_0_23, plain, (![X3]:~r2_xboole_0(X3,X3)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[irreflexivity_r2_xboole_0])])])).
cnf(c_0_24,negated_conjecture,(esk3_0=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_22]), c_0_12])])).
cnf(c_0_25,plain,(~r2_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_26,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_24]), c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 27
# Proof object clause steps            : 16
# Proof object formula steps           : 11
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 9
# Processed clauses                    : 32
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 29
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 13
# Generated clauses                    : 21
# ...of the previous two non-trivial   : 29
# Contextual simplify-reflections      : 1
# Paramodulations                      : 20
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 7
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 9
# ...number of literals in the above   : 17
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 25
# Rec. Clause-clause subsumption calls : 24
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 681

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 42992 pages
