# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k5_setfam_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2))), file('out/t5_pre_topc__5_premises__rerun.E_input', redefinition_k5_setfam_1)).
fof(t4_subset_1, axiom, (![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), file('out/t5_pre_topc__5_premises__rerun.E_input', t4_subset_1)).
fof(d1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1)))))))))), file('out/t5_pre_topc__5_premises__rerun.E_input', d1_pre_topc)).
fof(t2_zfmisc_1, axiom, (k3_tarski(k1_xboole_0)=k1_xboole_0), file('out/t5_pre_topc__5_premises__rerun.E_input', t2_zfmisc_1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t5_pre_topc__5_premises__rerun.E_input', t2_xboole_1)).
fof(t5_pre_topc, conjecture, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>r2_hidden(k1_xboole_0,u1_pre_topc(X1)))), file('out/t5_pre_topc__5_premises__rerun.E_input', t5_pre_topc)).
fof(c_0_6, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|k5_setfam_1(X3,X4)=k3_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
fof(c_0_7, plain, (![X2]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_8, plain, (![X4]:![X5]:![X6]:![X7]:((((r2_hidden(u1_struct_0(X4),u1_pre_topc(X4))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&((((~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|(~r1_tarski(X5,u1_pre_topc(X4))|r2_hidden(k5_setfam_1(u1_struct_0(X4),X5),u1_pre_topc(X4))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))&(((~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X4)))|((~r2_hidden(X6,u1_pre_topc(X4))|~r2_hidden(X7,u1_pre_topc(X4)))|r2_hidden(k9_subset_1(u1_struct_0(X4),X6,X7),u1_pre_topc(X4)))))|~v2_pre_topc(X4))|~l1_pre_topc(X4))))&((((((m1_subset_1(esk2_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk1_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk3_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|m1_subset_1(esk1_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk2_1(X4),u1_pre_topc(X4))|m1_subset_1(esk1_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk3_1(X4),u1_pre_topc(X4))|m1_subset_1(esk1_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk2_1(X4),esk3_1(X4)),u1_pre_topc(X4))|m1_subset_1(esk1_1(X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&((((((m1_subset_1(esk2_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk1_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk3_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|r1_tarski(esk1_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk2_1(X4),u1_pre_topc(X4))|r1_tarski(esk1_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk3_1(X4),u1_pre_topc(X4))|r1_tarski(esk1_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk2_1(X4),esk3_1(X4)),u1_pre_topc(X4))|r1_tarski(esk1_1(X4),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))))&(((((m1_subset_1(esk2_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk1_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&(((((m1_subset_1(esk3_1(X4),k1_zfmisc_1(u1_struct_0(X4)))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk1_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((((r2_hidden(esk2_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk1_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))&((((r2_hidden(esk3_1(X4),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk1_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4)))&((((~r2_hidden(k9_subset_1(u1_struct_0(X4),esk2_1(X4),esk3_1(X4)),u1_pre_topc(X4))|~r2_hidden(k5_setfam_1(u1_struct_0(X4),esk1_1(X4)),u1_pre_topc(X4)))|~r2_hidden(u1_struct_0(X4),u1_pre_topc(X4)))|v2_pre_topc(X4))|~l1_pre_topc(X4))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])])])).
cnf(c_0_9,plain,(k5_setfam_1(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(k3_tarski(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t2_zfmisc_1])).
fof(c_0_12, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_13, negated_conjecture, (~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>r2_hidden(k1_xboole_0,u1_pre_topc(X1))))), inference(assume_negation,[status(cth)],[t5_pre_topc])).
cnf(c_0_14,plain,(r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r1_tarski(X2,u1_pre_topc(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,plain,(k5_setfam_1(X1,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])).
cnf(c_0_16,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, negated_conjecture, (((v2_pre_topc(esk4_0)&l1_pre_topc(esk4_0))&~r2_hidden(k1_xboole_0,u1_pre_topc(esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_18,plain,(r2_hidden(k1_xboole_0,u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_10]), c_0_15]), c_0_16])])).
cnf(c_0_19,negated_conjecture,(v2_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,negated_conjecture,(l1_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21,negated_conjecture,(~r2_hidden(k1_xboole_0,u1_pre_topc(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 11
# Proof object formula steps           : 12
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 3
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 0
# Paramodulations                      : 3
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 32
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 112
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2619

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45328 pages
