# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_setfam_1, axiom, (![X1]:![X2]:((~(X1=k1_xboole_0)=>(X2=k1_setfam_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>![X4]:(r2_hidden(X4,X1)=>r2_hidden(X3,X4)))))&(X1=k1_xboole_0=>(X2=k1_setfam_1(X1)<=>X2=k1_xboole_0)))), file('out/t5_setfam_1__7_premises__rerun__rerun.E_input', d1_setfam_1)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t5_setfam_1__7_premises__rerun__rerun.E_input', t3_xboole_0)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t5_setfam_1__7_premises__rerun__rerun.E_input', d1_xboole_0)).
fof(t66_xboole_1, axiom, (![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1))))), file('out/t5_setfam_1__7_premises__rerun__rerun.E_input', t66_xboole_1)).
fof(t5_setfam_1, conjecture, (![X1]:(r2_hidden(k1_xboole_0,X1)=>k1_setfam_1(X1)=k1_xboole_0)), file('out/t5_setfam_1__7_premises__rerun__rerun.E_input', t5_setfam_1)).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:![X6]:![X12]:![X5]:![X6]:![X6]:((((((~r2_hidden(X7,X6)|(~r2_hidden(X8,X5)|r2_hidden(X7,X8)))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&((((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(X7,esk1_3(X5,X6,X7))|r2_hidden(X7,X6))|X6!=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((((r2_hidden(esk3_2(X5,X6),X5)|~r2_hidden(esk2_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)&(((~r2_hidden(esk2_2(X5,X6),esk3_2(X5,X6))|~r2_hidden(esk2_2(X5,X6),X6))|X6=k1_setfam_1(X5))|X5=k1_xboole_0))&(((r2_hidden(esk2_2(X5,X6),X6)|(~r2_hidden(X12,X5)|r2_hidden(esk2_2(X5,X6),X12)))|X6=k1_setfam_1(X5))|X5=k1_xboole_0)))&(((X6!=k1_setfam_1(X5)|X6=k1_xboole_0)|X5!=k1_xboole_0)&((X6!=k1_xboole_0|X6=k1_setfam_1(X5))|X5!=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_setfam_1])])])])])])])])).
fof(c_0_6, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk4_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk4_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_7, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk5_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_8,plain,(X1=k1_xboole_0|r2_hidden(X3,X4)|X2!=k1_setfam_1(X1)|~r2_hidden(X4,X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(v1_xboole_0(X1)|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X2]:![X2]:((r1_xboole_0(X2,X2)|X2!=k1_xboole_0)&(X2=k1_xboole_0|~r1_xboole_0(X2,X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])])])).
cnf(c_0_12,plain,(X1=k1_xboole_0|r2_hidden(X2,X3)|~r2_hidden(X2,k1_setfam_1(X1))|~r2_hidden(X3,X1)), inference(er,[status(thm)],[c_0_8])).
fof(c_0_13, negated_conjecture, (~(![X1]:(r2_hidden(k1_xboole_0,X1)=>k1_setfam_1(X1)=k1_xboole_0))), inference(assume_negation,[status(cth)],[t5_setfam_1])).
cnf(c_0_14,plain,(v1_xboole_0(X1)|~r2_hidden(esk5_1(X1),X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_15,plain,(r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,plain,(X1=k1_xboole_0|v1_xboole_0(k1_setfam_1(X1))|r2_hidden(esk5_1(k1_setfam_1(X1)),X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_12, c_0_10])).
fof(c_0_18, negated_conjecture, ((r2_hidden(k1_xboole_0,esk6_0)&k1_setfam_1(esk6_0)!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_19,plain,(v1_xboole_0(X1)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_14, c_0_10])).
cnf(c_0_20,plain,(r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22,plain,(X1=k1_xboole_0|v1_xboole_0(k1_setfam_1(X1))|~v1_xboole_0(X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,negated_conjecture,(r2_hidden(k1_xboole_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25,plain,(r1_xboole_0(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_16, c_0_21])).
cnf(c_0_26,negated_conjecture,(esk6_0=k1_xboole_0|v1_xboole_0(k1_setfam_1(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])])).
cnf(c_0_27,plain,(X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28,negated_conjecture,(esk6_0=k1_xboole_0|r1_xboole_0(k1_setfam_1(esk6_0),X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29,negated_conjecture,(k1_setfam_1(esk6_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30,negated_conjecture,(~v1_xboole_0(esk6_0)), inference(spm,[status(thm)],[c_0_16, c_0_23])).
cnf(c_0_31,negated_conjecture,(esk6_0=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_32,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_24])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 22
# Proof object formula steps           : 11
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 55
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 55
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 52
# ...of the previous two non-trivial   : 42
# Contextual simplify-reflections      : 0
# Paramodulations                      : 46
# Factorizations                       : 0
# Equation resolutions                 : 8
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 59
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 127
# Rec. Clause-clause subsumption calls : 99
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 28
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1792

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 42952 pages
