# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1))))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', t5_waybel_0)).
fof(reflexivity_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', reflexivity_r3_orders_2)).
fof(d2_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X5,X3)&r1_orders_2(X1,X5,X4)))))))))))))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', d2_waybel_0)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', t4_subset)).
fof(redefinition_r3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3)))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', redefinition_r3_orders_2)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', d1_tarski)).
fof(redefinition_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', redefinition_k6_domain_1)).
fof(d1_waybel_0, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X3,X5)&r1_orders_2(X1,X4,X5)))))))))))))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', d1_waybel_0)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', fc2_struct_0)).
fof(dt_k6_domain_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1)))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', dt_k6_domain_1)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t5_waybel_0__94_premises__rerun__rerun.E_input', dt_l1_orders_2)).
fof(c_0_11, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)))))), inference(assume_negation,[status(cth)],[t5_waybel_0])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))|r3_orders_2(X4,X5,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])])])).
fof(c_0_13, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])).
fof(c_0_14, plain, (![X6]:![X7]:![X8]:![X9]:![X13]:(((((((((m1_subset_1(esk3_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((((r2_hidden(esk3_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((((r1_orders_2(X6,esk3_4(X6,X7,X8,X9),X8)|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((((r1_orders_2(X6,esk3_4(X6,X7,X8,X9),X9)|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))&((((m1_subset_1(esk4_2(X6,X7),u1_struct_0(X6))|v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((m1_subset_1(esk5_2(X6,X7),u1_struct_0(X6))|v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((r2_hidden(esk4_2(X6,X7),X7)|v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((r2_hidden(esk5_2(X6,X7),X7)|v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((~m1_subset_1(X13,u1_struct_0(X6))|(~r2_hidden(X13,X7)|(~r1_orders_2(X6,X13,esk4_2(X6,X7))|~r1_orders_2(X6,X13,esk5_2(X6,X7)))))|v2_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_waybel_0])])])])])])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(((~r3_orders_2(X4,X5,X6)|r1_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4)))))&((~r1_orders_2(X4,X5,X6)|r3_orders_2(X4,X5,X6))|((v2_struct_0(X4)|(~v3_orders_2(X4)|~l1_orders_2(X4)))|(~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_17,plain,(r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk6_2(X4,X5),X5)|esk6_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk6_2(X4,X5),X5)|esk6_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_23,plain,(v2_waybel_0(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,X3,esk5_2(X1,X2))|~r1_orders_2(X1,X3,esk4_2(X1,X2))|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(v2_struct_0(X2)|r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~r3_orders_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26,negated_conjecture,(r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_27,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|k6_domain_1(X3,X4)=k1_tarski(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_29, plain, (![X6]:![X7]:![X8]:![X9]:![X13]:(((((((((m1_subset_1(esk7_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((((r2_hidden(esk7_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((((((r1_orders_2(X6,X8,esk7_4(X6,X7,X8,X9))|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&(((((((r1_orders_2(X6,X9,esk7_4(X6,X7,X8,X9))|~r2_hidden(X9,X7))|~r2_hidden(X8,X7))|~m1_subset_1(X9,u1_struct_0(X6)))|~m1_subset_1(X8,u1_struct_0(X6)))|~v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))&((((m1_subset_1(esk8_2(X6,X7),u1_struct_0(X6))|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((m1_subset_1(esk9_2(X6,X7),u1_struct_0(X6))|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((r2_hidden(esk8_2(X6,X7),X7)|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((r2_hidden(esk9_2(X6,X7),X7)|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6))&((((~m1_subset_1(X13,u1_struct_0(X6))|(~r2_hidden(X13,X7)|(~r1_orders_2(X6,esk8_2(X6,X7),X13)|~r1_orders_2(X6,esk9_2(X6,X7),X13))))|v1_waybel_0(X7,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel_0])])])])])])])).
cnf(c_0_30,plain,(v2_waybel_0(X1,X2)|~r1_orders_2(X2,X3,esk4_2(X2,X1))|~r1_orders_2(X2,X3,esk5_2(X2,X1))|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31,negated_conjecture,(r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_32,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_33,plain,(v2_waybel_0(X2,X1)|r2_hidden(esk5_2(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_34,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_35, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_36,plain,(k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37,plain,(v1_waybel_0(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,esk9_2(X1,X2),X3)|~r1_orders_2(X1,esk8_2(X1,X2),X3)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38,plain,(v1_waybel_0(X2,X1)|r2_hidden(esk9_2(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39,negated_conjecture,(v2_waybel_0(X1,esk1_0)|~r1_orders_2(esk1_0,esk5_2(esk1_0,X1),esk4_2(esk1_0,X1))|~r2_hidden(esk5_2(esk1_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_19])]), c_0_24])).
cnf(c_0_40,plain,(esk5_2(X1,k1_tarski(X2))=X2|v2_waybel_0(k1_tarski(X2),X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_34])])).
cnf(c_0_42,plain,(v2_waybel_0(X2,X1)|r2_hidden(esk4_2(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_43, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~m1_subset_1(X4,X3))|m1_subset_1(k6_domain_1(X3,X4),k1_zfmisc_1(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_44,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45,negated_conjecture,(k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k1_tarski(esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_36, c_0_18])).
cnf(c_0_46,plain,(v1_waybel_0(X1,X2)|~r1_orders_2(X2,esk8_2(X2,X1),X3)|~r1_orders_2(X2,esk9_2(X2,X1),X3)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_37, c_0_24])).
cnf(c_0_47,plain,(esk9_2(X1,k1_tarski(X2))=X2|v1_waybel_0(k1_tarski(X2),X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_38])).
cnf(c_0_48,plain,(v1_waybel_0(X2,X1)|r2_hidden(esk8_2(X1,X2),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_49,negated_conjecture,(v2_waybel_0(k1_tarski(X1),esk1_0)|~r1_orders_2(esk1_0,X1,esk4_2(esk1_0,k1_tarski(X1)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_19])])).
cnf(c_0_50,plain,(esk4_2(X1,k1_tarski(X2))=X2|v2_waybel_0(k1_tarski(X2),X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_42])).
cnf(c_0_51,plain,(m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52,negated_conjecture,(k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k1_tarski(esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_21])).
cnf(c_0_53,plain,(v1_waybel_0(k1_tarski(X1),X2)|~r1_orders_2(X2,esk8_2(X2,k1_tarski(X1)),X3)|~r1_orders_2(X2,X1,X3)|~r2_hidden(X3,k1_tarski(X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54,plain,(esk8_2(X1,k1_tarski(X2))=X2|v1_waybel_0(k1_tarski(X2),X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_48])).
cnf(c_0_55,negated_conjecture,(~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_56,negated_conjecture,(v2_waybel_0(k1_tarski(X1),esk1_0)|~r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_19])])).
cnf(c_0_57,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_18])])).
cnf(c_0_58,plain,(v1_waybel_0(k1_tarski(X1),X2)|~r1_orders_2(X2,X1,X3)|~r2_hidden(X3,k1_tarski(X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59,negated_conjecture,(~l1_struct_0(esk1_0)|~v2_waybel_0(k1_tarski(esk2_0),esk1_0)|~v1_waybel_0(k1_tarski(esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_52])).
cnf(c_0_60,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|v2_waybel_0(k1_tarski(esk2_0),esk1_0)|~l1_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|v1_waybel_0(k1_tarski(esk2_0),esk1_0)|~l1_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r2_hidden(X1,k1_tarski(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_57]), c_0_19])])).
cnf(c_0_62,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~v1_waybel_0(k1_tarski(esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_63,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_41]), c_0_62])).
fof(c_0_64, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_65,negated_conjecture,(~l1_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_63]), c_0_21])).
cnf(c_0_66,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_67,negated_conjecture,(~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_19])])).
cnf(c_0_68,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_31]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 46
# Proof object formula steps           : 23
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 114
# ...of these trivial                  : 0
# ...subsumed                          : 12
# ...remaining for further processing  : 102
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 117
# ...of the previous two non-trivial   : 93
# Contextual simplify-reflections      : 48
# Paramodulations                      : 109
# Factorizations                       : 2
# Equation resolutions                 : 7
# Current number of processed clauses  : 61
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 55
# Current number of unprocessed clauses: 47
# ...number of literals in the above   : 305
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 1951
# Rec. Clause-clause subsumption calls : 342
# Non-unit clause-clause subsumptions  : 64
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6872

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 43152 pages
