# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k3_orders_2(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', dt_k3_orders_2)).
fof(t4_subset_1, axiom, (![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', t4_subset_1)).
fof(t60_orders_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k3_orders_2(X1,k1_struct_0(X1),X2)=k1_xboole_0))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', t60_orders_2)).
fof(redefinition_m2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1))))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2)))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, (![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2)))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', cc1_subset_1)).
fof(existence_m2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1))))=>?[X3]:m2_subset_1(X3,X1,X2))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', existence_m2_subset_1)).
fof(dt_m2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1))))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1)))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', dt_m2_subset_1)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', t2_subset)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', d3_tarski)).
fof(t57_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,k3_orders_2(X1,X4,X3))<=>(r2_orders_2(X1,X2,X3)&r2_hidden(X2,X4)))))))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', t57_orders_2)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', dt_l1_orders_2)).
fof(d2_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k1_struct_0(X1)=k1_xboole_0)), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', d2_struct_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', d10_xboole_0)).
fof(fc1_xboole_0, axiom, (v1_xboole_0(k1_xboole_0)), file('out/t60_orders_2__16_premises__rerun__rerun.E_input', fc1_xboole_0)).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|~l1_orders_2(X4)))))|(~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k3_orders_2(X4,X5,X6),k1_zfmisc_1(u1_struct_0(X4))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_orders_2])])])).
fof(c_0_16, plain, (![X2]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_17, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k3_orders_2(X1,k1_struct_0(X1),X2)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t60_orders_2])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:![X6]:(((~m2_subset_1(X6,X4,X5)|m1_subset_1(X6,X5))|(v1_xboole_0(X4)|(v1_xboole_0(X5)|~m1_subset_1(X5,k1_zfmisc_1(X4)))))&((~m1_subset_1(X6,X5)|m2_subset_1(X6,X4,X5))|(v1_xboole_0(X4)|(v1_xboole_0(X5)|~m1_subset_1(X5,k1_zfmisc_1(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])])])).
fof(c_0_19, plain, (![X3]:![X4]:(~v1_xboole_0(X3)|(~m1_subset_1(X4,k1_zfmisc_1(X3))|v1_xboole_0(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])])])).
cnf(c_0_20,plain,(m1_subset_1(k3_orders_2(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&k3_orders_2(esk1_0,k1_struct_0(esk1_0),esk2_0)!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
fof(c_0_23, plain, (![X4]:![X5]:((v1_xboole_0(X4)|(v1_xboole_0(X5)|~m1_subset_1(X5,k1_zfmisc_1(X4))))|m2_subset_1(esk3_2(X4,X5),X4,X5))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[existence_m2_subset_1])])])])])])).
cnf(c_0_24,plain,(v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m2_subset_1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,plain,(v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(m1_subset_1(k3_orders_2(X1,k1_xboole_0,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33,plain,(m2_subset_1(esk3_2(X1,X2),X1,X2)|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:((v1_xboole_0(X4)|(v1_xboole_0(X5)|~m1_subset_1(X5,k1_zfmisc_1(X4))))|(~m2_subset_1(X6,X4,X5)|m1_subset_1(X6,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])])])).
cnf(c_0_35,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(k3_orders_2(esk1_0,k1_xboole_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_37,plain,(m2_subset_1(esk3_2(X1,X2),X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_33, c_0_25])).
cnf(c_0_38,plain,(m1_subset_1(X1,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_39, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_40,negated_conjecture,(v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))|m1_subset_1(X1,k3_orders_2(esk1_0,k1_xboole_0,esk2_0))|~m2_subset_1(X1,u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41,negated_conjecture,(m2_subset_1(esk3_2(u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0)),u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0))|v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))), inference(spm,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_42,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_38, c_0_25])).
fof(c_0_43, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk4_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
fof(c_0_44, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_45, plain, (![X5]:![X6]:![X7]:![X8]:(((((((r2_orders_2(X5,X6,X7)|~r2_hidden(X6,k3_orders_2(X5,X8,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(((((r2_hidden(X6,X8)|~r2_hidden(X6,k3_orders_2(X5,X8,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))&((((((~r2_orders_2(X5,X6,X7)|~r2_hidden(X6,X8))|r2_hidden(X6,k3_orders_2(X5,X8,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_orders_2])])])])])])])).
cnf(c_0_46,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47,negated_conjecture,(v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))|m1_subset_1(esk3_2(u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0)),k3_orders_2(esk1_0,k1_xboole_0,esk2_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48,negated_conjecture,(v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))|m1_subset_1(X1,u1_struct_0(esk1_0))|~m2_subset_1(X1,u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_36])).
fof(c_0_49, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_50,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52,plain,(v2_struct_0(X1)|r2_hidden(X2,X4)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,k3_orders_2(X1,X4,X3))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk3_2(u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0)),k3_orders_2(esk1_0,k1_xboole_0,esk2_0))|v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54,negated_conjecture,(v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))|m1_subset_1(esk3_2(u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_48, c_0_41])).
fof(c_0_55, plain, (![X2]:(~l1_struct_0(X2)|k1_struct_0(X2)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_struct_0])])).
cnf(c_0_56,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_57, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_58,plain,(r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_59,plain,(v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_60,negated_conjecture,(r2_hidden(esk3_2(u1_struct_0(esk1_0),k3_orders_2(esk1_0,k1_xboole_0,esk2_0)),k1_xboole_0)|v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_21]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31])]), c_0_32]), c_0_54])).
cnf(c_0_61,plain,(k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62,negated_conjecture,(l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_56, c_0_28])).
cnf(c_0_63,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65,negated_conjecture,(v1_xboole_0(k3_orders_2(esk1_0,k1_xboole_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_60]), c_0_59])])).
cnf(c_0_66,negated_conjecture,(k3_orders_2(esk1_0,k1_struct_0(esk1_0),esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_67,negated_conjecture,(k1_struct_0(esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68,plain,(X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_69,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,k1_xboole_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_58, c_0_65])).
cnf(c_0_70,negated_conjecture,(k3_orders_2(esk1_0,k1_xboole_0,esk2_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 42
# Proof object formula steps           : 30
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 109
# ...of these trivial                  : 1
# ...subsumed                          : 1
# ...remaining for further processing  : 107
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 11
# Generated clauses                    : 119
# ...of the previous two non-trivial   : 101
# Contextual simplify-reflections      : 6
# Paramodulations                      : 117
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 66
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 44
# Current number of unprocessed clauses: 38
# ...number of literals in the above   : 125
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 805
# Rec. Clause-clause subsumption calls : 253
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 31
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 41
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6541

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 43172 pages
