# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C45_F1_PI_AE_Q7_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t57_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_xtuple_0(k5_relat_1(X3,X2)))<=>(r2_hidden(X1,X2)&r2_hidden(X1,k9_xtuple_0(X3)))))), file('out/t60_relat_1__2_premises__rerun.E_input', t57_relat_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t60_relat_1__2_premises__rerun.E_input', d3_tarski)).
fof(t60_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>r1_tarski(k9_xtuple_0(k5_relat_1(X2,X1)),k9_xtuple_0(X2)))), file('out/t60_relat_1__2_premises__rerun.E_input', t60_relat_1)).
fof(c_0_3, plain, (![X4]:![X5]:![X6]:((((r2_hidden(X4,X5)|~r2_hidden(X4,k9_xtuple_0(k5_relat_1(X6,X5))))|~v1_relat_1(X6))&((r2_hidden(X4,k9_xtuple_0(X6))|~r2_hidden(X4,k9_xtuple_0(k5_relat_1(X6,X5))))|~v1_relat_1(X6)))&(((~r2_hidden(X4,X5)|~r2_hidden(X4,k9_xtuple_0(X6)))|r2_hidden(X4,k9_xtuple_0(k5_relat_1(X6,X5))))|~v1_relat_1(X6)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t57_relat_1])])])).
fof(c_0_4, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk1_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk1_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_5, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>r1_tarski(k9_xtuple_0(k5_relat_1(X2,X1)),k9_xtuple_0(X2))))), inference(assume_negation,[status(cth)],[t60_relat_1])).
cnf(c_0_6,plain,(r2_hidden(X2,k9_xtuple_0(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k9_xtuple_0(k5_relat_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_7,plain,(r1_tarski(X1,X2)|r2_hidden(esk1_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_8, negated_conjecture, ((v1_relat_1(esk3_0)&~r1_tarski(k9_xtuple_0(k5_relat_1(esk3_0,esk2_0)),k9_xtuple_0(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_9,plain,(r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,plain,(r1_tarski(k9_xtuple_0(k5_relat_1(X1,X2)),X3)|r2_hidden(esk1_2(k9_xtuple_0(k5_relat_1(X1,X2)),X3),k9_xtuple_0(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11,negated_conjecture,(~r1_tarski(k9_xtuple_0(k5_relat_1(esk3_0,esk2_0)),k9_xtuple_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(r1_tarski(k9_xtuple_0(k5_relat_1(X1,X2)),k9_xtuple_0(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 15
# Proof object clause steps            : 8
# Proof object formula steps           : 7
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 8
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 8
# Processed clauses                    : 22
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 21
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 16
# ...of the previous two non-trivial   : 12
# Contextual simplify-reflections      : 0
# Paramodulations                      : 16
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 8
# Clause-clause subsumption calls (NU) : 22
# Rec. Clause-clause subsumption calls : 16
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 844

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 43044 pages
