# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t62_relat_1__8_premises__rerun.E_input', t3_xboole_0)).
fof(t7_tarski, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3)))))))), file('out/t62_relat_1__8_premises__rerun.E_input', t7_tarski)).
fof(d8_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4]:![X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2)))))))), file('out/t62_relat_1__8_premises__rerun.E_input', d8_relat_1)).
fof(t56_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>(![X2]:![X3]:~(r2_hidden(k4_tarski(X2,X3),X1))=>X1=k1_xboole_0))), file('out/t62_relat_1__8_premises__rerun.E_input', t56_relat_1)).
fof(fc6_relat_1, axiom, (![X1]:![X2]:v1_relat_1(k2_zfmisc_1(X1,X2))), file('out/t62_relat_1__8_premises__rerun.E_input', fc6_relat_1)).
fof(t113_zfmisc_1, axiom, (![X1]:![X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), file('out/t62_relat_1__8_premises__rerun.E_input', t113_zfmisc_1)).
fof(dt_k5_relat_1, axiom, (![X1]:![X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2)))), file('out/t62_relat_1__8_premises__rerun.E_input', dt_k5_relat_1)).
fof(t65_xboole_1, axiom, (![X1]:r1_xboole_0(X1,k1_xboole_0)), file('out/t62_relat_1__8_premises__rerun.E_input', t65_xboole_1)).
fof(t62_relat_1, conjecture, (![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0))), file('out/t62_relat_1__8_premises__rerun.E_input', t62_relat_1)).
fof(c_0_9, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk1_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk1_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_10, plain, (![X5]:![X6]:![X8]:((r2_hidden(esk2_2(X5,X6),X6)|~r2_hidden(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,esk2_2(X5,X6)))|~r2_hidden(X5,X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])])])).
cnf(c_0_11,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(r2_hidden(esk2_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13,plain,(~r2_hidden(esk2_2(X1,X2),X3)|~r2_hidden(X1,X2)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X10]:![X11]:![X13]:![X16]:((((((((r2_hidden(k4_tarski(X10,esk3_5(X7,X8,X9,X10,X11)),X7)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7))&(((((r2_hidden(k4_tarski(esk3_5(X7,X8,X9,X10,X11),X11),X8)|~r2_hidden(k4_tarski(X10,X11),X9))|X9!=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7)))&((((((~r2_hidden(k4_tarski(X10,X13),X7)|~r2_hidden(k4_tarski(X13,X11),X8))|r2_hidden(k4_tarski(X10,X11),X9))|X9!=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7)))&((((((~r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk5_3(X7,X8,X9)),X9)|(~r2_hidden(k4_tarski(esk4_3(X7,X8,X9),X16),X7)|~r2_hidden(k4_tarski(X16,esk5_3(X7,X8,X9)),X8)))|X9=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7))&((((((r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk6_3(X7,X8,X9)),X7)|r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk5_3(X7,X8,X9)),X9))|X9=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7))&(((((r2_hidden(k4_tarski(esk6_3(X7,X8,X9),esk5_3(X7,X8,X9)),X8)|r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk5_3(X7,X8,X9)),X9))|X9=k5_relat_1(X7,X8))|~v1_relat_1(X9))|~v1_relat_1(X8))|~v1_relat_1(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])])).
cnf(c_0_15,plain,(~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_13, c_0_12])).
cnf(c_0_16,plain,(r2_hidden(k4_tarski(esk3_5(X1,X2,X3,X4,X5),X5),X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)|X3!=k5_relat_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:(~v1_relat_1(X4)|(r2_hidden(k4_tarski(esk7_1(X4),esk8_1(X4)),X4)|X4=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t56_relat_1])])])])])])).
fof(c_0_18, plain, (![X3]:![X4]:v1_relat_1(k2_zfmisc_1(X3,X4))), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_19, plain, (![X3]:![X4]:![X3]:![X4]:((k2_zfmisc_1(X3,X4)!=k1_xboole_0|(X3=k1_xboole_0|X4=k1_xboole_0))&((X3!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)&(X4!=k1_xboole_0|k2_zfmisc_1(X3,X4)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])])])).
cnf(c_0_20,plain,(X1!=k5_relat_1(X2,X3)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X4,X5),X1)|~r1_xboole_0(X3,X3)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21,plain,(X1=k1_xboole_0|r2_hidden(k4_tarski(esk7_1(X1),esk8_1(X1)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, (![X3]:![X4]:((~v1_relat_1(X3)|~v1_relat_1(X4))|v1_relat_1(k5_relat_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_23,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_24, plain, (![X2]:r1_xboole_0(X2,k1_xboole_0)), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_25,plain,(v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(r2_hidden(k4_tarski(X4,esk3_5(X1,X2,X3,X4,X5)),X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)|X3!=k5_relat_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_28, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t62_relat_1])).
cnf(c_0_29,plain,(X1=k1_xboole_0|X1!=k5_relat_1(X2,X3)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r1_xboole_0(X3,X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30,plain,(v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_15, c_0_23])).
cnf(c_0_32,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,plain,(v1_relat_1(k1_xboole_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34,plain,(X1!=k5_relat_1(X2,X3)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X4,X5),X1)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_15, c_0_27])).
fof(c_0_35, negated_conjecture, ((v1_relat_1(esk9_0)&(k5_relat_1(k1_xboole_0,esk9_0)!=k1_xboole_0|k5_relat_1(esk9_0,k1_xboole_0)!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])).
cnf(c_0_36,plain,(k5_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xboole_0(X2,X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_29]), c_0_30])).
cnf(c_0_37,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38,plain,(v1_relat_1(k1_xboole_0)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_39,plain,(X1=k1_xboole_0|X1!=k5_relat_1(X2,X3)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_34, c_0_21])).
cnf(c_0_40,negated_conjecture,(k5_relat_1(esk9_0,k1_xboole_0)!=k1_xboole_0|k5_relat_1(k1_xboole_0,esk9_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41,plain,(k5_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
cnf(c_0_42,negated_conjecture,(v1_relat_1(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43,plain,(k5_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xboole_0(X1,X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_39]), c_0_30])).
cnf(c_0_44,negated_conjecture,(k5_relat_1(k1_xboole_0,esk9_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
cnf(c_0_45,plain,(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_37]), c_0_38])])).
cnf(c_0_46,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_42])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 28
# Proof object formula steps           : 19
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 68
# ...of these trivial                  : 1
# ...subsumed                          : 6
# ...remaining for further processing  : 61
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 58
# ...of the previous two non-trivial   : 43
# Contextual simplify-reflections      : 2
# Paramodulations                      : 55
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 40
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 99
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 453
# Rec. Clause-clause subsumption calls : 179
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2666

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 42940 pages
