# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_pre_topc, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3]:![X4]:(g1_pre_topc(X1,X2)=g1_pre_topc(X3,X4)=>(X1=X3&X2=X4)))), file('out/t63_tex_2__20_premises__rerun.E_input', free_g1_pre_topc)).
fof(dt_u1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t63_tex_2__20_premises__rerun.E_input', dt_u1_pre_topc)).
fof(abstractness_v1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))), file('out/t63_tex_2__20_premises__rerun.E_input', abstractness_v1_pre_topc)).
fof(d5_compts_1, axiom, (![X1]:k1_compts_1(X1)=g1_pre_topc(X1,k2_subset_1(k9_setfam_1(X1)))), file('out/t63_tex_2__20_premises__rerun.E_input', d5_compts_1)).
fof(d3_subset_1, axiom, (![X1]:k2_subset_1(X1)=X1), file('out/t63_tex_2__20_premises__rerun.E_input', d3_subset_1)).
fof(redefinition_k9_setfam_1, axiom, (![X1]:k9_setfam_1(X1)=k1_zfmisc_1(X1)), file('out/t63_tex_2__20_premises__rerun.E_input', redefinition_k9_setfam_1)).
fof(t63_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2)))=>v1_tdlat_3(X1)))), file('out/t63_tex_2__20_premises__rerun.E_input', t63_tex_2)).
fof(fc1_compts_1, axiom, (![X1]:(v1_pre_topc(k1_compts_1(X1))&v2_pre_topc(k1_compts_1(X1)))), file('out/t63_tex_2__20_premises__rerun.E_input', fc1_compts_1)).
fof(dt_k1_compts_1, axiom, (![X1]:l1_pre_topc(k1_compts_1(X1))), file('out/t63_tex_2__20_premises__rerun.E_input', dt_k1_compts_1)).
fof(cc1_funct_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_partfun1(X3,X1)=>v1_funct_2(X3,X1,X2)))), file('out/t63_tex_2__20_premises__rerun.E_input', cc1_funct_2)).
fof(dt_k6_partfun1, axiom, (![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))), file('out/t63_tex_2__20_premises__rerun.E_input', dt_k6_partfun1)).
fof(fc3_funct_1, axiom, (![X1]:(v1_relat_1(k4_relat_1(X1))&v1_funct_1(k4_relat_1(X1)))), file('out/t63_tex_2__20_premises__rerun.E_input', fc3_funct_1)).
fof(redefinition_k6_partfun1, axiom, (![X1]:k6_partfun1(X1)=k4_relat_1(X1)), file('out/t63_tex_2__20_premises__rerun.E_input', redefinition_k6_partfun1)).
fof(d6_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v4_pre_topc(X4,X2)=>v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('out/t63_tex_2__20_premises__rerun.E_input', d6_pre_topc)).
fof(t16_tdlat_3, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v4_pre_topc(X2,X1))))), file('out/t63_tex_2__20_premises__rerun.E_input', t16_tdlat_3)).
fof(cc1_tdlat_3, axiom, (![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1)))), file('out/t63_tex_2__20_premises__rerun.E_input', cc1_tdlat_3)).
fof(fc4_tex_1, axiom, (![X1]:v1_tdlat_3(k1_compts_1(X1))), file('out/t63_tex_2__20_premises__rerun.E_input', fc4_tex_1)).
fof(t94_funct_2, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k8_relset_1(X1,X1,k6_partfun1(X1),X2)=X2)), file('out/t63_tex_2__20_premises__rerun.E_input', t94_funct_2)).
fof(fc2_compts_1, axiom, (![X1]:(~(v1_xboole_0(X1))=>~(v2_struct_0(k1_compts_1(X1))))), file('out/t63_tex_2__20_premises__rerun.E_input', fc2_compts_1)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t63_tex_2__20_premises__rerun.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t63_tex_2__20_premises__rerun.E_input', dt_l1_pre_topc)).
fof(c_0_21, plain, (![X5]:![X6]:![X7]:![X8]:(((X5=X7|g1_pre_topc(X5,X6)!=g1_pre_topc(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k1_zfmisc_1(X5))))&((X6=X8|g1_pre_topc(X5,X6)!=g1_pre_topc(X7,X8))|~m1_subset_1(X6,k1_zfmisc_1(k1_zfmisc_1(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_pre_topc])])])])])])).
fof(c_0_22, plain, (![X2]:(~l1_pre_topc(X2)|m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_23,plain,(X1=X4|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|g1_pre_topc(X2,X1)!=g1_pre_topc(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_24,plain,(m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, (![X2]:(~l1_pre_topc(X2)|(~v1_pre_topc(X2)|X2=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_26,plain,(u1_pre_topc(X1)=X2|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=g1_pre_topc(X3,X2)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27,plain,(X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~v1_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_28,plain,(X2=X3|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|g1_pre_topc(X2,X1)!=g1_pre_topc(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,plain,(u1_pre_topc(g1_pre_topc(X1,X2))=X2|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27])])).
cnf(c_0_30,plain,(u1_struct_0(X1)=X2|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=g1_pre_topc(X2,X3)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_28, c_0_24])).
fof(c_0_31, plain, (![X2]:k1_compts_1(X2)=g1_pre_topc(X2,k2_subset_1(k9_setfam_1(X2)))), inference(variable_rename,[status(thm)],[d5_compts_1])).
fof(c_0_32, plain, (![X2]:k2_subset_1(X2)=X2), inference(variable_rename,[status(thm)],[d3_subset_1])).
fof(c_0_33, plain, (![X2]:k9_setfam_1(X2)=k1_zfmisc_1(X2)), inference(variable_rename,[status(thm)],[redefinition_k9_setfam_1])).
fof(c_0_34, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2)))=>v1_tdlat_3(X1))))), inference(assume_negation,[status(cth)],[t63_tex_2])).
cnf(c_0_35,plain,(m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(g1_pre_topc(X2,X1)))))|~v1_pre_topc(g1_pre_topc(X2,X1))|~l1_pre_topc(g1_pre_topc(X2,X1))), inference(spm,[status(thm)],[c_0_24, c_0_29])).
cnf(c_0_36,plain,(u1_struct_0(g1_pre_topc(X1,X2))=X1|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_27])])).
fof(c_0_37, plain, (![X2]:![X2]:(v1_pre_topc(k1_compts_1(X2))&v2_pre_topc(k1_compts_1(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc1_compts_1])])])).
cnf(c_0_38,plain,(k1_compts_1(X1)=g1_pre_topc(X1,k2_subset_1(k9_setfam_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39,plain,(k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40,plain,(k9_setfam_1(X1)=k1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_41, plain, (![X2]:l1_pre_topc(k1_compts_1(X2))), inference(variable_rename,[status(thm)],[dt_k1_compts_1])).
fof(c_0_42, negated_conjecture, (![X5]:![X6]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(((v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))|((~v1_funct_1(X6)|(~v1_funct_2(X6,u1_struct_0(esk1_0),u1_struct_0(X5))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X5))))))|v5_pre_topc(X6,esk1_0,X5)))&~v1_tdlat_3(esk1_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_34])])])])])])])).
fof(c_0_43, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|(~v1_partfun1(X6,X4)|v1_funct_2(X6,X4,X5)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])).
fof(c_0_44, plain, (![X2]:![X2]:(v1_partfun1(k6_partfun1(X2),X2)&m1_subset_1(k6_partfun1(X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[dt_k6_partfun1])])])).
fof(c_0_45, plain, (![X2]:![X2]:(v1_relat_1(k4_relat_1(X2))&v1_funct_1(k4_relat_1(X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[fc3_funct_1])])])).
fof(c_0_46, plain, (![X2]:k6_partfun1(X2)=k4_relat_1(X2)), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
cnf(c_0_47,plain,(m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v1_pre_topc(g1_pre_topc(X2,X1))|~l1_pre_topc(g1_pre_topc(X2,X1))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_48,plain,(v1_pre_topc(k1_compts_1(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49,plain,(k1_compts_1(X1)=g1_pre_topc(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_50,plain,(l1_pre_topc(k1_compts_1(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_51, plain, (![X5]:![X6]:![X7]:![X8]:(((((~v5_pre_topc(X7,X5,X6)|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|(~v4_pre_topc(X8,X6)|v4_pre_topc(k8_relset_1(u1_struct_0(X5),u1_struct_0(X6),X7,X8),X5))))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&(((((m1_subset_1(esk2_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(X6)))|v5_pre_topc(X7,X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&(((((v4_pre_topc(esk2_3(X5,X6,X7),X6)|v5_pre_topc(X7,X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|~l1_pre_topc(X6))|~l1_pre_topc(X5))&((((~v4_pre_topc(k8_relset_1(u1_struct_0(X5),u1_struct_0(X6),X7,esk2_3(X5,X6,X7)),X5)|v5_pre_topc(X7,X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|~l1_pre_topc(X6))|~l1_pre_topc(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_pre_topc])])])])])])])).
cnf(c_0_52,negated_conjecture,(v5_pre_topc(X1,esk1_0,X2)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53,plain,(v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54,plain,(m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55,plain,(v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56,plain,(v1_funct_1(k4_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57,plain,(k6_partfun1(X1)=k4_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_58, plain, (![X3]:![X4]:(((~v1_tdlat_3(X3)|(~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|v4_pre_topc(X4,X3)))|(~v2_pre_topc(X3)|~l1_pre_topc(X3)))&(((m1_subset_1(esk3_1(X3),k1_zfmisc_1(u1_struct_0(X3)))|v1_tdlat_3(X3))|(~v2_pre_topc(X3)|~l1_pre_topc(X3)))&((~v4_pre_topc(esk3_1(X3),X3)|v1_tdlat_3(X3))|(~v2_pre_topc(X3)|~l1_pre_topc(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_tdlat_3])])])])])])])).
fof(c_0_59, plain, (![X2]:(~l1_pre_topc(X2)|(~v1_tdlat_3(X2)|v2_pre_topc(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
cnf(c_0_60,plain,(X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~v1_pre_topc(g1_pre_topc(X1,X3))|~l1_pre_topc(g1_pre_topc(X1,X3))), inference(spm,[status(thm)],[c_0_28, c_0_47])).
cnf(c_0_61,plain,(v1_pre_topc(g1_pre_topc(X1,k1_zfmisc_1(X1)))), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_62,plain,(l1_pre_topc(g1_pre_topc(X1,k1_zfmisc_1(X1)))), inference(rw,[status(thm)],[c_0_50, c_0_49])).
fof(c_0_63, plain, (![X2]:v1_tdlat_3(k1_compts_1(X2))), inference(variable_rename,[status(thm)],[fc4_tex_1])).
cnf(c_0_64,plain,(v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v4_pre_topc(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v5_pre_topc(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_65,negated_conjecture,(v5_pre_topc(X1,esk1_0,g1_pre_topc(X2,X3))|v2_struct_0(g1_pre_topc(X2,X3))|~v1_funct_2(X1,u1_struct_0(esk1_0),X2)|~v1_funct_1(X1)|~v2_pre_topc(g1_pre_topc(X2,X3))|~v1_pre_topc(g1_pre_topc(X2,X3))|~l1_pre_topc(g1_pre_topc(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))), inference(spm,[status(thm)],[c_0_52, c_0_36])).
cnf(c_0_66,plain,(v1_funct_2(k6_partfun1(X1),X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
cnf(c_0_67,plain,(v1_funct_1(k6_partfun1(X1))), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68,plain,(v4_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_69,plain,(v2_pre_topc(X1)|~v1_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70,plain,(X1=X2|g1_pre_topc(X1,k1_zfmisc_1(X1))!=g1_pre_topc(X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])])).
cnf(c_0_71,plain,(v1_tdlat_3(k1_compts_1(X1))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_72,plain,(v4_pre_topc(k8_relset_1(u1_struct_0(X1),X2,X3,X4),X1)|~v4_pre_topc(X4,g1_pre_topc(X2,X5))|~v5_pre_topc(X3,X1,g1_pre_topc(X2,X5))|~v1_funct_2(X3,u1_struct_0(X1),X2)|~v1_funct_1(X3)|~v1_pre_topc(g1_pre_topc(X2,X5))|~l1_pre_topc(g1_pre_topc(X2,X5))|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X2)))|~m1_subset_1(X4,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_64, c_0_36])).
cnf(c_0_73,negated_conjecture,(v5_pre_topc(k6_partfun1(u1_struct_0(esk1_0)),esk1_0,g1_pre_topc(u1_struct_0(esk1_0),X1))|v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),X1))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),X1))|~v1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),X1))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_54]), c_0_66]), c_0_67])])).
cnf(c_0_74,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_75,plain,(v4_pre_topc(X1,X2)|~v1_tdlat_3(X2)|~l1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(csr,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76,plain,(u1_struct_0(g1_pre_topc(X1,k1_zfmisc_1(X1)))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_27])]), c_0_61]), c_0_62])])).
cnf(c_0_77,plain,(v1_tdlat_3(g1_pre_topc(X1,k1_zfmisc_1(X1)))), inference(rw,[status(thm)],[c_0_71, c_0_49])).
cnf(c_0_78,plain,(v2_pre_topc(k1_compts_1(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_79,negated_conjecture,(v4_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k6_partfun1(u1_struct_0(esk1_0)),X1),esk1_0)|v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),X2))|~v4_pre_topc(X1,g1_pre_topc(u1_struct_0(esk1_0),X2))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),X2))|~v1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),X2))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_66]), c_0_67]), c_0_74]), c_0_54])])).
cnf(c_0_80,plain,(v4_pre_topc(X1,g1_pre_topc(X2,k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_62])])).
cnf(c_0_81,plain,(v2_pre_topc(g1_pre_topc(X1,k1_zfmisc_1(X1)))), inference(rw,[status(thm)],[c_0_78, c_0_49])).
fof(c_0_82, plain, (![X3]:![X4]:(~m1_subset_1(X4,k1_zfmisc_1(X3))|k8_relset_1(X3,X3,k6_partfun1(X3),X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_funct_2])])).
fof(c_0_83, plain, (![X2]:(v1_xboole_0(X2)|~v2_struct_0(k1_compts_1(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_compts_1])])])).
cnf(c_0_84,negated_conjecture,(v4_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k6_partfun1(u1_struct_0(esk1_0)),X1),esk1_0)|v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81]), c_0_61]), c_0_62])])).
cnf(c_0_85,plain,(k8_relset_1(X1,X1,k6_partfun1(X1),X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_86,plain,(v1_xboole_0(X1)|~v2_struct_0(k1_compts_1(X1))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_87,negated_conjecture,(v4_pre_topc(X1,esk1_0)|v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_88,plain,(v1_tdlat_3(X1)|m1_subset_1(esk3_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_89,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_90,negated_conjecture,(~v1_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_91,plain,(v1_xboole_0(X1)|~v2_struct_0(g1_pre_topc(X1,k1_zfmisc_1(X1)))), inference(rw,[status(thm)],[c_0_86, c_0_49])).
cnf(c_0_92,negated_conjecture,(v4_pre_topc(esk3_1(esk1_0),esk1_0)|v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89]), c_0_74])]), c_0_90])).
fof(c_0_93, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_94,plain,(v1_tdlat_3(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v4_pre_topc(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_95,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|v4_pre_topc(esk3_1(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_96,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_97,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_89]), c_0_74])]), c_0_90])).
cnf(c_0_98,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_99, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_100,negated_conjecture,(~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])).
cnf(c_0_101,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_102,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_74])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 103
# Proof object clause steps            : 60
# Proof object formula steps           : 43
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 30
# Processed clauses                    : 229
# ...of these trivial                  : 3
# ...subsumed                          : 106
# ...remaining for further processing  : 120
# Other redundant clauses eliminated   : 16
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 2
# Generated clauses                    : 318
# ...of the previous two non-trivial   : 267
# Contextual simplify-reflections      : 61
# Paramodulations                      : 292
# Factorizations                       : 0
# Equation resolutions                 : 26
# Current number of processed clauses  : 84
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 64
# Current number of unprocessed clauses: 88
# ...number of literals in the above   : 700
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 3307
# Rec. Clause-clause subsumption calls : 806
# Non-unit clause-clause subsumptions  : 171
# Unit Clause-clause subsumption calls : 33
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10997

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.004 s
# Total time               : 0.028 s
# Maximum resident set size: 44892 pages
