# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), file('out/t63_enumset1__8_premises__rerun.E_input', t44_enumset1)).
fof(t43_enumset1, axiom, (![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3))), file('out/t63_enumset1__8_premises__rerun.E_input', t43_enumset1)).
fof(t47_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5))), file('out/t63_enumset1__8_premises__rerun.E_input', t47_enumset1)).
fof(l62_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6))), file('out/t63_enumset1__8_premises__rerun.E_input', l62_enumset1)).
fof(t50_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), file('out/t63_enumset1__8_premises__rerun.E_input', t50_enumset1)).
fof(t55_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6))), file('out/t63_enumset1__8_premises__rerun.E_input', t55_enumset1)).
fof(t4_xboole_1, axiom, (![X1]:![X2]:![X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t63_enumset1__8_premises__rerun.E_input', t4_xboole_1)).
fof(t63_enumset1, conjecture, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8))), file('out/t63_enumset1__8_premises__rerun.E_input', t63_enumset1)).
fof(l75_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), file('out/t63_enumset1__8_premises__rerun.E_input', l75_enumset1)).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X5),k1_enumset1(X6,X7,X8))), inference(variable_rename,[status(thm)],[t44_enumset1])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:k1_enumset1(X4,X5,X6)=k2_xboole_0(k2_tarski(X4,X5),k1_tarski(X6))), inference(variable_rename,[status(thm)],[t43_enumset1])).
fof(c_0_11, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k3_enumset1(X6,X7,X8,X9,X10)=k2_xboole_0(k1_tarski(X6),k2_enumset1(X7,X8,X9,X10))), inference(variable_rename,[status(thm)],[t47_enumset1])).
cnf(c_0_12,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(k1_enumset1(X1,X2,X3)=k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:k4_enumset1(X7,X8,X9,X10,X11,X12)=k2_xboole_0(k1_enumset1(X7,X8,X9),k1_enumset1(X10,X11,X12))), inference(variable_rename,[status(thm)],[l62_enumset1])).
fof(c_0_15, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:k3_enumset1(X6,X7,X8,X9,X10)=k2_xboole_0(k2_enumset1(X6,X7,X8,X9),k1_tarski(X10))), inference(variable_rename,[status(thm)],[t50_enumset1])).
cnf(c_0_16,plain,(k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_tarski(X2,X3),k1_tarski(X4)))), inference(rw,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_18, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:k4_enumset1(X7,X8,X9,X10,X11,X12)=k2_xboole_0(k3_enumset1(X7,X8,X9,X10,X11),k1_tarski(X12))), inference(variable_rename,[status(thm)],[t55_enumset1])).
cnf(c_0_19,plain,(k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_tarski(X3,X4),k1_tarski(X5))))), inference(rw,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:k2_xboole_0(k2_xboole_0(X4,X5),X6)=k2_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t4_xboole_1])).
fof(c_0_23, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:![X7]:![X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8)))), inference(assume_negation,[status(cth)],[t63_enumset1])).
fof(c_0_24, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:k6_enumset1(X9,X10,X11,X12,X13,X14,X15,X16)=k2_xboole_0(k2_enumset1(X9,X10,X11,X12),k2_enumset1(X13,X14,X15,X16))), inference(variable_rename,[status(thm)],[l75_enumset1])).
cnf(c_0_25,plain,(k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3)),k2_xboole_0(k2_tarski(X4,X5),k1_tarski(X6)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_13]), c_0_13])).
cnf(c_0_27,plain,(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_tarski(X3,X4),k1_tarski(X5))))=k2_xboole_0(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_tarski(X2,X3),k1_tarski(X4))),k1_tarski(X5))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_17]), c_0_21])).
cnf(c_0_28,plain,(k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, negated_conjecture, (k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0)!=k2_xboole_0(k2_tarski(esk1_0,esk2_0),k4_enumset1(esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
cnf(c_0_30,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(k2_xboole_0(k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3)),k2_xboole_0(k2_tarski(X4,X5),k1_tarski(X6)))=k2_xboole_0(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_tarski(X3,X4),k1_tarski(X5)))),k1_tarski(X6))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_21]), c_0_26])).
cnf(c_0_32,plain,(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_tarski(X2,X3),k2_xboole_0(k1_tarski(X4),k1_tarski(X5))))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_tarski(X3,X4),k1_tarski(X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_28])).
cnf(c_0_33,negated_conjecture,(k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0)!=k2_xboole_0(k2_tarski(esk1_0,esk2_0),k4_enumset1(esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34,plain,(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_tarski(X2,X3),k1_tarski(X4))),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k2_tarski(X6,X7),k1_tarski(X8))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_17]), c_0_17])).
cnf(c_0_35,plain,(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k2_tarski(X4,X5),k1_tarski(X6)))))=k2_xboole_0(k2_tarski(X1,X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k2_tarski(X4,X5),k1_tarski(X6))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_28]), c_0_28]), c_0_28]), c_0_32]), c_0_28])).
cnf(c_0_36,negated_conjecture,(k2_xboole_0(k2_xboole_0(k1_tarski(esk1_0),k2_xboole_0(k2_tarski(esk2_0,esk3_0),k1_tarski(esk4_0))),k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k2_tarski(esk6_0,esk7_0),k1_tarski(esk8_0))))!=k2_xboole_0(k2_tarski(esk1_0,esk2_0),k2_xboole_0(k2_xboole_0(k2_tarski(esk3_0,esk4_0),k1_tarski(esk5_0)),k2_xboole_0(k2_tarski(esk6_0,esk7_0),k1_tarski(esk8_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_26]), c_0_34])).
cnf(c_0_37,plain,(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k2_tarski(X4,X5),k2_xboole_0(k1_tarski(X6),X7)))))=k2_xboole_0(k2_tarski(X1,X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k2_tarski(X4,X5),k2_xboole_0(k1_tarski(X6),X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_35]), c_0_28]), c_0_28]), c_0_28]), c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_38,plain,(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_tarski(X2,X3),k2_xboole_0(k1_tarski(X4),k2_xboole_0(k1_tarski(X5),X6))))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_tarski(X3,X4),k2_xboole_0(k1_tarski(X5),X6))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_32]), c_0_28]), c_0_28]), c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_39,negated_conjecture,(k2_xboole_0(k1_tarski(esk1_0),k2_xboole_0(k2_tarski(esk2_0,esk3_0),k2_xboole_0(k1_tarski(esk4_0),k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k2_tarski(esk6_0,esk7_0),k1_tarski(esk8_0))))))!=k2_xboole_0(k2_tarski(esk1_0,esk2_0),k2_xboole_0(k2_tarski(esk3_0,esk4_0),k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k2_tarski(esk6_0,esk7_0),k1_tarski(esk8_0)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_40,plain,(k2_xboole_0(k2_tarski(X1,X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k1_tarski(X4),k2_xboole_0(k2_tarski(X5,X6),k2_xboole_0(k1_tarski(X7),X8)))))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_tarski(X2,X3),k2_xboole_0(k1_tarski(X4),k2_xboole_0(k2_tarski(X5,X6),k2_xboole_0(k1_tarski(X7),X8)))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_37])).
cnf(c_0_41,plain,(k2_xboole_0(k2_tarski(X1,X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k1_tarski(X4),k2_xboole_0(k2_tarski(X5,X6),k1_tarski(X7)))))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_tarski(X2,X3),k2_xboole_0(k1_tarski(X4),k2_xboole_0(k2_tarski(X5,X6),k1_tarski(X7)))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_32]), c_0_35])).
cnf(c_0_42,negated_conjecture,(k2_xboole_0(k1_tarski(esk1_0),k2_xboole_0(k1_tarski(esk2_0),k2_xboole_0(k2_tarski(esk3_0,esk4_0),k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k2_tarski(esk6_0,esk7_0),k1_tarski(esk8_0))))))!=k2_xboole_0(k2_tarski(esk1_0,esk2_0),k2_xboole_0(k2_tarski(esk3_0,esk4_0),k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k2_tarski(esk6_0,esk7_0),k1_tarski(esk8_0)))))), inference(rw,[status(thm)],[c_0_39, c_0_38])).
cnf(c_0_43,plain,(k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_tarski(X3,X4),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k2_tarski(X6,X7),k1_tarski(X8))))))=k2_xboole_0(k2_tarski(X1,X2),k2_xboole_0(k2_tarski(X3,X4),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k2_tarski(X6,X7),k1_tarski(X8)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_32]), c_0_35]), c_0_41])).
cnf(c_0_44,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 26
# Proof object formula steps           : 19
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 40
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 4
# Processed clauses                    : 15
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 14
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 36
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 0
# Paramodulations                      : 36
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 8
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 16
# ...number of literals in the above   : 16
# Current number of archived formulas  : 0
# Current number of archived clauses   : 11
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 33
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2334

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 42896 pages
