# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2))))), file('out/t65_subset_1__6_premises__rerun.E_input', d3_xboole_0)).
fof(t65_subset_1, conjecture, (![X1]:![X2]:![X3]:![X4]:~((r2_hidden(X4,X3)&(r1_tarski(X3,k2_zfmisc_1(X1,X2))&![X5]:(m1_subset_1(X5,X1)=>![X6]:(m1_subset_1(X6,X2)=>~(X4=k4_tarski(X5,X6)))))))), file('out/t65_subset_1__6_premises__rerun.E_input', t65_subset_1)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t65_subset_1__6_premises__rerun.E_input', t12_xboole_1)).
fof(l139_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,k2_zfmisc_1(X2,X3))&![X4]:![X5]:~(k4_tarski(X4,X5)=X1)))), file('out/t65_subset_1__6_premises__rerun.E_input', l139_zfmisc_1)).
fof(l54_zfmisc_1, axiom, (![X1]:![X2]:![X3]:![X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4)))), file('out/t65_subset_1__6_premises__rerun.E_input', l54_zfmisc_1)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t65_subset_1__6_premises__rerun.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t65_subset_1__6_premises__rerun.E_input', d1_xboole_0)).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(r2_hidden(X8,X5)|r2_hidden(X8,X6)))|X7!=k2_xboole_0(X5,X6))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X7))|X7!=k2_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|r2_hidden(X8,X7))|X7!=k2_xboole_0(X5,X6))))&((((~r2_hidden(esk1_3(X5,X6,X7),X5)|~r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k2_xboole_0(X5,X6))&((~r2_hidden(esk1_3(X5,X6,X7),X6)|~r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k2_xboole_0(X5,X6)))&((r2_hidden(esk1_3(X5,X6,X7),X7)|(r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k2_xboole_0(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])])).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:~((r2_hidden(X4,X3)&(r1_tarski(X3,k2_zfmisc_1(X1,X2))&![X5]:(m1_subset_1(X5,X1)=>![X6]:(m1_subset_1(X6,X2)=>~(X4=k4_tarski(X5,X6))))))))), inference(assume_negation,[status(cth)],[t65_subset_1])).
cnf(c_0_9,plain,(r2_hidden(X4,X1)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, negated_conjecture, (![X11]:![X12]:(r2_hidden(esk5_0,esk4_0)&(r1_tarski(esk4_0,k2_zfmisc_1(esk2_0,esk3_0))&(~m1_subset_1(X11,esk2_0)|(~m1_subset_1(X12,esk3_0)|esk5_0!=k4_tarski(X11,X12)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])])).
fof(c_0_11, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_12,plain,(r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(r2_hidden(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,negated_conjecture,(r1_tarski(esk4_0,k2_zfmisc_1(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, plain, (![X6]:![X7]:![X8]:(~r2_hidden(X6,k2_zfmisc_1(X7,X8))|k4_tarski(esk6_1(X6),esk7_1(X6))=X6)), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l139_zfmisc_1])])])])])])])).
cnf(c_0_17,negated_conjecture,(r2_hidden(esk5_0,k2_xboole_0(esk4_0,X1))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18,negated_conjecture,(k2_xboole_0(esk4_0,k2_zfmisc_1(esk2_0,esk3_0))=k2_zfmisc_1(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_19, plain, (![X5]:![X6]:![X7]:![X8]:![X5]:![X6]:![X7]:![X8]:(((r2_hidden(X5,X7)|~r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X7,X8)))&(r2_hidden(X6,X8)|~r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X7,X8))))&((~r2_hidden(X5,X7)|~r2_hidden(X6,X8))|r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])])])).
cnf(c_0_20,plain,(k4_tarski(esk6_1(X1),esk7_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,negated_conjecture,(r2_hidden(esk5_0,k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_22, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
fof(c_0_23, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk8_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_24,plain,(r2_hidden(X1,X3)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,negated_conjecture,(k4_tarski(esk6_1(esk5_0),esk7_1(esk5_0))=esk5_0), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,negated_conjecture,(r2_hidden(esk6_1(esk5_0),X1)|~r2_hidden(esk5_0,k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,plain,(r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,negated_conjecture,(esk5_0!=k4_tarski(X1,X2)|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_31,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,negated_conjecture,(r2_hidden(esk6_1(esk5_0),esk2_0)), inference(spm,[status(thm)],[c_0_28, c_0_21])).
cnf(c_0_33,negated_conjecture,(r2_hidden(esk7_1(esk5_0),X1)|~r2_hidden(esk5_0,k2_zfmisc_1(X2,X1))), inference(spm,[status(thm)],[c_0_29, c_0_25])).
cnf(c_0_34,negated_conjecture,(~m1_subset_1(esk7_1(esk5_0),esk3_0)|~m1_subset_1(esk6_1(esk5_0),esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_25])).
cnf(c_0_35,negated_conjecture,(m1_subset_1(esk6_1(esk5_0),esk2_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36,negated_conjecture,(r2_hidden(esk7_1(esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_37,negated_conjecture,(~m1_subset_1(esk7_1(esk5_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35])])).
cnf(c_0_38,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_36]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 24
# Proof object formula steps           : 15
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 299
# ...of these trivial                  : 1
# ...subsumed                          : 52
# ...remaining for further processing  : 246
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 1557
# ...of the previous two non-trivial   : 1305
# Contextual simplify-reflections      : 8
# Paramodulations                      : 1515
# Factorizations                       : 14
# Equation resolutions                 : 28
# Current number of processed clauses  : 243
#    Positive orientable unit clauses  : 113
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 48
#    Non-unit-clauses                  : 82
# Current number of unprocessed clauses: 1006
# ...number of literals in the above   : 2645
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 898
# Rec. Clause-clause subsumption calls : 472
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 616
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 450
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 28303

# -------------------------------------------------
# User time                : 0.036 s
# System time              : 0.000 s
# Total time               : 0.036 s
# Maximum resident set size: 43032 pages
