# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:?[X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6)))))), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', d2_zfmisc_1)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', d1_tarski)).
fof(t7_mcart_1, axiom, (![X1]:![X2]:(k1_xtuple_0(k4_tarski(X1,X2))=X1&k2_xtuple_0(k4_tarski(X1,X2))=X2)), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', t7_mcart_1)).
fof(t65_mcart_1, conjecture, (![X1]:![X2]:![X3]:![X4]:k4_zfmisc_1(k1_tarski(X1),k1_tarski(X2),k1_tarski(X3),k1_tarski(X4))=k1_tarski(k4_mcart_1(X1,X2,X3,X4))), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', t65_mcart_1)).
fof(d4_zfmisc_1, axiom, (![X1]:![X2]:![X3]:![X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', d4_zfmisc_1)).
fof(d3_zfmisc_1, axiom, (![X1]:![X2]:![X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', d3_zfmisc_1)).
fof(d4_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', d4_mcart_1)).
fof(d3_mcart_1, axiom, (![X1]:![X2]:![X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), file('out/t65_mcart_1__11_premises__rerun__rerun.E_input', d3_mcart_1)).
fof(c_0_8, plain, (![X7]:![X8]:![X9]:![X10]:![X10]:![X13]:![X14]:![X7]:![X8]:![X9]:![X16]:![X17]:(((((r2_hidden(esk1_4(X7,X8,X9,X10),X7)|~r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8))&(((r2_hidden(esk2_4(X7,X8,X9,X10),X8)|~r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8))&((X10=k4_tarski(esk1_4(X7,X8,X9,X10),esk2_4(X7,X8,X9,X10))|~r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8))))&(((~r2_hidden(X13,X7)|(~r2_hidden(X14,X8)|X10!=k4_tarski(X13,X14)))|r2_hidden(X10,X9))|X9!=k2_zfmisc_1(X7,X8)))&(((~r2_hidden(esk3_3(X7,X8,X9),X9)|(~r2_hidden(X16,X7)|(~r2_hidden(X17,X8)|esk3_3(X7,X8,X9)!=k4_tarski(X16,X17))))|X9=k2_zfmisc_1(X7,X8))&(((r2_hidden(esk4_3(X7,X8,X9),X7)|r2_hidden(esk3_3(X7,X8,X9),X9))|X9=k2_zfmisc_1(X7,X8))&(((r2_hidden(esk5_3(X7,X8,X9),X8)|r2_hidden(esk3_3(X7,X8,X9),X9))|X9=k2_zfmisc_1(X7,X8))&((esk3_3(X7,X8,X9)=k4_tarski(esk4_3(X7,X8,X9),esk5_3(X7,X8,X9))|r2_hidden(esk3_3(X7,X8,X9),X9))|X9=k2_zfmisc_1(X7,X8))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])])).
cnf(c_0_9,plain,(r2_hidden(X4,X1)|X1!=k2_zfmisc_1(X2,X3)|X4!=k4_tarski(X5,X6)|~r2_hidden(X6,X3)|~r2_hidden(X5,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_10,plain,(X4=k4_tarski(esk1_4(X2,X3,X1,X4),esk2_4(X2,X3,X1,X4))|X1!=k2_zfmisc_1(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11,plain,(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_9])])).
cnf(c_0_12,plain,(k4_tarski(esk1_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk2_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_13,plain,(r2_hidden(esk2_4(X2,X3,X1,X4),X3)|X1!=k2_zfmisc_1(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk6_2(X4,X5),X5)|esk6_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk6_2(X4,X5),X5)|esk6_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_15,plain,(r2_hidden(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(esk2_4(X4,X5,k2_zfmisc_1(X4,X5),X1),X3)|~r2_hidden(esk1_4(X4,X5,k2_zfmisc_1(X4,X5),X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X4,X5))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16,plain,(r2_hidden(esk2_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, (![X3]:![X4]:![X3]:![X4]:(k1_xtuple_0(k4_tarski(X3,X4))=X3&k2_xtuple_0(k4_tarski(X3,X4))=X4)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t7_mcart_1])])])).
cnf(c_0_19,plain,(r2_hidden(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(esk1_4(X4,X3,k2_zfmisc_1(X4,X3),X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X4,X3))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_17])])).
cnf(c_0_21,plain,(k1_xtuple_0(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22,plain,(r2_hidden(X1,k2_zfmisc_1(k1_tarski(esk1_4(X2,X3,k2_zfmisc_1(X2,X3),X1)),X3))|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23,plain,(esk1_4(X1,X2,k2_zfmisc_1(X1,X2),X3)=k1_xtuple_0(X3)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_21, c_0_12])).
cnf(c_0_24,plain,(r2_hidden(X1,k2_zfmisc_1(k1_tarski(k1_xtuple_0(X1)),X2))|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26,plain,(r2_hidden(esk1_4(X2,X3,X1,X4),X2)|X1!=k2_zfmisc_1(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_27,plain,(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X1),X3))|~r2_hidden(X2,X3)|~r2_hidden(X1,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_11]), c_0_21])).
cnf(c_0_28,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_29,plain,(r2_hidden(esk1_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_30,plain,(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X1),X3))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_27, c_0_20])).
cnf(c_0_31,plain,(k4_tarski(k1_xtuple_0(X1),esk2_4(X2,X3,k2_zfmisc_1(X2,X3),X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_12, c_0_23])).
cnf(c_0_32,plain,(esk1_4(k1_tarski(X1),X2,k2_zfmisc_1(k1_tarski(X1),X2),X3)=X1|~r2_hidden(X3,k2_zfmisc_1(k1_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33,plain,(r2_hidden(X1,k2_zfmisc_1(k1_tarski(k1_xtuple_0(X1)),X2))|~r2_hidden(esk2_4(X3,X4,k2_zfmisc_1(X3,X4),X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X4))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,plain,(k2_xtuple_0(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_35,plain,(k4_tarski(X1,esk2_4(k1_tarski(X1),X2,k2_zfmisc_1(k1_tarski(X1),X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(k1_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_12, c_0_32])).
cnf(c_0_36,plain,(esk2_4(X1,k1_tarski(X2),k2_zfmisc_1(X1,k1_tarski(X2)),X3)=X2|~r2_hidden(X3,k2_zfmisc_1(X1,k1_tarski(X2)))), inference(spm,[status(thm)],[c_0_28, c_0_16])).
cnf(c_0_37,plain,(r2_hidden(X1,k2_zfmisc_1(k1_tarski(k1_xtuple_0(X1)),k1_tarski(esk2_4(X2,X3,k2_zfmisc_1(X2,X3),X1))))|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_33, c_0_20])).
cnf(c_0_38,plain,(esk2_4(X1,X2,k2_zfmisc_1(X1,X2),X3)=k2_xtuple_0(X3)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_34, c_0_12])).
fof(c_0_39, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:k4_zfmisc_1(k1_tarski(X1),k1_tarski(X2),k1_tarski(X3),k1_tarski(X4))=k1_tarski(k4_mcart_1(X1,X2,X3,X4)))), inference(assume_negation,[status(cth)],[t65_mcart_1])).
fof(c_0_40, plain, (![X5]:![X6]:![X7]:![X8]:k4_zfmisc_1(X5,X6,X7,X8)=k2_zfmisc_1(k3_zfmisc_1(X5,X6,X7),X8)), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_41, plain, (![X4]:![X5]:![X6]:k3_zfmisc_1(X4,X5,X6)=k2_zfmisc_1(k2_zfmisc_1(X4,X5),X6)), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_42, plain, (![X5]:![X6]:![X7]:![X8]:k4_mcart_1(X5,X6,X7,X8)=k4_tarski(k3_mcart_1(X5,X6,X7),X8)), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_43, plain, (![X4]:![X5]:![X6]:k3_mcart_1(X4,X5,X6)=k4_tarski(k4_tarski(X4,X5),X6)), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_44,plain,(k4_tarski(X1,X2)=X3|~r2_hidden(X3,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2)))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45,plain,(X1=k1_tarski(X2)|esk6_2(X2,X1)=X2|r2_hidden(esk6_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_46,plain,(r2_hidden(X1,k2_zfmisc_1(k1_tarski(k1_xtuple_0(X1)),k1_tarski(k2_xtuple_0(X1))))|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_47, negated_conjecture, (k4_zfmisc_1(k1_tarski(esk7_0),k1_tarski(esk8_0),k1_tarski(esk9_0),k1_tarski(esk10_0))!=k1_tarski(k4_mcart_1(esk7_0,esk8_0,esk9_0,esk10_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])).
cnf(c_0_48,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49,plain,(k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50,plain,(k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51,plain,(k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52,plain,(esk6_2(X1,k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3)))=k4_tarski(X2,X3)|esk6_2(X1,k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3)))=X1|k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))=k1_tarski(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53,plain,(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2)))|~r2_hidden(X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_30]), c_0_21]), c_0_34])).
cnf(c_0_54,negated_conjecture,(k4_zfmisc_1(k1_tarski(esk7_0),k1_tarski(esk8_0),k1_tarski(esk9_0),k1_tarski(esk10_0))!=k1_tarski(k4_mcart_1(esk7_0,esk8_0,esk9_0,esk10_0))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_55,plain,(k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_56,plain,(k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57,plain,(X1=k1_tarski(X2)|esk6_2(X2,X1)!=X2|~r2_hidden(esk6_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_58,plain,(esk6_2(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2)))=k4_tarski(X1,X2)|k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2))), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_52])])).
cnf(c_0_59,plain,(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2)))), inference(spm,[status(thm)],[c_0_53, c_0_20])).
cnf(c_0_60,negated_conjecture,(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_tarski(esk7_0),k1_tarski(esk8_0)),k1_tarski(esk9_0)),k1_tarski(esk10_0))!=k1_tarski(k4_tarski(k4_tarski(k4_tarski(esk7_0,esk8_0),esk9_0),esk10_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55]), c_0_56])).
cnf(c_0_61,plain,(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59])])).
cnf(c_0_62,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61]), c_0_61]), c_0_61])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 46
# Proof object formula steps           : 17
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 15
# Processed clauses                    : 3475
# ...of these trivial                  : 1
# ...subsumed                          : 2681
# ...remaining for further processing  : 793
# Other redundant clauses eliminated   : 131
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 68
# Generated clauses                    : 70529
# ...of the previous two non-trivial   : 68777
# Contextual simplify-reflections      : 1382
# Paramodulations                      : 70392
# Factorizations                       : 8
# Equation resolutions                 : 131
# Current number of processed clauses  : 688
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 684
# Current number of unprocessed clauses: 48662
# ...number of literals in the above   : 242140
# Current number of archived formulas  : 0
# Current number of archived clauses   : 103
# Clause-clause subsumption calls (NU) : 75178
# Rec. Clause-clause subsumption calls : 43269
# Non-unit clause-clause subsumptions  : 4079
# Unit Clause-clause subsumption calls : 96
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 40
# BW rewrite match successes           : 40
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1648606

# -------------------------------------------------
# User time                : 1.476 s
# System time              : 0.052 s
# Total time               : 1.528 s
# Maximum resident set size: 63408 pages
