# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t66_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X2))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>![X7]:(m1_subset_1(X7,u1_struct_0(X4))=>((v3_pre_topc(X5,X2)&(r2_hidden(X6,X5)&(r1_tarski(X5,u1_struct_0(X4))&X6=X7)))=>(r1_tmap_1(X2,X1,X3,X6)<=>r1_tmap_1(X4,X1,k2_tmap_1(X2,X1,X3,X4),X7))))))))))), file('out/t66_tmap_1__6_premises__rerun__rerun.E_input', t66_tmap_1)).
fof(t65_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X2))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>![X7]:(m1_subset_1(X7,u1_struct_0(X4))=>((r1_tarski(X5,u1_struct_0(X4))&(m1_connsp_2(X5,X2,X6)&X6=X7))=>(r1_tmap_1(X2,X1,X3,X6)<=>r1_tmap_1(X4,X1,k2_tmap_1(X2,X1,X3,X4),X7))))))))))), file('out/t66_tmap_1__6_premises__rerun__rerun.E_input', t65_tmap_1)).
fof(dt_m1_connsp_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>![X3]:(m1_connsp_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t66_tmap_1__6_premises__rerun__rerun.E_input', dt_m1_connsp_2)).
fof(t55_pre_topc, axiom, (![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1)))))), file('out/t66_tmap_1__6_premises__rerun__rerun.E_input', t55_pre_topc)).
fof(t9_connsp_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((m1_connsp_2(X4,X1,X3)&r1_tarski(X4,X2))))))))))), file('out/t66_tmap_1__6_premises__rerun__rerun.E_input', t9_connsp_2)).
fof(t1_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3))), file('out/t66_tmap_1__6_premises__rerun__rerun.E_input', t1_xboole_1)).
fof(c_0_6, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X2))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>![X7]:(m1_subset_1(X7,u1_struct_0(X4))=>((v3_pre_topc(X5,X2)&(r2_hidden(X6,X5)&(r1_tarski(X5,u1_struct_0(X4))&X6=X7)))=>(r1_tmap_1(X2,X1,X3,X6)<=>r1_tmap_1(X4,X1,k2_tmap_1(X2,X1,X3,X4),X7)))))))))))), inference(assume_negation,[status(cth)],[t66_tmap_1])).
fof(c_0_7, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((((~r1_tmap_1(X9,X8,X10,X13)|r1_tmap_1(X11,X8,k2_tmap_1(X9,X8,X10,X11),X14))|(~r1_tarski(X12,u1_struct_0(X11))|(~m1_connsp_2(X12,X9,X13)|X13!=X14)))|~m1_subset_1(X14,u1_struct_0(X11)))|~m1_subset_1(X13,u1_struct_0(X9)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X9))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))&(((((((((~r1_tmap_1(X11,X8,k2_tmap_1(X9,X8,X10,X11),X14)|r1_tmap_1(X9,X8,X10,X13))|(~r1_tarski(X12,u1_struct_0(X11))|(~m1_connsp_2(X12,X9,X13)|X13!=X14)))|~m1_subset_1(X14,u1_struct_0(X11)))|~m1_subset_1(X13,u1_struct_0(X9)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X9))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_tmap_1])])])])])])])).
fof(c_0_8, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|(~m1_connsp_2(X6,X4,X5)|m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_connsp_2])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|~l1_pre_topc(X4))|((v2_struct_0(X5)|~m1_pre_topc(X5,X4))|(~m1_subset_1(X6,u1_struct_0(X5))|m1_subset_1(X6,u1_struct_0(X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])])])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X10]:(((((((m1_subset_1(esk8_3(X5,X6,X7),k1_zfmisc_1(u1_struct_0(X5)))|~r2_hidden(X7,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~v3_pre_topc(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((((m1_connsp_2(esk8_3(X5,X6,X7),X5,X7)|~r2_hidden(X7,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~v3_pre_topc(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((r1_tarski(esk8_3(X5,X6,X7),X6)|~r2_hidden(X7,X6))|~m1_subset_1(X7,u1_struct_0(X5)))|~v3_pre_topc(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))))&((((m1_subset_1(esk9_2(X5,X6),u1_struct_0(X5))|v3_pre_topc(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((r2_hidden(esk9_2(X5,X6),X6)|v3_pre_topc(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X5)))|(~m1_connsp_2(X10,X5,esk9_2(X5,X6))|~r1_tarski(X10,X6)))|v3_pre_topc(X6,X5))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_connsp_2])])])])])])])])).
fof(c_0_11, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk2_0))&(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(m1_subset_1(esk6_0,u1_struct_0(esk2_0))&(m1_subset_1(esk7_0,u1_struct_0(esk4_0))&((v3_pre_topc(esk5_0,esk2_0)&(r2_hidden(esk6_0,esk5_0)&(r1_tarski(esk5_0,u1_struct_0(esk4_0))&esk6_0=esk7_0)))&((~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk6_0)|~r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),esk7_0))&(r1_tmap_1(esk2_0,esk1_0,esk3_0,esk6_0)|r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),esk7_0)))))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
cnf(c_0_12,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|r1_tmap_1(X2,X1,X3,X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_pre_topc(X4,X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X4))|X6!=X7|~m1_connsp_2(X5,X2,X6)|~r1_tarski(X5,u1_struct_0(X4))|~r1_tmap_1(X4,X1,k2_tmap_1(X2,X1,X3,X4),X7)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m1_connsp_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X5,X6))|r1_tarski(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_16,plain,(v2_struct_0(X1)|r1_tarski(esk8_3(X1,X2,X3),X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,negated_conjecture,(r2_hidden(esk6_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,plain,(r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_connsp_2(X6,X1,X4)|~r1_tmap_1(X5,X2,k2_tmap_1(X1,X2,X3,X5),X4)|~r1_tarski(X6,u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X5))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_12, c_0_13])]), c_0_14])).
cnf(c_0_19,negated_conjecture,(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24,negated_conjecture,(v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_26,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_27,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28,plain,(v2_struct_0(X1)|m1_connsp_2(esk8_3(X1,X2,X3),X1,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29,plain,(r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30,negated_conjecture,(r1_tarski(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31,negated_conjecture,(r1_tarski(esk8_3(X1,esk5_0,esk6_0),esk5_0)|v2_struct_0(X1)|~v3_pre_topc(esk5_0,X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk6_0,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_32,negated_conjecture,(v3_pre_topc(esk5_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_33,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_34,negated_conjecture,(m1_subset_1(esk6_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|r1_tmap_1(X4,X1,k2_tmap_1(X2,X1,X3,X4),X7)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_pre_topc(X4,X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X4))|X6!=X7|~m1_connsp_2(X5,X2,X6)|~r1_tarski(X5,u1_struct_0(X4))|~r1_tmap_1(X2,X1,X3,X6)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_36,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,X1)|v2_struct_0(X2)|~m1_connsp_2(X3,esk2_0,X1)|~r1_tmap_1(X2,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,X2),X1)|~r1_tarski(X3,u1_struct_0(X2))|~m1_pre_topc(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(X2))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])]), c_0_26]), c_0_27])).
cnf(c_0_37,negated_conjecture,(m1_pre_topc(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39,negated_conjecture,(m1_connsp_2(esk8_3(X1,esk5_0,esk6_0),X1,esk6_0)|v2_struct_0(X1)|~v3_pre_topc(esk5_0,X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk6_0,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_28, c_0_17])).
cnf(c_0_40,negated_conjecture,(r1_tarski(X1,u1_struct_0(esk4_0))|~r1_tarski(X1,esk5_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_41,negated_conjecture,(r1_tarski(esk8_3(esk2_0,esk5_0,esk6_0),esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34]), c_0_22]), c_0_24])]), c_0_27])).
cnf(c_0_42,negated_conjecture,(m1_subset_1(esk7_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_43,negated_conjecture,(esk6_0=esk7_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_44,negated_conjecture,(r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),esk7_0)|r1_tmap_1(esk2_0,esk1_0,esk3_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_45,plain,(r1_tmap_1(X1,X2,k2_tmap_1(X3,X2,X4,X1),X5)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_connsp_2(X6,X3,X5)|~r1_tmap_1(X3,X2,X4,X5)|~r1_tarski(X6,u1_struct_0(X1))|~m1_pre_topc(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_35, c_0_13])]), c_0_14])).
cnf(c_0_46,negated_conjecture,(~r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),esk7_0)|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_47,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,X1)|~m1_connsp_2(X2,esk2_0,X1)|~r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),X1)|~r1_tarski(X2,u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_48,negated_conjecture,(m1_connsp_2(esk8_3(esk2_0,esk5_0,esk6_0),esk2_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_32]), c_0_33]), c_0_34]), c_0_22]), c_0_24])]), c_0_27])).
cnf(c_0_49,negated_conjecture,(r1_tarski(esk8_3(esk2_0,esk5_0,esk6_0),u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50,negated_conjecture,(m1_subset_1(esk6_0,u1_struct_0(esk4_0))), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51,negated_conjecture,(r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),esk6_0)|r1_tmap_1(esk2_0,esk1_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[c_0_44, c_0_43])).
cnf(c_0_52,negated_conjecture,(r1_tmap_1(X1,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,X1),X2)|v2_struct_0(X1)|~m1_connsp_2(X3,esk2_0,X2)|~r1_tmap_1(esk2_0,esk1_0,esk3_0,X2)|~r1_tarski(X3,u1_struct_0(X1))|~m1_pre_topc(X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])]), c_0_26]), c_0_27])).
cnf(c_0_53,negated_conjecture,(~r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),esk6_0)|~r1_tmap_1(esk2_0,esk1_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[c_0_46, c_0_43])).
cnf(c_0_54,negated_conjecture,(r1_tmap_1(esk2_0,esk1_0,esk3_0,esk6_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50])]), c_0_51])).
cnf(c_0_55,negated_conjecture,(r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),X1)|~m1_connsp_2(X2,esk2_0,X1)|~r1_tmap_1(esk2_0,esk1_0,esk3_0,X1)|~r1_tarski(X2,u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_37]), c_0_38])).
cnf(c_0_56,negated_conjecture,(~r1_tmap_1(esk4_0,esk1_0,k2_tmap_1(esk2_0,esk1_0,esk3_0,esk4_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_57,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_48]), c_0_54]), c_0_49]), c_0_50])]), c_0_56]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 45
# Proof object formula steps           : 13
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 85
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 85
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 28
# ...of the previous two non-trivial   : 24
# Contextual simplify-reflections      : 5
# Paramodulations                      : 26
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 26
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 8
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 751
# Rec. Clause-clause subsumption calls : 12
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4605

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45272 pages
