# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(s3_funct_2__e3_108__tex_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r2_hidden(X4,X3)=>X5=X4)))=>?[X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_hidden(X5,X3)=>k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X4,X5)=X5)))))), file('out/t66_tex_2__9_premises__rerun.E_input', s3_funct_2__e3_108__tex_2)).
fof(cc1_tdlat_3, axiom, (![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1)))), file('out/t66_tex_2__9_premises__rerun.E_input', cc1_tdlat_3)).
fof(existence_m1_subset_1, axiom, (![X1]:?[X2]:m1_subset_1(X2,X1)), file('out/t66_tex_2__9_premises__rerun.E_input', existence_m1_subset_1)).
fof(t1_tsep_1, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t66_tex_2__9_premises__rerun.E_input', t1_tsep_1)).
fof(t66_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>?[X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))&v3_borsuk_1(X3,X1,X2))))), file('out/t66_tex_2__9_premises__rerun.E_input', t66_tex_2)).
fof(t62_tex_2, axiom, (![X1]:((v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2))))), file('out/t66_tex_2__9_premises__rerun.E_input', t62_tex_2)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t66_tex_2__9_premises__rerun.E_input', dt_m1_pre_topc)).
fof(cc1_pre_topc, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2)))), file('out/t66_tex_2__9_premises__rerun.E_input', cc1_pre_topc)).
fof(t1_subset, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2))), file('out/t66_tex_2__9_premises__rerun.E_input', t1_subset)).
fof(d16_borsuk_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_borsuk_1(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4)=X4))))))), file('out/t66_tex_2__9_premises__rerun.E_input', d16_borsuk_1)).
fof(c_0_10, plain, (![X6]:![X7]:![X8]:![X10]:![X12]:(((((v1_funct_1(esk2_3(X6,X7,X8))|m1_subset_1(esk1_3(X6,X7,X8),u1_struct_0(X6)))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))&(((v1_funct_2(esk2_3(X6,X7,X8),u1_struct_0(X6),u1_struct_0(X7))|m1_subset_1(esk1_3(X6,X7,X8),u1_struct_0(X6)))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))&((m1_subset_1(esk2_3(X6,X7,X8),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X7))))|m1_subset_1(esk1_3(X6,X7,X8),u1_struct_0(X6)))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))))&(((~m1_subset_1(X12,u1_struct_0(X6))|(~r2_hidden(X12,X8)|k3_funct_2(u1_struct_0(X6),u1_struct_0(X7),esk2_3(X6,X7,X8),X12)=X12))|m1_subset_1(esk1_3(X6,X7,X8),u1_struct_0(X6)))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))))))&(((((v1_funct_1(esk2_3(X6,X7,X8))|(r2_hidden(esk1_3(X6,X7,X8),X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))&(((v1_funct_2(esk2_3(X6,X7,X8),u1_struct_0(X6),u1_struct_0(X7))|(r2_hidden(esk1_3(X6,X7,X8),X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))&((m1_subset_1(esk2_3(X6,X7,X8),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X7))))|(r2_hidden(esk1_3(X6,X7,X8),X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))))&(((~m1_subset_1(X12,u1_struct_0(X6))|(~r2_hidden(X12,X8)|k3_funct_2(u1_struct_0(X6),u1_struct_0(X7),esk2_3(X6,X7,X8),X12)=X12))|(r2_hidden(esk1_3(X6,X7,X8),X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))))))&((((v1_funct_1(esk2_3(X6,X7,X8))|(X10!=esk1_3(X6,X7,X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))&(((v1_funct_2(esk2_3(X6,X7,X8),u1_struct_0(X6),u1_struct_0(X7))|(X10!=esk1_3(X6,X7,X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))&((m1_subset_1(esk2_3(X6,X7,X8),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X7))))|(X10!=esk1_3(X6,X7,X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6))))))))&(((~m1_subset_1(X12,u1_struct_0(X6))|(~r2_hidden(X12,X8)|k3_funct_2(u1_struct_0(X6),u1_struct_0(X7),esk2_3(X6,X7,X8),X12)=X12))|(X10!=esk1_3(X6,X7,X8)|~m1_subset_1(X10,u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~l1_pre_topc(X6))))|((v2_struct_0(X7)|~m1_pre_topc(X7,X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s3_funct_2__e3_108__tex_2])])])])])])])])).
fof(c_0_11, plain, (![X2]:(~l1_pre_topc(X2)|(~v1_tdlat_3(X2)|v2_pre_topc(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
cnf(c_0_12,plain,(v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(esk1_3(X2,X3,X1),X1)|m1_subset_1(esk2_3(X2,X3,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13,plain,(v2_pre_topc(X1)|~v1_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, (![X3]:m1_subset_1(esk3_1(X3),X3)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_15,plain,(v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(esk1_3(X2,X3,X1),X1)|v1_funct_2(esk2_3(X2,X3,X1),u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(r2_hidden(esk1_3(X1,X2,X3),X3)|m1_subset_1(esk2_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(csr,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17,plain,(m1_subset_1(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(X3)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])])])).
fof(c_0_19, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>?[X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))&v3_borsuk_1(X3,X1,X2)))))), inference(assume_negation,[status(cth)],[t66_tex_2])).
cnf(c_0_20,plain,(v1_funct_2(esk2_3(X1,X2,X3),u1_struct_0(X1),u1_struct_0(X2))|r2_hidden(esk1_3(X1,X2,X3),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(csr,[status(thm)],[c_0_15, c_0_13])).
cnf(c_0_21,plain,(r2_hidden(esk1_3(X1,X2,X3),X3)|m1_subset_1(esk2_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22,plain,(m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, negated_conjecture, (![X6]:((~v2_struct_0(esk4_0)&(v2_pre_topc(esk4_0)&(v1_tdlat_3(esk4_0)&l1_pre_topc(esk4_0))))&((~v2_struct_0(esk5_0)&m1_pre_topc(esk5_0,esk4_0))&((~v1_funct_1(X6)|(~v1_funct_2(X6,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|(~v5_pre_topc(X6,esk4_0,esk5_0)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))))))|~v3_borsuk_1(X6,esk4_0,esk5_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])])])])).
cnf(c_0_24,plain,(v1_funct_2(esk2_3(X1,X2,X3),u1_struct_0(X1),u1_struct_0(X2))|r2_hidden(esk1_3(X1,X2,X3),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_20, c_0_17])).
cnf(c_0_25,plain,(r2_hidden(esk1_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X3))|m1_subset_1(esk2_3(X1,X2,u1_struct_0(X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26,negated_conjecture,(m1_pre_topc(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27,negated_conjecture,(l1_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,negated_conjecture,(v1_tdlat_3(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29,negated_conjecture,(~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,plain,(v1_funct_2(esk2_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X1),u1_struct_0(X2))|r2_hidden(esk1_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_24, c_0_22])).
fof(c_0_32, plain, (![X4]:![X5]:![X6]:((~v2_pre_topc(X4)|(~v1_tdlat_3(X4)|~l1_pre_topc(X4)))|((~v2_pre_topc(X5)|~l1_pre_topc(X5))|((~v1_funct_1(X6)|(~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X5))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))))|v5_pre_topc(X6,X4,X5))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t62_tex_2])])])])])).
fof(c_0_33, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
fof(c_0_34, plain, (![X3]:![X4]:((~v2_pre_topc(X3)|~l1_pre_topc(X3))|(~m1_pre_topc(X4,X3)|v2_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])])])).
cnf(c_0_35,plain,(v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(esk1_3(X2,X3,X1),X1)|v1_funct_1(esk2_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36,negated_conjecture,(r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(X1)),u1_struct_0(X1))|m1_subset_1(esk2_3(esk4_0,esk5_0,u1_struct_0(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk5_0))))|~m1_pre_topc(X1,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])]), c_0_29]), c_0_30])).
cnf(c_0_37,negated_conjecture,(v1_funct_2(esk2_3(esk4_0,esk5_0,u1_struct_0(X1)),u1_struct_0(esk4_0),u1_struct_0(esk5_0))|r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(X1)),u1_struct_0(X1))|~m1_pre_topc(X1,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_26]), c_0_27]), c_0_28])]), c_0_29]), c_0_30])).
cnf(c_0_38,plain,(v5_pre_topc(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40,plain,(v2_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41,negated_conjecture,(v2_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_42,plain,(v1_funct_1(esk2_3(X1,X2,X3))|r2_hidden(esk1_3(X1,X2,X3),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(csr,[status(thm)],[c_0_35, c_0_13])).
cnf(c_0_43,plain,(v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(esk1_3(X2,X3,X1),X1)|k3_funct_2(u1_struct_0(X2),u1_struct_0(X3),esk2_3(X2,X3,X1),X5)=X5|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))|~r2_hidden(X5,X1)|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_44,negated_conjecture,(~v3_borsuk_1(X1,esk4_0,esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk5_0))))|~v5_pre_topc(X1,esk4_0,esk5_0)|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_45,negated_conjecture,(r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|m1_subset_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk5_0))))), inference(spm,[status(thm)],[c_0_36, c_0_26])).
cnf(c_0_46,negated_conjecture,(v1_funct_2(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk4_0),u1_struct_0(esk5_0))|r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))), inference(spm,[status(thm)],[c_0_37, c_0_26])).
cnf(c_0_47,plain,(v5_pre_topc(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[c_0_38, c_0_13])).
cnf(c_0_48,negated_conjecture,(l1_pre_topc(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_26]), c_0_27])])).
cnf(c_0_49,negated_conjecture,(v2_pre_topc(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_26]), c_0_27]), c_0_41])])).
fof(c_0_50, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|m1_subset_1(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_51,plain,(v1_funct_1(esk2_3(X1,X2,X3))|r2_hidden(esk1_3(X1,X2,X3),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_42, c_0_17])).
cnf(c_0_52,plain,(k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),esk2_3(X1,X2,X3),X4)=X4|r2_hidden(esk1_3(X1,X2,X3),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(csr,[status(thm)],[c_0_43, c_0_13])).
fof(c_0_53, plain, (![X5]:![X6]:![X7]:![X8]:(((((~v3_borsuk_1(X7,X5,X6)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,u1_struct_0(X6))|k3_funct_2(u1_struct_0(X5),u1_struct_0(X6),X7,X8)=X8)))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((m1_subset_1(esk6_3(X5,X6,X7),u1_struct_0(X5))|v3_borsuk_1(X7,X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((((r2_hidden(esk6_3(X5,X6,X7),u1_struct_0(X6))|v3_borsuk_1(X7,X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&((((k3_funct_2(u1_struct_0(X5),u1_struct_0(X6),X7,esk6_3(X5,X6,X7))!=esk6_3(X5,X6,X7)|v3_borsuk_1(X7,X5,X6))|(~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6)))))))|(v2_struct_0(X6)|~m1_pre_topc(X6,X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_borsuk_1])])])])])])])])).
cnf(c_0_54,negated_conjecture,(r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|~v3_borsuk_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),esk4_0,esk5_0)|~v5_pre_topc(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),esk4_0,esk5_0)|~v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
cnf(c_0_55,negated_conjecture,(v5_pre_topc(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),esk4_0,esk5_0)|r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|~v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_45]), c_0_48]), c_0_27]), c_0_28]), c_0_49])]), c_0_46])).
cnf(c_0_56,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57,plain,(v1_funct_1(esk2_3(X1,X2,u1_struct_0(X3)))|r2_hidden(esk1_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_51, c_0_22])).
cnf(c_0_58,plain,(k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),esk2_3(X1,X2,X3),X4)=X4|r2_hidden(esk1_3(X1,X2,X3),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_52, c_0_17])).
cnf(c_0_59,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v3_borsuk_1(X3,X1,X2)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60,negated_conjecture,(r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|~v3_borsuk_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),esk4_0,esk5_0)|~v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61,plain,(v1_funct_1(esk2_3(X1,X2,u1_struct_0(X3)))|m1_subset_1(esk1_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62,plain,(k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),esk2_3(X1,X2,u1_struct_0(X3)),X4)=X4|r2_hidden(esk1_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(X4,u1_struct_0(X3))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_58, c_0_22])).
cnf(c_0_63,negated_conjecture,(r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|m1_subset_1(esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0))),u1_struct_0(esk4_0))|~v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_45]), c_0_26]), c_0_27]), c_0_41])]), c_0_29]), c_0_30]), c_0_46]), c_0_60])).
cnf(c_0_64,negated_conjecture,(v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(X1)))|m1_subset_1(esk1_3(esk4_0,esk5_0,u1_struct_0(X1)),u1_struct_0(X1))|~m1_pre_topc(X1,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_26]), c_0_27]), c_0_28])]), c_0_30]), c_0_29])).
cnf(c_0_65,negated_conjecture,(k3_funct_2(u1_struct_0(esk4_0),u1_struct_0(X1),esk2_3(esk4_0,X1,u1_struct_0(X2)),esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0))))=esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))|r2_hidden(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|r2_hidden(esk1_3(esk4_0,X1,u1_struct_0(X2)),u1_struct_0(X2))|v2_struct_0(X1)|~v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))|~r2_hidden(esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0))),u1_struct_0(X2))|~m1_pre_topc(X1,esk4_0)|~m1_pre_topc(X2,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_27]), c_0_28])]), c_0_30])).
cnf(c_0_66,negated_conjecture,(v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))|m1_subset_1(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))), inference(spm,[status(thm)],[c_0_64, c_0_26])).
cnf(c_0_67,plain,(v2_struct_0(X3)|v2_struct_0(X2)|m1_subset_1(esk2_3(X2,X3,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))|X4!=esk1_3(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_68,plain,(v2_struct_0(X3)|v2_struct_0(X2)|v1_funct_1(esk2_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))|X4!=esk1_3(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_69,plain,(v2_struct_0(X3)|v2_struct_0(X2)|v1_funct_2(esk2_3(X2,X3,X1),u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))|X4!=esk1_3(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_70,negated_conjecture,(k3_funct_2(u1_struct_0(esk4_0),u1_struct_0(X1),esk2_3(esk4_0,X1,u1_struct_0(X2)),esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0))))=esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))|r2_hidden(esk1_3(esk4_0,X1,u1_struct_0(X2)),u1_struct_0(X2))|m1_subset_1(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|v2_struct_0(X1)|~r2_hidden(esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0))),u1_struct_0(X2))|~m1_pre_topc(X1,esk4_0)|~m1_pre_topc(X2,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_56])).
cnf(c_0_71,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v3_borsuk_1(X3,X1,X2)|r2_hidden(esk6_3(X1,X2,X3),u1_struct_0(X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_72,plain,(v2_struct_0(X3)|v2_struct_0(X2)|k3_funct_2(u1_struct_0(X2),u1_struct_0(X3),esk2_3(X2,X3,X1),X5)=X5|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,u1_struct_0(X3))|X4!=esk1_3(X2,X3,X1)|~r2_hidden(X5,X1)|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_73,plain,(m1_subset_1(esk2_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_67, c_0_13])])).
cnf(c_0_74,plain,(v1_funct_1(esk2_3(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_68, c_0_13])])).
cnf(c_0_75,plain,(v1_funct_2(esk2_3(X1,X2,X3),u1_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_69, c_0_13])])).
cnf(c_0_76,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v3_borsuk_1(X3,X1,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,esk6_3(X1,X2,X3))!=esk6_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_77,negated_conjecture,(k3_funct_2(u1_struct_0(esk4_0),u1_struct_0(X1),esk2_3(esk4_0,X1,u1_struct_0(esk5_0)),esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0))))=esk6_3(esk4_0,esk5_0,esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))|v3_borsuk_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),esk4_0,esk5_0)|r2_hidden(esk1_3(esk4_0,X1,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|m1_subset_1(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))|v2_struct_0(X1)|~v1_funct_2(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk5_0))))|~m1_pre_topc(X1,esk4_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_26]), c_0_27]), c_0_41])]), c_0_29]), c_0_30]), c_0_66])).
cnf(c_0_78,plain,(k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),esk2_3(X1,X2,X3),X4)=X4|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(X4,X3)|~m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_72, c_0_13])])).
cnf(c_0_79,plain,(v3_borsuk_1(esk2_3(X1,X2,X3),X1,X2)|m1_subset_1(esk6_3(X1,X2,esk2_3(X1,X2,X3)),u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_73]), c_0_13]), c_0_74]), c_0_75])).
cnf(c_0_80,negated_conjecture,(v3_borsuk_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),esk4_0,esk5_0)|m1_subset_1(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_26]), c_0_27]), c_0_41])]), c_0_29]), c_0_30]), c_0_45]), c_0_66]), c_0_46]), c_0_56])).
cnf(c_0_81,negated_conjecture,(~v3_borsuk_1(esk2_3(esk4_0,esk5_0,X1),esk4_0,esk5_0)|~v5_pre_topc(esk2_3(esk4_0,esk5_0,X1),esk4_0,esk5_0)|~v1_funct_2(esk2_3(esk4_0,esk5_0,X1),u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~v1_funct_1(esk2_3(esk4_0,esk5_0,X1))|~m1_subset_1(esk1_3(esk4_0,esk5_0,X1),u1_struct_0(esk5_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_73]), c_0_26]), c_0_27]), c_0_28])]), c_0_30]), c_0_29])).
cnf(c_0_82,plain,(v5_pre_topc(esk2_3(X1,X2,X3),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_73]), c_0_39]), c_0_74]), c_0_75])).
cnf(c_0_83,plain,(v3_borsuk_1(esk2_3(X1,X2,X3),X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk6_3(X1,X2,esk2_3(X1,X2,X3)),X3)|~m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_78]), c_0_79]), c_0_13]), c_0_73]), c_0_74]), c_0_75])).
cnf(c_0_84,negated_conjecture,(m1_subset_1(esk1_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk5_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_80]), c_0_66]), c_0_56])).
cnf(c_0_85,negated_conjecture,(~v3_borsuk_1(esk2_3(esk4_0,esk5_0,X1),esk4_0,esk5_0)|~v1_funct_2(esk2_3(esk4_0,esk5_0,X1),u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~v1_funct_1(esk2_3(esk4_0,esk5_0,X1))|~m1_subset_1(esk1_3(esk4_0,esk5_0,X1),u1_struct_0(esk5_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_26]), c_0_27]), c_0_28]), c_0_49])]), c_0_29]), c_0_30])).
cnf(c_0_86,plain,(v3_borsuk_1(esk2_3(X1,X2,u1_struct_0(X2)),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk1_3(X1,X2,u1_struct_0(X2)),u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_71]), c_0_13]), c_0_73]), c_0_74]), c_0_75]), c_0_22])).
cnf(c_0_87,negated_conjecture,(v1_funct_1(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_84]), c_0_26]), c_0_27]), c_0_28])]), c_0_30]), c_0_29])).
cnf(c_0_88,negated_conjecture,(~v1_funct_2(esk2_3(esk4_0,esk5_0,u1_struct_0(esk5_0)),u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_84]), c_0_26]), c_0_27]), c_0_28])]), c_0_29]), c_0_30]), c_0_87])).
cnf(c_0_89,negated_conjecture,(~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_75]), c_0_84]), c_0_26]), c_0_27]), c_0_28])]), c_0_30]), c_0_29])).
cnf(c_0_90,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_22]), c_0_26]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 91
# Proof object clause steps            : 70
# Proof object formula steps           : 21
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 125
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 214
# ...of these trivial                  : 0
# ...subsumed                          : 15
# ...remaining for further processing  : 199
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 46
# Backward-rewritten                   : 13
# Generated clauses                    : 216
# ...of the previous two non-trivial   : 209
# Contextual simplify-reflections      : 135
# Paramodulations                      : 212
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 106
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 95
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 329
# Current number of archived formulas  : 0
# Current number of archived clauses   : 89
# Clause-clause subsumption calls (NU) : 10750
# Rec. Clause-clause subsumption calls : 975
# Non-unit clause-clause subsumptions  : 195
# Unit Clause-clause subsumption calls : 78
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 26047

# -------------------------------------------------
# User time                : 0.044 s
# System time              : 0.000 s
# Total time               : 0.044 s
# Maximum resident set size: 45156 pages
