# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t70_enumset1)).
fof(t67_zfmisc_1, conjecture, (![X1]:![X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)<=>~(r2_hidden(X1,X2)))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t67_zfmisc_1)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t95_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t71_enumset1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t100_xboole_1)).
fof(commutativity_k5_xboole_0, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', commutativity_k5_xboole_0)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t91_xboole_1)).
fof(t125_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t125_enumset1)).
fof(t92_xboole_1, axiom, (![X1]:k5_xboole_0(X1,X1)=k1_xboole_0), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t92_xboole_1)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t5_boole)).
fof(t79_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t79_xboole_1)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t3_xboole_0)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t7_xboole_1)).
fof(t1_boole, axiom, (![X1]:k2_xboole_0(X1,k1_xboole_0)=X1), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t1_boole)).
fof(l27_zfmisc_1, axiom, (![X1]:![X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', l27_zfmisc_1)).
fof(t88_xboole_1, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t88_xboole_1)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(t31_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t31_zfmisc_1)).
fof(t6_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X1,X2)), file('out/t67_zfmisc_1__33_premises__rerun__rerun__rerun.E_input', t6_xboole_1)).
fof(c_0_20, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_21, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_22, negated_conjecture, (~(![X1]:![X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)<=>~(r2_hidden(X1,X2))))), inference(assume_negation,[status(cth)],[t67_zfmisc_1])).
fof(c_0_23, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_24,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_27, negated_conjecture, (((k4_xboole_0(k1_tarski(esk1_0),esk2_0)!=k1_tarski(esk1_0)|r2_hidden(esk1_0,esk2_0))&(k4_xboole_0(k1_tarski(esk1_0),esk2_0)=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_28, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_29,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32,negated_conjecture,(k4_xboole_0(k1_tarski(esk1_0),esk2_0)=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
fof(c_0_35, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k5_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
fof(c_0_36, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_37,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k5_xboole_0(k1_tarski(esk1_0),esk2_0),k3_tarski(k2_enumset1(k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0),esk2_0))))=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_38,plain,(k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_39,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_40, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X8,X7,X6,X5)), inference(variable_rename,[status(thm)],[t125_enumset1])).
cnf(c_0_41,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0),esk2_0)))))=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_42,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_43,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(esk2_0,k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0))))))=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_44, plain, (![X2]:k5_xboole_0(X2,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_45, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_46,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k3_tarski(k2_enumset1(esk2_0,k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0))),X1))))=k5_xboole_0(k1_tarski(esk1_0),X1)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_43]), c_0_39]), c_0_39])).
cnf(c_0_47,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_48,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_49, plain, (![X3]:![X4]:r1_xboole_0(k4_xboole_0(X3,X4),X4)), inference(variable_rename,[status(thm)],[t79_xboole_1])).
cnf(c_0_50,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(esk2_0,k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0))))=k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k1_tarski(esk1_0)))|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_51,plain,(r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_52,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k1_tarski(esk1_0)))))=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_50])).
fof(c_0_53, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk3_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk3_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
cnf(c_0_54,plain,(r1_xboole_0(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2)))),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_33]), c_0_34])).
cnf(c_0_55,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),X1)))))=k5_xboole_0(k1_tarski(esk1_0),X1)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_52]), c_0_39]), c_0_39]), c_0_39])).
cnf(c_0_56,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_57,plain,(r1_xboole_0(k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X2))))),X2)), inference(rw,[status(thm)],[c_0_54, c_0_39])).
cnf(c_0_58,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_48, c_0_38])).
cnf(c_0_59,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(esk2_0,k1_tarski(esk1_0))))=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_47]), c_0_48]), c_0_38])).
fof(c_0_60, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_61, plain, (![X2]:k2_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_62,plain,(~r2_hidden(X1,k5_xboole_0(X2,k5_xboole_0(X2,k5_xboole_0(X3,k3_tarski(k2_enumset1(X2,X2,X2,X3))))))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_47]), c_0_58])).
cnf(c_0_64,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),X1))))=X1|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_59]), c_0_58]), c_0_39]), c_0_39])).
fof(c_0_65, plain, (![X3]:![X4]:(r2_hidden(X3,X4)|r1_xboole_0(k1_tarski(X3),X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_66,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_68, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|k4_xboole_0(k2_xboole_0(X3,X4),X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t88_xboole_1])])).
fof(c_0_69, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_70,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|k4_xboole_0(k1_tarski(esk1_0),esk2_0)!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_71,plain,(~r2_hidden(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X3,X3,X3,X2))))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_72,negated_conjecture,(k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k1_tarski(esk1_0))))))=k3_tarski(k2_enumset1(esk2_0,k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0)))|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_64, c_0_50])).
cnf(c_0_73,plain,(r1_xboole_0(k1_tarski(X1),X2)|r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_74, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_zfmisc_1])])])])).
cnf(c_0_75,plain,(r1_tarski(X1,k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_30]), c_0_31])).
cnf(c_0_76,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_30]), c_0_31])).
cnf(c_0_77,plain,(k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_78,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
fof(c_0_79, plain, (![X3]:![X4]:k2_xboole_0(X3,k2_xboole_0(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t6_xboole_1])).
cnf(c_0_80,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k5_xboole_0(k1_tarski(esk1_0),esk2_0),k3_tarski(k2_enumset1(k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0),esk2_0))))!=k1_tarski(esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_33]), c_0_34])).
cnf(c_0_81,plain,(~r2_hidden(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(X2,X3,X3,X3))))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_71, c_0_42])).
cnf(c_0_82,negated_conjecture,(k3_tarski(k2_enumset1(esk2_0,k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0)))=k5_xboole_0(esk2_0,k1_tarski(esk1_0))|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_64, c_0_72])).
cnf(c_0_83,plain,(r2_hidden(X1,X2)|~r2_hidden(X3,k1_tarski(X1))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_56, c_0_73])).
cnf(c_0_84,plain,(r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_85,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_86,plain,(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k5_xboole_0(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),X2),k3_tarski(k2_enumset1(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,X2)),X2))))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_30]), c_0_33]), c_0_31]), c_0_31]), c_0_34])).
cnf(c_0_87,plain,(k2_enumset1(X1,X1,X1,X2)=k2_enumset1(X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_25]), c_0_25]), c_0_31]), c_0_31])).
cnf(c_0_88,plain,(k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_89,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0),esk2_0)))))!=k1_tarski(esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_38]), c_0_39])).
cnf(c_0_90,negated_conjecture,(~r2_hidden(X1,k1_tarski(esk1_0))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_63]), c_0_83])).
cnf(c_0_91,plain,(r2_hidden(X1,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92,plain,(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k5_xboole_0(X2,k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X2,X2,X2,k3_tarski(k2_enumset1(X1,X1,X1,X2)))))))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_38]), c_0_87]), c_0_39])).
cnf(c_0_93,plain,(k5_xboole_0(X1,k5_xboole_0(X2,X3))=k5_xboole_0(X3,k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_94,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X1,X1,X1,X2))))=k3_tarski(k2_enumset1(X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_30]), c_0_30]), c_0_30]), c_0_31]), c_0_31]), c_0_31])).
cnf(c_0_95,negated_conjecture,(r2_hidden(esk1_0,esk2_0)|k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(esk2_0,k5_xboole_0(k1_tarski(esk1_0),k3_tarski(k2_enumset1(esk2_0,k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0))))))!=k1_tarski(esk1_0)), inference(rw,[status(thm)],[c_0_89, c_0_42])).
cnf(c_0_96,plain,(k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_63, c_0_38])).
cnf(c_0_97,negated_conjecture,(~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_98,plain,(k5_xboole_0(X1,k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X1)))))=X2|~r1_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93]), c_0_39]), c_0_38]), c_0_63])).
cnf(c_0_99,plain,(k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X1))))=k3_tarski(k2_enumset1(X2,X2,X2,X1))), inference(spm,[status(thm)],[c_0_94, c_0_87])).
cnf(c_0_100,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_101,negated_conjecture,(k5_xboole_0(esk2_0,k3_tarski(k2_enumset1(esk2_0,k1_tarski(esk1_0),k1_tarski(esk1_0),k1_tarski(esk1_0))))!=k1_tarski(esk1_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_93]), c_0_39]), c_0_96]), c_0_97])).
cnf(c_0_102,plain,(k2_enumset1(X1,X1,X1,X2)=k2_enumset1(X1,X2,X2,X2)), inference(spm,[status(thm)],[c_0_42, c_0_87])).
cnf(c_0_103,plain,(k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X1)))=X2|~r1_xboole_0(X2,X1)), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_104,negated_conjecture,(r1_xboole_0(X1,k1_tarski(esk1_0))|~r2_hidden(esk3_2(X1,k1_tarski(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_90, c_0_100])).
cnf(c_0_105,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk3_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_106,negated_conjecture,(k5_xboole_0(esk2_0,k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k1_tarski(esk1_0))))!=k1_tarski(esk1_0)), inference(rw,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_107,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X2))=k5_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_63, c_0_103])).
cnf(c_0_108,negated_conjecture,(r1_xboole_0(esk2_0,k1_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_109,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_38]), c_0_63]), c_0_108])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 69
# Proof object formula steps           : 41
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 67
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 19
# Processed clauses                    : 1349
# ...of these trivial                  : 93
# ...subsumed                          : 951
# ...remaining for further processing  : 305
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 24
# Generated clauses                    : 3555
# ...of the previous two non-trivial   : 2952
# Contextual simplify-reflections      : 130
# Paramodulations                      : 3555
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 261
#    Positive orientable unit clauses  : 40
#    Positive unorientable unit clauses: 7
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 211
# Current number of unprocessed clauses: 1541
# ...number of literals in the above   : 4432
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 29530
# Rec. Clause-clause subsumption calls : 23884
# Non-unit clause-clause subsumptions  : 781
# Unit Clause-clause subsumption calls : 106
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 307
# BW rewrite match successes           : 139
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 49890

# -------------------------------------------------
# User time                : 0.072 s
# System time              : 0.008 s
# Total time               : 0.080 s
# Maximum resident set size: 45348 pages
