# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t69_yellow_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v3_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2))=>k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3)))))), file('out/t69_yellow_0__2_premises__rerun.E_input', t69_yellow_0)).
fof(t17_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(r1_yellow_0(X1,X2)&r2_yellow_0(X1,X2)))), file('out/t69_yellow_0__2_premises__rerun.E_input', t17_yellow_0)).
fof(t65_yellow_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>((r1_yellow_0(X1,X3)&r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2)))=>(r1_yellow_0(X2,X3)&k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3))))))), file('out/t69_yellow_0__2_premises__rerun.E_input', t65_yellow_0)).
fof(c_0_3, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v3_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2))=>k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3))))))), inference(assume_negation,[status(cth)],[t69_yellow_0])).
fof(c_0_4, plain, (![X3]:![X4]:![X4]:((r1_yellow_0(X3,X4)|(v2_struct_0(X3)|(~v5_orders_2(X3)|(~v3_lattice3(X3)|~l1_orders_2(X3)))))&(r2_yellow_0(X3,X4)|(v2_struct_0(X3)|(~v5_orders_2(X3)|(~v3_lattice3(X3)|~l1_orders_2(X3))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_yellow_0])])])])])])])).
fof(c_0_5, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v3_lattice3(esk1_0)&l1_orders_2(esk1_0))))))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(r2_hidden(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk2_0))&k1_yellow_0(esk2_0,esk3_0)!=k1_yellow_0(esk1_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:(((((r1_yellow_0(X5,X6)|(~r1_yellow_0(X4,X6)|~r2_hidden(k1_yellow_0(X4,X6),u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v4_yellow_0(X5,X4)|~m1_yellow_0(X5,X4))))|(v2_struct_0(X4)|(~v4_orders_2(X4)|~l1_orders_2(X4))))&((((k1_yellow_0(X5,X6)=k1_yellow_0(X4,X6)|(~r1_yellow_0(X4,X6)|~r2_hidden(k1_yellow_0(X4,X6),u1_struct_0(X5))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X5)|(~v4_yellow_0(X5,X4)|~m1_yellow_0(X5,X4))))|(v2_struct_0(X4)|(~v4_orders_2(X4)|~l1_orders_2(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_yellow_0])])])])])])])).
cnf(c_0_7,plain,(v2_struct_0(X1)|r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,negated_conjecture,(v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,plain,(v2_struct_0(X1)|v2_struct_0(X2)|k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3)|~l1_orders_2(X1)|~v4_orders_2(X1)|~m1_yellow_0(X2,X1)|~v4_yellow_0(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2))|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,negated_conjecture,(r2_hidden(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,negated_conjecture,(r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])]), c_0_11])).
cnf(c_0_15,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16,negated_conjecture,(m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17,negated_conjecture,(v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19,negated_conjecture,(k1_yellow_0(esk2_0,esk3_0)!=k1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_20,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_21,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_9]), c_0_18])]), c_0_19]), c_0_20]), c_0_11]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 22
# Proof object clause steps            : 15
# Proof object formula steps           : 7
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 34
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 34
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 3
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 0
# Paramodulations                      : 3
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 18
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 4
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 2
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1220

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 45312 pages
