# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_yellow_0, conjecture, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2))))), file('out/t6_yellow_0__7_premises__rerun.E_input', t6_yellow_0)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t6_yellow_0__7_premises__rerun.E_input', d9_lattice3)).
fof(d8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4))))))), file('out/t6_yellow_0__7_premises__rerun.E_input', d8_lattice3)).
fof(l3_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1)))), file('out/t6_yellow_0__7_premises__rerun.E_input', l3_subset_1)).
fof(t4_subset_1, axiom, (![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), file('out/t6_yellow_0__7_premises__rerun.E_input', t4_subset_1)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t6_yellow_0__7_premises__rerun.E_input', d1_tarski)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t6_yellow_0__7_premises__rerun.E_input', t4_subset)).
fof(d9_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))))))), file('out/t6_yellow_0__7_premises__rerun.E_input', d9_orders_2)).
fof(c_0_8, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))))), inference(assume_negation,[status(cth)],[t6_yellow_0])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk3_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk3_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk3_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
fof(c_0_10, negated_conjecture, ((l1_orders_2(esk1_0)&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~r2_lattice3(esk1_0,k1_xboole_0,esk2_0)|~r1_lattice3(esk1_0,k1_xboole_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X8]:((((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk4_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,X7,esk4_3(X5,X6,X7))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X5,k1_zfmisc_1(X4))|(~r2_hidden(X6,X5)|r2_hidden(X6,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])])])).
fof(c_0_13, plain, (![X2]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_14,plain,(r2_lattice3(X1,X3,X2)|r2_hidden(esk3_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(r1_lattice3(X1,X3,X2)|r2_hidden(esk4_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk5_2(X4,X5),X5)|esk5_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk5_2(X4,X5),X5)|esk5_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_18,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(r2_hidden(esk3_3(esk1_0,X1,X2),X1)|r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22,negated_conjecture,(r2_hidden(esk4_3(esk1_0,X1,X2),X1)|r1_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_16, c_0_15])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_24,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26,negated_conjecture,(r2_hidden(esk3_3(esk1_0,X1,esk2_0),X1)|r2_lattice3(esk1_0,X1,esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,negated_conjecture,(r2_hidden(esk4_3(esk1_0,X1,esk2_0),X1)|r1_lattice3(esk1_0,X1,esk2_0)), inference(spm,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_28,plain,(r1_lattice3(X1,X3,X2)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_29,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_31,negated_conjecture,(r2_hidden(esk3_3(esk1_0,k1_xboole_0,esk2_0),X1)|r2_lattice3(esk1_0,k1_xboole_0,esk2_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32,negated_conjecture,(r2_hidden(esk4_3(esk1_0,k1_xboole_0,esk2_0),X1)|r1_lattice3(esk1_0,k1_xboole_0,esk2_0)), inference(spm,[status(thm)],[c_0_25, c_0_27])).
cnf(c_0_33,negated_conjecture,(r1_lattice3(esk1_0,X1,X2)|m1_subset_1(esk4_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_15])).
cnf(c_0_34,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_29, c_0_19])).
cnf(c_0_35,negated_conjecture,(X1=esk3_3(esk1_0,k1_xboole_0,esk2_0)|r2_lattice3(esk1_0,k1_xboole_0,esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36,negated_conjecture,(X1=esk4_3(esk1_0,k1_xboole_0,esk2_0)|r1_lattice3(esk1_0,k1_xboole_0,esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_32])).
cnf(c_0_37,negated_conjecture,(r1_lattice3(esk1_0,X1,esk2_0)|m1_subset_1(esk4_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_38,negated_conjecture,(r2_lattice3(esk1_0,k1_xboole_0,esk2_0)|m1_subset_1(esk3_3(esk1_0,k1_xboole_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_34, c_0_26])).
cnf(c_0_39,negated_conjecture,(X1=X2|r2_lattice3(esk1_0,k1_xboole_0,esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_35])).
cnf(c_0_40,negated_conjecture,(~r1_lattice3(esk1_0,k1_xboole_0,esk2_0)|~r2_lattice3(esk1_0,k1_xboole_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_41,negated_conjecture,(X1=X2|r1_lattice3(esk1_0,k1_xboole_0,esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_36])).
cnf(c_0_42,negated_conjecture,(r1_lattice3(esk1_0,k1_xboole_0,esk2_0)|m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_43,negated_conjecture,(r2_lattice3(esk1_0,k1_xboole_0,esk2_0)|m1_subset_1(X1,X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44,negated_conjecture,(X1=X2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_39])).
cnf(c_0_45,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_42]), c_0_43])).
cnf(c_0_46,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_47,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_48,negated_conjecture,(l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_15, c_0_44])).
cnf(c_0_49,negated_conjecture,(m1_subset_1(X1,X2)), inference(spm,[status(thm)],[c_0_45, c_0_44])).
fof(c_0_50, plain, (![X4]:![X5]:![X6]:(((((~r1_orders_2(X4,X5,X6)|r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4)))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4))&((((~r2_hidden(k4_tarski(X5,X6),u1_orders_2(X4))|r1_orders_2(X4,X5,X6))|~m1_subset_1(X6,u1_struct_0(X4)))|~m1_subset_1(X5,u1_struct_0(X4)))|~l1_orders_2(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])])])).
cnf(c_0_51,plain,(r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_52,negated_conjecture,(~r1_lattice3(esk1_0,k1_xboole_0,X1)|~r2_lattice3(esk1_0,k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_40, c_0_44])).
cnf(c_0_53,plain,(r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,esk4_3(X1,X2,X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])]), c_0_49])])).
cnf(c_0_54,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_55,plain,(r2_hidden(X1,X2)), inference(sr,[status(thm)],[c_0_51, c_0_44])).
cnf(c_0_56,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk3_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_57,negated_conjecture,(~r1_lattice3(esk1_0,X1,X2)|~r2_lattice3(esk1_0,X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_44])).
cnf(c_0_58,negated_conjecture,(r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,X4)), inference(spm,[status(thm)],[c_0_53, c_0_44])).
cnf(c_0_59,plain,(r1_orders_2(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_48])]), c_0_55]), c_0_49]), c_0_49])])).
cnf(c_0_60,plain,(r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk3_3(X1,X2,X3),X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_48])]), c_0_49])])).
cnf(c_0_61,negated_conjecture,(~r1_lattice3(X1,X2,X3)|~r2_lattice3(X1,X2,X3)), inference(spm,[status(thm)],[c_0_57, c_0_44])).
cnf(c_0_62,negated_conjecture,(r1_lattice3(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_63,negated_conjecture,(r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X4,X3)), inference(spm,[status(thm)],[c_0_60, c_0_44])).
cnf(c_0_64,negated_conjecture,(~r2_lattice3(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_65,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_59])]), c_0_64]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 49
# Proof object formula steps           : 17
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 192
# ...of these trivial                  : 8
# ...subsumed                          : 79
# ...remaining for further processing  : 104
# Other redundant clauses eliminated   : 20
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 57
# Backward-rewritten                   : 38
# Generated clauses                    : 1003
# ...of the previous two non-trivial   : 964
# Contextual simplify-reflections      : 45
# Paramodulations                      : 976
# Factorizations                       : 2
# Equation resolutions                 : 24
# Current number of processed clauses  : 7
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 10
# Current number of archived formulas  : 0
# Current number of archived clauses   : 96
# Clause-clause subsumption calls (NU) : 1467
# Rec. Clause-clause subsumption calls : 891
# Non-unit clause-clause subsumptions  : 162
# Unit Clause-clause subsumption calls : 109
# Rewrite failures with RHS unbound    : 82
# BW rewrite match attempts            : 77
# BW rewrite match successes           : 76
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11613

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 45328 pages
