# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t70_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)<=>(~(r2_hidden(X1,X3))&(r2_hidden(X2,X3)|X1=X2)))), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t70_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t71_enumset1)).
fof(t60_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(r2_hidden(X1,X2)=>((r2_hidden(X3,X2)&~(X1=X3))|k3_xboole_0(k2_tarski(X1,X3),X2)=k1_tarski(X1)))), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t60_zfmisc_1)).
fof(d1_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3))))))), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', d1_enumset1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t112_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t112_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X1)=X1), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', idempotence_k3_xboole_0)).
fof(commutativity_k5_xboole_0, axiom, (![X1]:![X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', commutativity_k5_xboole_0)).
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', d5_xboole_0)).
fof(t59_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((k3_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)&(r2_hidden(X2,X3)&~(X1=X2))))), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t59_zfmisc_1)).
fof(t1_xboole_0, axiom, (![X1]:![X2]:![X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3))))), file('out/t70_zfmisc_1__16_premises__rerun__rerun.E_input', t1_xboole_0)).
fof(c_0_13, negated_conjecture, (~(![X1]:![X2]:![X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)<=>(~(r2_hidden(X1,X3))&(r2_hidden(X2,X3)|X1=X2))))), inference(assume_negation,[status(cth)],[t70_zfmisc_1])).
fof(c_0_14, negated_conjecture, (((((~r2_hidden(esk2_0,esk3_0)|r2_hidden(esk1_0,esk3_0))|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk1_0))&((esk1_0!=esk2_0|r2_hidden(esk1_0,esk3_0))|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk1_0)))&((~r2_hidden(esk1_0,esk3_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_tarski(esk1_0))&((r2_hidden(esk2_0,esk3_0)|esk1_0=esk2_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_tarski(esk1_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])])).
fof(c_0_15, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:(((r2_hidden(X6,X5)|k3_xboole_0(k2_tarski(X4,X6),X5)=k1_tarski(X4))|~r2_hidden(X4,X5))&((X4!=X6|k3_xboole_0(k2_tarski(X4,X6),X5)=k1_tarski(X4))|~r2_hidden(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t60_zfmisc_1])])])])])])).
fof(c_0_19, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X10]:![X6]:![X7]:![X8]:![X9]:((((~r2_hidden(X10,X9)|(X10=X6|(X10=X7|X10=X8)))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X6|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X7|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&((X10!=X8|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8)))))&((((esk4_4(X6,X7,X8,X9)!=X6|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&(((esk4_4(X6,X7,X8,X9)!=X7|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&((esk4_4(X6,X7,X8,X9)!=X8|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))))&((r2_hidden(esk4_4(X6,X7,X8,X9),X9)|(esk4_4(X6,X7,X8,X9)=X6|(esk4_4(X6,X7,X8,X9)=X7|esk4_4(X6,X7,X8,X9)=X8)))|X9=k1_enumset1(X6,X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])])).
cnf(c_0_20,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk1_0)|~r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k3_xboole_0(X4,X5),k3_xboole_0(X6,X5))=k3_xboole_0(k5_xboole_0(X4,X6),X5)), inference(variable_rename,[status(thm)],[t112_xboole_1])).
fof(c_0_26, plain, (![X3]:k3_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_27, plain, (![X3]:![X4]:k5_xboole_0(X3,X4)=k5_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_28,plain,(k3_xboole_0(k2_tarski(X1,X3),X2)=k1_tarski(X1)|~r2_hidden(X1,X2)|X1!=X3), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X2), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),esk3_0))!=k1_tarski(esk1_0)|~r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), c_0_23])).
cnf(c_0_31,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,plain,(k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,plain,(k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35,plain,(k3_xboole_0(k2_enumset1(X1,X1,X1,X3),X2)=k1_tarski(X1)|X1!=X3|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_21]), c_0_23])).
cnf(c_0_36,plain,(r2_hidden(X5,X1)|X2!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_29, c_0_23])).
fof(c_0_37, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk5_3(X5,X6,X7),X7)|(~r2_hidden(esk5_3(X5,X6,X7),X5)|r2_hidden(esk5_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk5_3(X5,X6,X7),X5)|r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk5_3(X5,X6,X7),X6)|r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
cnf(c_0_38,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))!=k1_tarski(esk1_0)|~r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39,plain,(k5_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,k5_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_31])).
cnf(c_0_40,plain,(k3_xboole_0(k2_enumset1(X1,X1,X1,X1),X2)=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_41,plain,(r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_36])])).
cnf(c_0_42,negated_conjecture,(k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_tarski(esk1_0)|esk1_0=esk2_0|r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_43,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))!=k1_tarski(esk1_0)|~r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45,plain,(k1_tarski(X1)=k2_enumset1(X1,X1,X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_40]), c_0_41])])).
cnf(c_0_46,plain,(k3_xboole_0(k2_tarski(X1,X3),X2)=k1_tarski(X1)|r2_hidden(X3,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_47,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X4), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_48, plain, (![X4]:![X5]:![X6]:(k3_xboole_0(k2_tarski(X4,X5),X6)!=k1_tarski(X4)|(~r2_hidden(X5,X6)|X4=X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t59_zfmisc_1])])])).
cnf(c_0_49,negated_conjecture,(esk2_0=esk1_0|k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),esk3_0))=k1_tarski(esk1_0)|r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_21]), c_0_22]), c_0_23]), c_0_23])).
cnf(c_0_50,plain,(X1!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X4,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[c_0_43, c_0_22])).
cnf(c_0_51,negated_conjecture,(k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_52, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:((((~r2_hidden(X4,X5)|~r2_hidden(X4,X6))|~r2_hidden(X4,k5_xboole_0(X5,X6)))&((r2_hidden(X4,X5)|r2_hidden(X4,X6))|~r2_hidden(X4,k5_xboole_0(X5,X6))))&(((~r2_hidden(X4,X5)|r2_hidden(X4,X6))|r2_hidden(X4,k5_xboole_0(X5,X6)))&((~r2_hidden(X4,X6)|r2_hidden(X4,X5))|r2_hidden(X4,k5_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])])])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|~r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54,plain,(k3_xboole_0(k2_enumset1(X1,X1,X1,X3),X2)=k1_tarski(X1)|r2_hidden(X3,X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_21]), c_0_23])).
cnf(c_0_55,plain,(r2_hidden(X5,X1)|X4!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_47, c_0_23])).
cnf(c_0_56,plain,(X1=X2|~r2_hidden(X2,X3)|k3_xboole_0(k2_tarski(X1,X2),X3)!=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57,negated_conjecture,(k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))=k1_tarski(esk1_0)|esk2_0=esk1_0|r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_49, c_0_31])).
cnf(c_0_58,plain,(~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_59,negated_conjecture,(k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),esk3_0))=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_21]), c_0_22]), c_0_23]), c_0_23])).
cnf(c_0_60,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk1_0)|esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_61,plain,(~r2_hidden(X1,k5_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62,negated_conjecture,(r2_hidden(esk2_0,k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))|r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk1_0,k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))|~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_45])])).
cnf(c_0_63,plain,(r2_hidden(X1,k2_enumset1(X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_55])])).
cnf(c_0_64,plain,(X1=X2|k3_xboole_0(k2_enumset1(X1,X1,X1,X2),X3)!=k1_tarski(X1)|~r2_hidden(X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_21]), c_0_23])).
cnf(c_0_65,negated_conjecture,(k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))=k1_tarski(esk1_0)|esk2_0=esk1_0|r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_57, c_0_39])).
cnf(c_0_66,plain,(~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X3,X2)))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_58, c_0_31])).
cnf(c_0_67,negated_conjecture,(k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk3_0)), inference(rw,[status(thm)],[c_0_59, c_0_31])).
cnf(c_0_68,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|esk2_0!=esk1_0|k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),esk3_0))!=k1_tarski(esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_21]), c_0_22]), c_0_23]), c_0_23])).
cnf(c_0_69,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk1_0,k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))|~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])])).
cnf(c_0_70,plain,(r2_hidden(X1,k5_xboole_0(X2,X3))|r2_hidden(X1,X2)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_71,negated_conjecture,(esk2_0=esk1_0|r2_hidden(esk2_0,esk3_0)|~r2_hidden(esk2_0,k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_72,negated_conjecture,(~r2_hidden(X1,k1_tarski(esk1_0))|~r2_hidden(esk1_0,esk3_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_73,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))!=k1_tarski(esk1_0)|esk2_0!=esk1_0), inference(rw,[status(thm)],[c_0_68, c_0_31])).
cnf(c_0_74,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_41])])).
cnf(c_0_75,negated_conjecture,(esk2_0=esk1_0|r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_70]), c_0_63])])).
cnf(c_0_76,negated_conjecture,(~r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(esk1_0,esk3_0)|~r2_hidden(X1,esk3_0)), inference(rw,[status(thm)],[c_0_72, c_0_45])).
cnf(c_0_77,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))!=k1_tarski(esk1_0)|esk2_0!=esk1_0), inference(rw,[status(thm)],[c_0_73, c_0_39])).
cnf(c_0_78,negated_conjecture,(esk2_0=esk1_0|r2_hidden(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_79,negated_conjecture,(~r2_hidden(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_63])).
cnf(c_0_80,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|esk2_0!=esk1_0), inference(rw,[status(thm)],[c_0_77, c_0_45])).
cnf(c_0_81,negated_conjecture,(esk2_0=esk1_0), inference(sr,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_82,negated_conjecture,(k3_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_81]), c_0_81]), c_0_81])]), c_0_79])).
cnf(c_0_83,negated_conjecture,(~r2_hidden(esk1_0,k5_xboole_0(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_40]), c_0_45])])).
cnf(c_0_84,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_70]), c_0_41])]), c_0_79]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 85
# Proof object clause steps            : 58
# Proof object formula steps           : 27
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 64
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 29
# Processed clauses                    : 168
# ...of these trivial                  : 2
# ...subsumed                          : 51
# ...remaining for further processing  : 115
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 26
# Generated clauses                    : 298
# ...of the previous two non-trivial   : 274
# Contextual simplify-reflections      : 3
# Paramodulations                      : 289
# Factorizations                       : 0
# Equation resolutions                 : 11
# Current number of processed clauses  : 40
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 18
# Current number of unprocessed clauses: 132
# ...number of literals in the above   : 296
# Current number of archived formulas  : 0
# Current number of archived clauses   : 70
# Clause-clause subsumption calls (NU) : 844
# Rec. Clause-clause subsumption calls : 531
# Non-unit clause-clause subsumptions  : 30
# Unit Clause-clause subsumption calls : 156
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 41
# BW rewrite match successes           : 24
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5590

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 43024 pages
