# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_boole, axiom, (![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t2_boole)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t48_xboole_1)).
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t4_xboole_0)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t3_boole)).
fof(d7_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0)), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', d7_xboole_0)).
fof(t70_xboole_1, conjecture, (![X1]:![X2]:![X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3)))))), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t70_xboole_1)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t3_xboole_0)).
fof(d3_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2))))), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', d3_xboole_0)).
fof(t47_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t47_xboole_1)).
fof(t41_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', t41_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), file('out/t70_xboole_1__11_premises__rerun__rerun.E_input', commutativity_k2_xboole_0)).
fof(c_0_11, plain, (![X2]:k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_12, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_13, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk9_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
cnf(c_0_14,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_17, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k3_xboole_0(X3,X4)=k1_xboole_0)&(k3_xboole_0(X3,X4)!=k1_xboole_0|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])])])).
fof(c_0_18, negated_conjecture, (~(![X1]:![X2]:![X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))))), inference(assume_negation,[status(cth)],[t70_xboole_1])).
cnf(c_0_19,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk2_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk2_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_24, negated_conjecture, (((((~r1_xboole_0(esk6_0,esk7_0)|~r1_xboole_0(esk6_0,esk8_0))|~r1_xboole_0(esk3_0,k2_xboole_0(esk4_0,esk5_0)))&(r1_xboole_0(esk6_0,k2_xboole_0(esk7_0,esk8_0))|~r1_xboole_0(esk3_0,k2_xboole_0(esk4_0,esk5_0))))&((((~r1_xboole_0(esk6_0,esk7_0)|~r1_xboole_0(esk6_0,esk8_0))|r1_xboole_0(esk3_0,esk4_0))&(r1_xboole_0(esk6_0,k2_xboole_0(esk7_0,esk8_0))|r1_xboole_0(esk3_0,esk4_0)))&(((~r1_xboole_0(esk6_0,esk7_0)|~r1_xboole_0(esk6_0,esk8_0))|r1_xboole_0(esk3_0,esk5_0))&(r1_xboole_0(esk6_0,k2_xboole_0(esk7_0,esk8_0))|r1_xboole_0(esk3_0,esk5_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])])])).
fof(c_0_25, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(r2_hidden(X8,X5)|r2_hidden(X8,X6)))|X7!=k2_xboole_0(X5,X6))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X7))|X7!=k2_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|r2_hidden(X8,X7))|X7!=k2_xboole_0(X5,X6))))&((((~r2_hidden(esk1_3(X5,X6,X7),X5)|~r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k2_xboole_0(X5,X6))&((~r2_hidden(esk1_3(X5,X6,X7),X6)|~r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k2_xboole_0(X5,X6)))&((r2_hidden(esk1_3(X5,X6,X7),X7)|(r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k2_xboole_0(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])])).
fof(c_0_26, plain, (![X3]:![X4]:k4_xboole_0(X3,k3_xboole_0(X3,X4))=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_27,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_19, c_0_15])).
cnf(c_0_28,plain,(k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29,plain,(r1_xboole_0(X1,X2)|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_22, c_0_15])).
cnf(c_0_30,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)|r1_xboole_0(esk6_0,k2_xboole_0(esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(r2_hidden(X4,X1)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,plain,(r2_hidden(X4,X1)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35,plain,(~r1_xboole_0(X1,k1_xboole_0)|~r2_hidden(X2,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_21]), c_0_28])).
cnf(c_0_36,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_21]), c_0_28])])).
cnf(c_0_37,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)|~r2_hidden(X1,k2_xboole_0(esk7_0,esk8_0))|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38,plain,(r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_39,plain,(r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_40,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_34, c_0_15])).
cnf(c_0_41,plain,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_42,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_43,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)|~r1_xboole_0(esk6_0,esk8_0)|~r1_xboole_0(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_44,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)|~r2_hidden(X1,esk6_0)|~r2_hidden(X1,esk8_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_46,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)|~r2_hidden(X1,esk6_0)|~r2_hidden(X1,esk7_0)), inference(spm,[status(thm)],[c_0_37, c_0_39])).
cnf(c_0_47,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)|r1_xboole_0(esk6_0,k2_xboole_0(esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_48,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_49,plain,(~r1_xboole_0(X1,k4_xboole_0(X1,X2))|~r2_hidden(X3,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_27, c_0_40])).
cnf(c_0_50,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51,negated_conjecture,(~r1_xboole_0(esk6_0,esk7_0)|~r1_xboole_0(esk6_0,esk8_0)|~r2_hidden(X1,esk5_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_43])).
cnf(c_0_52,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)|r1_xboole_0(X1,esk8_0)|~r2_hidden(esk2_2(X1,esk8_0),esk6_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)|r1_xboole_0(X1,esk7_0)|~r2_hidden(esk2_2(X1,esk7_0),esk6_0)), inference(spm,[status(thm)],[c_0_46, c_0_45])).
cnf(c_0_54,plain,(r2_hidden(X4,X3)|r2_hidden(X4,X2)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_55,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)|~r2_hidden(X1,k2_xboole_0(esk7_0,esk8_0))|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_30, c_0_47])).
fof(c_0_56, plain, (![X4]:![X5]:![X6]:k4_xboole_0(k4_xboole_0(X4,X5),X6)=k4_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_57, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k2_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_58,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_48, c_0_15])).
cnf(c_0_59,plain,(~r2_hidden(X1,k4_xboole_0(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_60,negated_conjecture,(r1_xboole_0(X1,esk5_0)|~r1_xboole_0(esk6_0,esk7_0)|~r1_xboole_0(esk6_0,esk8_0)|~r2_hidden(esk2_2(X1,esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_45])).
cnf(c_0_61,negated_conjecture,(r1_xboole_0(esk6_0,esk8_0)|r1_xboole_0(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_52, c_0_42])).
cnf(c_0_62,negated_conjecture,(r1_xboole_0(esk6_0,esk7_0)|r1_xboole_0(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_53, c_0_42])).
cnf(c_0_63,plain,(r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_64,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)|~r2_hidden(X1,esk6_0)|~r2_hidden(X1,esk8_0)), inference(spm,[status(thm)],[c_0_55, c_0_38])).
cnf(c_0_65,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)|~r2_hidden(X1,esk6_0)|~r2_hidden(X1,esk7_0)), inference(spm,[status(thm)],[c_0_55, c_0_39])).
cnf(c_0_66,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_67,plain,(k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_68,plain,(X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_28]), c_0_21])).
cnf(c_0_69,plain,(r1_xboole_0(X1,k4_xboole_0(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_59, c_0_45])).
cnf(c_0_70,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_42]), c_0_61]), c_0_62])).
cnf(c_0_71,plain,(r1_xboole_0(X1,k2_xboole_0(X2,X3))|r2_hidden(esk2_2(X1,k2_xboole_0(X2,X3)),X3)|r2_hidden(esk2_2(X1,k2_xboole_0(X2,X3)),X2)), inference(spm,[status(thm)],[c_0_63, c_0_45])).
cnf(c_0_72,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)|r1_xboole_0(X1,esk8_0)|~r2_hidden(esk2_2(X1,esk8_0),esk6_0)), inference(spm,[status(thm)],[c_0_64, c_0_45])).
cnf(c_0_73,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)|r1_xboole_0(X1,esk7_0)|~r2_hidden(esk2_2(X1,esk7_0),esk6_0)), inference(spm,[status(thm)],[c_0_65, c_0_45])).
cnf(c_0_74,plain,(X1=k2_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X3)|r2_hidden(esk1_3(X2,X3,X1),X2)|r2_hidden(esk1_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_75,plain,(k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_58]), c_0_21])).
cnf(c_0_76,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k4_xboole_0(X1,X3),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_66])).
cnf(c_0_77,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_78,negated_conjecture,(~r2_hidden(X1,esk5_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_70])).
cnf(c_0_79,plain,(r1_xboole_0(X1,k2_xboole_0(X2,X2))|r2_hidden(esk2_2(X1,k2_xboole_0(X2,X2)),X2)), inference(ef,[status(thm)],[c_0_71])).
cnf(c_0_80,negated_conjecture,(r1_xboole_0(esk6_0,esk8_0)|r1_xboole_0(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_72, c_0_42])).
cnf(c_0_81,negated_conjecture,(r1_xboole_0(esk6_0,esk7_0)|r1_xboole_0(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_73, c_0_42])).
cnf(c_0_82,plain,(X1=k2_xboole_0(X2,X3)|~r2_hidden(esk1_3(X2,X3,X1),X1)|~r2_hidden(esk1_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_83,plain,(k2_xboole_0(X1,X2)=X2|r2_hidden(esk1_3(X1,X2,X2),X2)|r2_hidden(esk1_3(X1,X2,X2),X1)), inference(ef,[status(thm)],[c_0_74])).
cnf(c_0_84,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X3)=X1|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_66, c_0_75])).
cnf(c_0_85,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_28]), c_0_77])).
cnf(c_0_86,negated_conjecture,(r1_xboole_0(X1,k2_xboole_0(esk5_0,esk5_0))|~r2_hidden(esk2_2(X1,k2_xboole_0(esk5_0,esk5_0)),esk3_0)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_87,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)|~r1_xboole_0(esk6_0,esk8_0)|~r1_xboole_0(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_88,negated_conjecture,(r1_xboole_0(esk6_0,esk8_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_80])).
cnf(c_0_89,negated_conjecture,(r1_xboole_0(esk6_0,esk7_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_81])).
cnf(c_0_90,plain,(r1_xboole_0(X1,k2_xboole_0(X2,X3))|k4_xboole_0(X1,k4_xboole_0(k4_xboole_0(X1,X2),X3))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_66])).
cnf(c_0_91,plain,(k2_xboole_0(X1,X2)=X2|r2_hidden(esk1_3(X1,X2,X2),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_83])).
cnf(c_0_92,plain,(k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_84]), c_0_85]), c_0_21])).
cnf(c_0_93,negated_conjecture,(r1_xboole_0(esk3_0,k2_xboole_0(esk5_0,esk5_0))), inference(spm,[status(thm)],[c_0_86, c_0_42])).
cnf(c_0_94,negated_conjecture,(~r1_xboole_0(esk6_0,esk7_0)|~r1_xboole_0(esk6_0,esk8_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_87])).
cnf(c_0_95,negated_conjecture,(r1_xboole_0(esk6_0,esk8_0)|r1_xboole_0(X1,esk4_0)|~r2_hidden(esk2_2(X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_88, c_0_45])).
cnf(c_0_96,negated_conjecture,(r1_xboole_0(esk6_0,esk7_0)|r1_xboole_0(X1,esk4_0)|~r2_hidden(esk2_2(X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_89, c_0_45])).
cnf(c_0_97,plain,(r1_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X1,X2),X3))))|k4_xboole_0(X1,k4_xboole_0(k4_xboole_0(X1,X2),X3))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_40])).
cnf(c_0_98,plain,(k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_41, c_0_91])).
cnf(c_0_99,plain,(r1_xboole_0(X1,k2_xboole_0(X2,X3))|k4_xboole_0(X1,k4_xboole_0(X1,X3))!=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_90, c_0_75])).
cnf(c_0_100,negated_conjecture,(k4_xboole_0(esk3_0,esk5_0)=esk3_0), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_101,negated_conjecture,(r1_xboole_0(X1,esk4_0)|~r2_hidden(esk2_2(X1,esk4_0),esk3_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_45]), c_0_95]), c_0_96])).
cnf(c_0_102,plain,(r1_xboole_0(X1,k2_xboole_0(X2,k1_xboole_0))|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_84]), c_0_85]), c_0_28])])).
cnf(c_0_103,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_67, c_0_98])).
cnf(c_0_104,negated_conjecture,(r1_xboole_0(esk6_0,k2_xboole_0(esk7_0,esk8_0))|~r1_xboole_0(esk3_0,k2_xboole_0(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_105,negated_conjecture,(r1_xboole_0(esk3_0,k2_xboole_0(X1,esk5_0))|~r1_xboole_0(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_28])])).
cnf(c_0_106,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_101, c_0_42])).
cnf(c_0_107,negated_conjecture,(~r1_xboole_0(esk3_0,k2_xboole_0(esk4_0,esk5_0))|~r1_xboole_0(esk6_0,esk8_0)|~r1_xboole_0(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_108,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(rw,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_109,negated_conjecture,(r1_xboole_0(esk6_0,k2_xboole_0(esk7_0,esk8_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_106])])).
cnf(c_0_110,plain,(r1_xboole_0(X1,k2_xboole_0(X2,k1_xboole_0))|~r1_xboole_0(X1,k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_102, c_0_67])).
cnf(c_0_111,negated_conjecture,(~r1_xboole_0(esk6_0,esk7_0)|~r1_xboole_0(esk6_0,esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_105]), c_0_106])])).
cnf(c_0_112,negated_conjecture,(r1_xboole_0(esk6_0,esk7_0)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_113,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X3,X2))), inference(rw,[status(thm)],[c_0_110, c_0_103])).
cnf(c_0_114,negated_conjecture,(~r1_xboole_0(esk6_0,esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_115,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_109]), c_0_114]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 116
# Proof object clause steps            : 93
# Proof object formula steps           : 23
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 57
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 11883
# ...of these trivial                  : 110
# ...subsumed                          : 10541
# ...remaining for further processing  : 1231
# Other redundant clauses eliminated   : 1035
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 131
# Backward-rewritten                   : 171
# Generated clauses                    : 100078
# ...of the previous two non-trivial   : 90293
# Contextual simplify-reflections      : 2555
# Paramodulations                      : 98601
# Factorizations                       : 442
# Equation resolutions                 : 1035
# Current number of processed clauses  : 894
#    Positive orientable unit clauses  : 67
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 52
#    Non-unit-clauses                  : 772
# Current number of unprocessed clauses: 62451
# ...number of literals in the above   : 214368
# Current number of archived formulas  : 0
# Current number of archived clauses   : 327
# Clause-clause subsumption calls (NU) : 271737
# Rec. Clause-clause subsumption calls : 193346
# Non-unit clause-clause subsumptions  : 7697
# Unit Clause-clause subsumption calls : 9783
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 179
# BW rewrite match successes           : 61
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1736939

# -------------------------------------------------
# User time                : 1.624 s
# System time              : 0.032 s
# Total time               : 1.656 s
# Maximum resident set size: 65796 pages
