# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', d5_xboole_0)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', t95_xboole_1)).
fof(t70_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)<=>(~(r2_hidden(X1,X3))&(r2_hidden(X2,X3)|X1=X2)))), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', t70_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', t71_enumset1)).
fof(d1_tarski, axiom, (![X1]:![X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1))), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', d1_tarski)).
fof(d1_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3))))))), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', d1_enumset1)).
fof(d2_tarski, axiom, (![X1]:![X2]:![X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2)))), file('out/t70_zfmisc_1__13_premises__rerun__rerun.E_input', d2_tarski)).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
fof(c_0_10, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_11, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
fof(c_0_12, negated_conjecture, (~(![X1]:![X2]:![X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)<=>(~(r2_hidden(X1,X3))&(r2_hidden(X2,X3)|X1=X2))))), inference(assume_negation,[status(cth)],[t70_zfmisc_1])).
cnf(c_0_13,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, negated_conjecture, (((((~r2_hidden(esk3_0,esk4_0)|r2_hidden(esk2_0,esk4_0))|k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)!=k1_tarski(esk2_0))&((esk2_0!=esk3_0|r2_hidden(esk2_0,esk4_0))|k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)!=k1_tarski(esk2_0)))&((~r2_hidden(esk2_0,esk4_0)|k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)=k1_tarski(esk2_0))&((r2_hidden(esk3_0,esk4_0)|esk2_0=esk3_0)|k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)=k1_tarski(esk2_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])).
fof(c_0_17, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|X6=X4)|X5!=k1_tarski(X4))&((X6!=X4|r2_hidden(X6,X5))|X5!=k1_tarski(X4)))&(((~r2_hidden(esk5_2(X4,X5),X5)|esk5_2(X4,X5)!=X4)|X5=k1_tarski(X4))&((r2_hidden(esk5_2(X4,X5),X5)|esk5_2(X4,X5)=X4)|X5=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
fof(c_0_20, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X10]:![X6]:![X7]:![X8]:![X9]:((((~r2_hidden(X10,X9)|(X10=X6|(X10=X7|X10=X8)))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X6|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X7|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&((X10!=X8|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8)))))&((((esk6_4(X6,X7,X8,X9)!=X6|~r2_hidden(esk6_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&(((esk6_4(X6,X7,X8,X9)!=X7|~r2_hidden(esk6_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&((esk6_4(X6,X7,X8,X9)!=X8|~r2_hidden(esk6_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))))&((r2_hidden(esk6_4(X6,X7,X8,X9),X9)|(esk6_4(X6,X7,X8,X9)=X6|(esk6_4(X6,X7,X8,X9)=X7|esk6_4(X6,X7,X8,X9)=X8)))|X9=k1_enumset1(X6,X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])])).
cnf(c_0_21,plain,(X1!=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3)))|~r2_hidden(X4,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
cnf(c_0_22,negated_conjecture,(k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)=k1_tarski(esk2_0)|~r2_hidden(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,plain,(r2_hidden(X3,X1)|X1!=k1_tarski(X2)|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(r2_hidden(X4,X1)|r2_hidden(X4,X3)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27,negated_conjecture,(k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)=k1_tarski(esk2_0)|esk2_0=esk3_0|r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X4), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,plain,(~r2_hidden(X1,k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3))))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_30,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k5_xboole_0(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0),k2_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0)))=k1_tarski(esk2_0)|~r2_hidden(esk2_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_14]), c_0_24]), c_0_24]), c_0_15])).
cnf(c_0_31,plain,(r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_32,plain,(r2_hidden(X4,X3)|r2_hidden(X4,X1)|X1!=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3)))|~r2_hidden(X4,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_14]), c_0_15])).
cnf(c_0_33,negated_conjecture,(esk3_0=esk2_0|k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k5_xboole_0(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0),k2_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0)))=k1_tarski(esk2_0)|r2_hidden(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_23]), c_0_14]), c_0_24]), c_0_24]), c_0_15])).
cnf(c_0_34,plain,(r2_hidden(X5,X1)|X4!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_28, c_0_24])).
cnf(c_0_35,negated_conjecture,(r2_hidden(esk2_0,esk4_0)|k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)!=k1_tarski(esk2_0)|~r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_36,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|r2_hidden(esk1_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37,negated_conjecture,(~r2_hidden(X1,k1_tarski(esk2_0))|~r2_hidden(esk2_0,esk4_0)|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38,plain,(r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_39,negated_conjecture,(esk3_0=esk2_0|r2_hidden(esk3_0,esk4_0)|r2_hidden(X1,esk4_0)|r2_hidden(X1,X2)|X2!=k1_tarski(esk2_0)|~r2_hidden(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40,plain,(r2_hidden(X1,X2)|X2!=k2_enumset1(X3,X3,X4,X1)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_41,plain,(X3=X2|X1!=k1_tarski(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42,negated_conjecture,(r2_hidden(esk2_0,esk4_0)|k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk4_0)!=k1_tarski(esk2_0)|esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_43,negated_conjecture,(r2_hidden(esk2_0,esk4_0)|k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k5_xboole_0(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0),k2_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0)))!=k1_tarski(esk2_0)|~r2_hidden(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_23]), c_0_14]), c_0_24]), c_0_24]), c_0_15])).
cnf(c_0_44,plain,(X1=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3)))|r2_hidden(esk1_3(X2,X3,X1),X2)|r2_hidden(esk1_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_14]), c_0_15])).
cnf(c_0_45,negated_conjecture,(~r2_hidden(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46,negated_conjecture,(esk3_0=esk2_0|r2_hidden(X1,k1_tarski(esk2_0))|r2_hidden(esk3_0,esk4_0)|r2_hidden(X1,esk4_0)|~r2_hidden(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_47,plain,(r2_hidden(X1,k2_enumset1(X2,X2,X3,X1))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_48,plain,(X1=X2|~r2_hidden(X2,k1_tarski(X1))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_49,negated_conjecture,(r2_hidden(esk2_0,esk4_0)|esk3_0!=esk2_0|k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k5_xboole_0(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0),k2_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0)))!=k1_tarski(esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_23]), c_0_14]), c_0_24]), c_0_24]), c_0_15])).
fof(c_0_50, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(X8=X5|X8=X6))|X7!=k2_tarski(X5,X6))&(((X8!=X5|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))&((X8!=X6|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))))&((((esk7_3(X5,X6,X7)!=X5|~r2_hidden(esk7_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6))&((esk7_3(X5,X6,X7)!=X6|~r2_hidden(esk7_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6)))&((r2_hidden(esk7_3(X5,X6,X7),X7)|(esk7_3(X5,X6,X7)=X5|esk7_3(X5,X6,X7)=X6))|X7=k2_tarski(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])])).
cnf(c_0_51,negated_conjecture,(r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,X1),k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))|r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,X1),X1)|X1!=k1_tarski(esk2_0)|~r2_hidden(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
cnf(c_0_52,negated_conjecture,(esk3_0=esk2_0|r2_hidden(esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,X1),k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))|r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,X1),X1)|X1!=k1_tarski(esk2_0)|esk3_0!=esk2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_44]), c_0_45])).
cnf(c_0_54,plain,(X4=X3|X4=X2|X1!=k2_tarski(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_55,negated_conjecture,(r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,X1),k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))|r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,X1),X1)|X1!=k1_tarski(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_56,plain,(X3=X4|X2=X4|X1!=k2_enumset1(X2,X2,X2,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_23]), c_0_24])).
cnf(c_0_57,negated_conjecture,(r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0)),k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))|r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0)),k1_tarski(esk2_0))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_58,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X2), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_59,plain,(X1=X2|X3=X2|~r2_hidden(X2,k2_enumset1(X1,X1,X1,X3))), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_60,negated_conjecture,(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0))=esk2_0|r2_hidden(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0)),k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_48, c_0_57])).
cnf(c_0_61,plain,(r2_hidden(X5,X1)|X2!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_58, c_0_24])).
cnf(c_0_62,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X3)|~r2_hidden(esk1_3(X2,X3,X1),X2)|~r2_hidden(esk1_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_63,negated_conjecture,(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0))=esk3_0|esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0))=esk2_0), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_64,plain,(r2_hidden(X1,X2)|X2!=k2_enumset1(X1,X1,X3,X4)), inference(er,[status(thm)],[c_0_61])).
cnf(c_0_65,plain,(X1=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3)))|r2_hidden(esk1_3(X2,X3,X1),X3)|~r2_hidden(esk1_3(X2,X3,X1),X2)|~r2_hidden(esk1_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_14]), c_0_15])).
cnf(c_0_66,negated_conjecture,(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0))=esk2_0|esk3_0!=esk2_0), inference(ef,[status(thm)],[c_0_63])).
cnf(c_0_67,plain,(r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_68,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k5_xboole_0(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0),k2_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0)))=k1_tarski(esk2_0)|esk3_0!=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_38])]), c_0_45])).
cnf(c_0_69,negated_conjecture,(esk3_0!=esk2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_68]), c_0_45])).
cnf(c_0_70,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|~r2_hidden(esk1_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_71,negated_conjecture,(r2_hidden(esk3_0,esk4_0)), inference(sr,[status(thm)],[c_0_52, c_0_69])).
cnf(c_0_72,plain,(X1=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X3),k2_xboole_0(X2,X3)))|r2_hidden(esk1_3(X2,X3,X1),X1)|~r2_hidden(esk1_3(X2,X3,X1),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_14]), c_0_15])).
cnf(c_0_73,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k5_xboole_0(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0),k2_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0)))!=k1_tarski(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_71])]), c_0_45])).
cnf(c_0_74,negated_conjecture,(esk1_3(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0,k1_tarski(esk2_0))=esk2_0|r2_hidden(esk3_0,k1_tarski(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_63]), c_0_71])]), c_0_73])).
cnf(c_0_75,negated_conjecture,(r2_hidden(esk3_0,k1_tarski(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_74]), c_0_67]), c_0_38])]), c_0_73]), c_0_45])).
cnf(c_0_76,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_75]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 77
# Proof object clause steps            : 58
# Proof object formula steps           : 19
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 59
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 28
# Processed clauses                    : 135
# ...of these trivial                  : 0
# ...subsumed                          : 21
# ...remaining for further processing  : 114
# Other redundant clauses eliminated   : 22
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 6
# Generated clauses                    : 732
# ...of the previous two non-trivial   : 677
# Contextual simplify-reflections      : 3
# Paramodulations                      : 681
# Factorizations                       : 9
# Equation resolutions                 : 36
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 47
# Current number of unprocessed clauses: 423
# ...number of literals in the above   : 2059
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 1669
# Rec. Clause-clause subsumption calls : 862
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 36
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 16560

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 43008 pages
