# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t71_xboole_1__17_premises__rerun.E_input', t4_xboole_0)).
fof(t2_boole, axiom, (![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), file('out/t71_xboole_1__17_premises__rerun.E_input', t2_boole)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t71_xboole_1__17_premises__rerun.E_input', t3_xboole_0)).
fof(t66_xboole_1, axiom, (![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1))))), file('out/t71_xboole_1__17_premises__rerun.E_input', t66_xboole_1)).
fof(t36_xboole_1, axiom, (![X1]:![X2]:r1_tarski(k4_xboole_0(X1,X2),X1)), file('out/t71_xboole_1__17_premises__rerun.E_input', t36_xboole_1)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t71_xboole_1__17_premises__rerun.E_input', t3_boole)).
fof(t39_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), file('out/t71_xboole_1__17_premises__rerun.E_input', t39_xboole_1)).
fof(t37_xboole_1, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('out/t71_xboole_1__17_premises__rerun.E_input', t37_xboole_1)).
fof(t18_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k3_xboole_0(X2,X3))=>r1_tarski(X1,X2))), file('out/t71_xboole_1__17_premises__rerun.E_input', t18_xboole_1)).
fof(idempotence_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X1)=X1), file('out/t71_xboole_1__17_premises__rerun.E_input', idempotence_k2_xboole_0)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t71_xboole_1__17_premises__rerun.E_input', d10_xboole_0)).
fof(t24_xboole_1, axiom, (![X1]:![X2]:![X3]:k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3))), file('out/t71_xboole_1__17_premises__rerun.E_input', t24_xboole_1)).
fof(t71_xboole_1, conjecture, (![X1]:![X2]:![X3]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X2)&(r1_xboole_0(X1,X2)&r1_xboole_0(X3,X2)))=>X1=X3)), file('out/t71_xboole_1__17_premises__rerun.E_input', t71_xboole_1)).
fof(t16_xboole_1, axiom, (![X1]:![X2]:![X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), file('out/t71_xboole_1__17_premises__rerun.E_input', t16_xboole_1)).
fof(t31_xboole_1, axiom, (![X1]:![X2]:![X3]:r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3))), file('out/t71_xboole_1__17_premises__rerun.E_input', t31_xboole_1)).
fof(t20_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&(r1_tarski(X1,X3)&![X4]:((r1_tarski(X4,X2)&r1_tarski(X4,X3))=>r1_tarski(X4,X1))))=>X1=k3_xboole_0(X2,X3))), file('out/t71_xboole_1__17_premises__rerun.E_input', t20_xboole_1)).
fof(t44_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3)))), file('out/t71_xboole_1__17_premises__rerun.E_input', t44_xboole_1)).
fof(t45_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1)))), file('out/t71_xboole_1__17_premises__rerun.E_input', t45_xboole_1)).
fof(c_0_18, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk1_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
fof(c_0_19, plain, (![X2]:k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_20,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk2_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk2_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_23, plain, (![X2]:![X2]:((r1_xboole_0(X2,X2)|X2!=k1_xboole_0)&(X2=k1_xboole_0|~r1_xboole_0(X2,X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])])])).
cnf(c_0_24,plain,(~r2_hidden(X1,k1_xboole_0)|~r1_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,plain,(r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27,plain,(r1_xboole_0(X1,k1_xboole_0)|~r1_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28,plain,(r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_26])).
fof(c_0_29, plain, (![X3]:![X4]:r1_tarski(k4_xboole_0(X3,X4),X3)), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_30, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_31, plain, (![X3]:![X4]:k2_xboole_0(X3,k4_xboole_0(X4,X3))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_32, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|k4_xboole_0(X3,X4)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_xboole_1])])])])).
cnf(c_0_33,plain,(r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:(~r1_tarski(X4,k3_xboole_0(X5,X6))|r1_tarski(X4,X5))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_xboole_1])])])])).
cnf(c_0_35,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_37, plain, (![X3]:k2_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_38,plain,(k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40,plain,(X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_41,plain,(r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_20, c_0_25])).
cnf(c_0_42,plain,(~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_33])])).
cnf(c_0_43,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_44, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_45,plain,(r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_47, plain, (![X4]:![X5]:![X6]:k2_xboole_0(X4,k3_xboole_0(X5,X6))=k3_xboole_0(k2_xboole_0(X4,X5),k2_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[t24_xboole_1])).
cnf(c_0_48,plain,(k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49,plain,(k2_xboole_0(X1,k1_xboole_0)=k2_xboole_0(X1,X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51,plain,(r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_52, negated_conjecture, (~(![X1]:![X2]:![X3]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X2)&(r1_xboole_0(X1,X2)&r1_xboole_0(X3,X2)))=>X1=X3))), inference(assume_negation,[status(cth)],[t71_xboole_1])).
cnf(c_0_53,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54,plain,(r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_55, plain, (![X4]:![X5]:![X6]:k3_xboole_0(k3_xboole_0(X4,X5),X6)=k3_xboole_0(X4,k3_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_56,plain,(k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_46])])).
cnf(c_0_58,plain,(k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_59, negated_conjecture, (((k2_xboole_0(esk3_0,esk4_0)=k2_xboole_0(esk5_0,esk4_0)&(r1_xboole_0(esk3_0,esk4_0)&r1_xboole_0(esk5_0,esk4_0)))&esk3_0!=esk5_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_52])])])).
fof(c_0_60, plain, (![X4]:![X5]:![X6]:r1_tarski(k2_xboole_0(k3_xboole_0(X4,X5),k3_xboole_0(X4,X6)),k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t31_xboole_1])).
cnf(c_0_61,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62,plain,(k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63,plain,(k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_57])).
cnf(c_0_64,negated_conjecture,(k2_xboole_0(esk3_0,esk4_0)=k2_xboole_0(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65,plain,(r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66,plain,(k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X1,X2)|~r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67,plain,(k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_63, c_0_57])).
cnf(c_0_68,plain,(r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_54, c_0_62])).
cnf(c_0_69,negated_conjecture,(k3_xboole_0(k2_xboole_0(esk3_0,esk4_0),k2_xboole_0(esk5_0,X1))=k2_xboole_0(esk5_0,k3_xboole_0(esk4_0,X1))), inference(spm,[status(thm)],[c_0_56, c_0_64])).
cnf(c_0_70,plain,(r1_tarski(k3_xboole_0(X1,X2),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_48]), c_0_48])).
cnf(c_0_71,plain,(k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_62]), c_0_62]), c_0_68])])).
cnf(c_0_72,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_73,negated_conjecture,(k3_xboole_0(k2_xboole_0(esk3_0,esk4_0),esk5_0)=esk5_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_57]), c_0_21]), c_0_57])).
cnf(c_0_74,plain,(r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_75,plain,(r1_xboole_0(X1,X2)|~r2_hidden(esk2_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_72, c_0_43])).
cnf(c_0_76,negated_conjecture,(r1_tarski(k3_xboole_0(X1,esk5_0),k3_xboole_0(X1,k2_xboole_0(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_68, c_0_73])).
cnf(c_0_77,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_74]), c_0_74])])).
cnf(c_0_78,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_75, c_0_25])).
cnf(c_0_79,negated_conjecture,(r1_xboole_0(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_80,plain,(k3_xboole_0(k2_xboole_0(X1,X2),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_21]), c_0_57])).
cnf(c_0_81,negated_conjecture,(r1_tarski(k3_xboole_0(X1,esk5_0),k3_xboole_0(k2_xboole_0(esk3_0,esk4_0),X1))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_82,negated_conjecture,(r1_xboole_0(esk4_0,esk5_0)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
fof(c_0_83, plain, (![X5]:![X6]:![X7]:(((((r1_tarski(esk6_3(X5,X6,X7),X6)|~r1_tarski(X5,X7))|~r1_tarski(X5,X6))|X5=k3_xboole_0(X6,X7))&(((r1_tarski(esk6_3(X5,X6,X7),X7)|~r1_tarski(X5,X7))|~r1_tarski(X5,X6))|X5=k3_xboole_0(X6,X7)))&(((~r1_tarski(esk6_3(X5,X6,X7),X5)|~r1_tarski(X5,X7))|~r1_tarski(X5,X6))|X5=k3_xboole_0(X6,X7)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_xboole_1])])])])])])).
fof(c_0_84, plain, (![X4]:![X5]:![X6]:(~r1_tarski(k4_xboole_0(X4,X5),X6)|r1_tarski(X4,k2_xboole_0(X5,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
cnf(c_0_85,plain,(r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X1,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_68, c_0_80])).
cnf(c_0_86,negated_conjecture,(r1_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_87,negated_conjecture,(r1_tarski(k3_xboole_0(esk5_0,k2_xboole_0(esk3_0,X1)),k2_xboole_0(esk3_0,k3_xboole_0(esk4_0,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_56]), c_0_77])).
cnf(c_0_88,negated_conjecture,(k3_xboole_0(esk4_0,esk5_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_82])).
cnf(c_0_89,plain,(X1=k3_xboole_0(X2,X3)|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_tarski(esk6_3(X1,X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_90,plain,(X1=k3_xboole_0(X2,X3)|r1_tarski(esk6_3(X1,X2,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_91,plain,(r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_92,plain,(r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(k2_xboole_0(X2,X3),X1))), inference(spm,[status(thm)],[c_0_85, c_0_77])).
cnf(c_0_93,negated_conjecture,(r1_xboole_0(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_78, c_0_86])).
fof(c_0_94, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|X4=k2_xboole_0(X3,k4_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
cnf(c_0_95,negated_conjecture,(r1_tarski(k3_xboole_0(esk5_0,k2_xboole_0(esk3_0,esk5_0)),esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_57])).
cnf(c_0_96,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_46])])).
cnf(c_0_97,plain,(r1_tarski(X1,k2_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_46]), c_0_38])).
cnf(c_0_98,negated_conjecture,(r1_tarski(k3_xboole_0(esk3_0,k2_xboole_0(esk5_0,X1)),k2_xboole_0(esk5_0,k3_xboole_0(esk4_0,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_69]), c_0_77])).
cnf(c_0_99,negated_conjecture,(k3_xboole_0(esk4_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_93])).
cnf(c_0_100,plain,(X1=k2_xboole_0(X2,k4_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_101,negated_conjecture,(r1_tarski(esk5_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97])])).
cnf(c_0_102,negated_conjecture,(esk3_0!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_103,negated_conjecture,(r1_tarski(k3_xboole_0(esk3_0,k2_xboole_0(esk5_0,esk3_0)),esk5_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_57])).
cnf(c_0_104,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_100, c_0_38])).
cnf(c_0_105,negated_conjecture,(~r1_tarski(esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_101]), c_0_102])).
cnf(c_0_106,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_67]), c_0_101])]), c_0_105]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 107
# Proof object clause steps            : 70
# Proof object formula steps           : 37
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 42
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 3081
# ...of these trivial                  : 553
# ...subsumed                          : 1599
# ...remaining for further processing  : 929
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 37
# Backward-rewritten                   : 142
# Generated clauses                    : 49183
# ...of the previous two non-trivial   : 19072
# Contextual simplify-reflections      : 215
# Paramodulations                      : 49179
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 747
#    Positive orientable unit clauses  : 452
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 287
# Current number of unprocessed clauses: 14512
# ...number of literals in the above   : 25277
# Current number of archived formulas  : 0
# Current number of archived clauses   : 179
# Clause-clause subsumption calls (NU) : 24997
# Rec. Clause-clause subsumption calls : 16810
# Non-unit clause-clause subsumptions  : 1614
# Unit Clause-clause subsumption calls : 1318
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2322
# BW rewrite match successes           : 609
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 432450

# -------------------------------------------------
# User time                : 0.344 s
# System time              : 0.012 s
# Total time               : 0.356 s
# Maximum resident set size: 45328 pages
