# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>(X2=k6_relat_1(X1)<=>![X3]:![X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>(r2_hidden(X3,X1)&X3=X4))))), file('out/t71_relat_1__4_premises__rerun.E_input', d10_relat_1)).
fof(dt_k6_relat_1, axiom, (![X1]:v1_relat_1(k6_relat_1(X1))), file('out/t71_relat_1__4_premises__rerun.E_input', dt_k6_relat_1)).
fof(d4_relat_1, axiom, (![X1]:![X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1)))), file('out/t71_relat_1__4_premises__rerun.E_input', d4_relat_1)).
fof(d5_relat_1, axiom, (![X1]:![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1)))), file('out/t71_relat_1__4_premises__rerun.E_input', d5_relat_1)).
fof(t71_relat_1, conjecture, (![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1)), file('out/t71_relat_1__4_premises__rerun.E_input', t71_relat_1)).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:![X7]:![X8]:((((((r2_hidden(X7,X5)|~r2_hidden(k4_tarski(X7,X8),X6))|X6!=k6_relat_1(X5))|~v1_relat_1(X6))&(((X7=X8|~r2_hidden(k4_tarski(X7,X8),X6))|X6!=k6_relat_1(X5))|~v1_relat_1(X6)))&((((~r2_hidden(X7,X5)|X7!=X8)|r2_hidden(k4_tarski(X7,X8),X6))|X6!=k6_relat_1(X5))|~v1_relat_1(X6)))&((((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|(~r2_hidden(esk1_2(X5,X6),X5)|esk1_2(X5,X6)!=esk2_2(X5,X6)))|X6=k6_relat_1(X5))|~v1_relat_1(X6))&((((r2_hidden(esk1_2(X5,X6),X5)|r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6))|X6=k6_relat_1(X5))|~v1_relat_1(X6))&(((esk1_2(X5,X6)=esk2_2(X5,X6)|r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6))|X6=k6_relat_1(X5))|~v1_relat_1(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_relat_1])])])])])])])).
fof(c_0_6, plain, (![X2]:v1_relat_1(k6_relat_1(X2))), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_7,plain,(r2_hidden(X3,X2)|~v1_relat_1(X1)|X1!=k6_relat_1(X2)|~r2_hidden(k4_tarski(X3,X4),X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(X7,esk3_3(X5,X6,X7)),X5))|X6!=k1_relat_1(X5))&((~r2_hidden(k4_tarski(X7,X9),X5)|r2_hidden(X7,X6))|X6!=k1_relat_1(X5)))&(((~r2_hidden(esk4_2(X5,X6),X6)|~r2_hidden(k4_tarski(esk4_2(X5,X6),X11),X5))|X6=k1_relat_1(X5))&((r2_hidden(esk4_2(X5,X6),X6)|r2_hidden(k4_tarski(esk4_2(X5,X6),esk5_2(X5,X6)),X5))|X6=k1_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X5]:![X6]:![X11]:((((~r2_hidden(X7,X6)|r2_hidden(k4_tarski(esk6_3(X5,X6,X7),X7),X5))|X6!=k2_relat_1(X5))&((~r2_hidden(k4_tarski(X9,X7),X5)|r2_hidden(X7,X6))|X6!=k2_relat_1(X5)))&(((~r2_hidden(esk7_2(X5,X6),X6)|~r2_hidden(k4_tarski(X11,esk7_2(X5,X6)),X5))|X6=k2_relat_1(X5))&((r2_hidden(esk7_2(X5,X6),X6)|r2_hidden(k4_tarski(esk8_2(X5,X6),esk7_2(X5,X6)),X5))|X6=k2_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])])).
cnf(c_0_11,plain,(X3=X4|~v1_relat_1(X1)|X1!=k6_relat_1(X2)|~r2_hidden(k4_tarski(X3,X4),X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,plain,(r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)|X1!=k6_relat_1(X2)|X3!=X4|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13,plain,(r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_7]), c_0_8])])).
cnf(c_0_14,plain,(X1=k1_relat_1(X2)|r2_hidden(k4_tarski(esk4_2(X2,X1),esk5_2(X2,X1)),X2)|r2_hidden(esk4_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(X1=k2_relat_1(X2)|r2_hidden(k4_tarski(esk8_2(X2,X1),esk7_2(X2,X1)),X2)|r2_hidden(esk7_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(X1=X2|~r2_hidden(k4_tarski(X1,X2),k6_relat_1(X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_11]), c_0_8])])).
cnf(c_0_17,plain,(r2_hidden(k4_tarski(X1,X1),X2)|X2!=k6_relat_1(X3)|~r2_hidden(X1,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(X1=k1_relat_1(k6_relat_1(X2))|r2_hidden(esk4_2(k6_relat_1(X2),X1),X1)|r2_hidden(esk4_2(k6_relat_1(X2),X1),X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19,plain,(X1=k2_relat_1(k6_relat_1(X2))|r2_hidden(esk7_2(k6_relat_1(X2),X1),X1)|r2_hidden(esk8_2(k6_relat_1(X2),X1),X2)), inference(spm,[status(thm)],[c_0_13, c_0_15])).
cnf(c_0_20,plain,(esk8_2(k6_relat_1(X1),X2)=esk7_2(k6_relat_1(X1),X2)|X2=k2_relat_1(k6_relat_1(X1))|r2_hidden(esk7_2(k6_relat_1(X1),X2),X2)), inference(spm,[status(thm)],[c_0_16, c_0_15])).
fof(c_0_21, negated_conjecture, (~(![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1))), inference(assume_negation,[status(cth)],[t71_relat_1])).
cnf(c_0_22,plain,(r2_hidden(k4_tarski(X1,X1),k6_relat_1(X2))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_17]), c_0_8])])).
cnf(c_0_23,plain,(k1_relat_1(k6_relat_1(X1))=X1|r2_hidden(esk4_2(k6_relat_1(X1),X1),X1)), inference(ef,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(X1=k2_relat_1(k6_relat_1(X2))|r2_hidden(esk7_2(k6_relat_1(X2),X1),X2)|r2_hidden(esk7_2(k6_relat_1(X2),X1),X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_25, negated_conjecture, ((k1_relat_1(k6_relat_1(esk9_0))!=esk9_0|k2_relat_1(k6_relat_1(esk10_0))!=esk10_0)), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])).
cnf(c_0_26,plain,(X1=k1_relat_1(X2)|~r2_hidden(k4_tarski(esk4_2(X2,X1),X3),X2)|~r2_hidden(esk4_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27,plain,(k1_relat_1(k6_relat_1(X1))=X1|r2_hidden(k4_tarski(esk4_2(k6_relat_1(X1),X1),esk4_2(k6_relat_1(X1),X1)),k6_relat_1(X1))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28,plain,(k2_relat_1(k6_relat_1(X1))=X1|r2_hidden(esk7_2(k6_relat_1(X1),X1),X1)), inference(ef,[status(thm)],[c_0_24])).
cnf(c_0_29,negated_conjecture,(k2_relat_1(k6_relat_1(esk10_0))!=esk10_0|k1_relat_1(k6_relat_1(esk9_0))!=esk9_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30,plain,(k1_relat_1(k6_relat_1(X1))=X1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_23])).
cnf(c_0_31,plain,(X1=k2_relat_1(X2)|~r2_hidden(k4_tarski(X3,esk7_2(X2,X1)),X2)|~r2_hidden(esk7_2(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_32,plain,(k2_relat_1(k6_relat_1(X1))=X1|r2_hidden(k4_tarski(esk7_2(k6_relat_1(X1),X1),esk7_2(k6_relat_1(X1),X1)),k6_relat_1(X1))), inference(spm,[status(thm)],[c_0_22, c_0_28])).
cnf(c_0_33,negated_conjecture,(k2_relat_1(k6_relat_1(esk10_0))!=esk10_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])])).
cnf(c_0_34,plain,(k2_relat_1(k6_relat_1(X1))=X1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_28])).
cnf(c_0_35,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 25
# Proof object formula steps           : 11
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 113
# ...of these trivial                  : 0
# ...subsumed                          : 31
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 22
# Generated clauses                    : 357
# ...of the previous two non-trivial   : 341
# Contextual simplify-reflections      : 10
# Paramodulations                      : 301
# Factorizations                       : 10
# Equation resolutions                 : 46
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 54
# Current number of unprocessed clauses: 131
# ...number of literals in the above   : 459
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 960
# Rec. Clause-clause subsumption calls : 556
# Non-unit clause-clause subsumptions  : 43
# Unit Clause-clause subsumption calls : 95
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7586

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45328 pages
