# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t71_orders_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,X3)&m1_orders_2(X3,X1,X4))=>k3_orders_2(X1,X3,X2)=k3_orders_2(X1,X4,X2))))))), file('out/t71_orders_2__8_premises__rerun.E_input', t71_orders_2)).
fof(dt_k3_orders_2, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k3_orders_2(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))), file('out/t71_orders_2__8_premises__rerun.E_input', dt_k3_orders_2)).
fof(t70_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_orders_2(X1,X2,X3)&(r2_hidden(X2,X4)&(r2_hidden(X3,X5)&m1_orders_2(X5,X1,X4))))=>r2_hidden(X2,X5)))))))), file('out/t71_orders_2__8_premises__rerun.E_input', t70_orders_2)).
fof(dt_m1_orders_2, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>![X3]:(m1_orders_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t71_orders_2__8_premises__rerun.E_input', dt_m1_orders_2)).
fof(t4_subset, axiom, (![X1]:![X2]:![X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('out/t71_orders_2__8_premises__rerun.E_input', t4_subset)).
fof(t57_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,k3_orders_2(X1,X4,X3))<=>(r2_orders_2(X1,X2,X3)&r2_hidden(X2,X4)))))))), file('out/t71_orders_2__8_premises__rerun.E_input', t57_orders_2)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t71_orders_2__8_premises__rerun.E_input', d3_tarski)).
fof(t67_orders_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_orders_2(X3,X1,X2)=>r1_tarski(X3,X2))))), file('out/t71_orders_2__8_premises__rerun.E_input', t67_orders_2)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t71_orders_2__8_premises__rerun.E_input', d10_xboole_0)).
fof(c_0_9, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,X3)&m1_orders_2(X3,X1,X4))=>k3_orders_2(X1,X3,X2)=k3_orders_2(X1,X4,X2)))))))), inference(assume_negation,[status(cth)],[t71_orders_2])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|~l1_orders_2(X4)))))|(~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X6,u1_struct_0(X4))))|m1_subset_1(k3_orders_2(X4,X5,X6),k1_zfmisc_1(u1_struct_0(X4))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_orders_2])])])).
fof(c_0_11, negated_conjecture, (((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((r2_hidden(esk2_0,esk3_0)&m1_orders_2(esk3_0,esk1_0,esk4_0))&k3_orders_2(esk1_0,esk3_0,esk2_0)!=k3_orders_2(esk1_0,esk4_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_12,plain,(m1_subset_1(k3_orders_2(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_19, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((v2_struct_0(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|~l1_orders_2(X6)))))|(~m1_subset_1(X7,u1_struct_0(X6))|(~m1_subset_1(X8,u1_struct_0(X6))|(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X6)))|(~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X6)))|((~r2_orders_2(X6,X7,X8)|(~r2_hidden(X7,X9)|(~r2_hidden(X8,X10)|~m1_orders_2(X10,X6,X9))))|r2_hidden(X7,X10)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_orders_2])])])])])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|~l1_orders_2(X4)))))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4))))|(~m1_orders_2(X6,X4,X5)|m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_orders_2])])])])])])).
fof(c_0_21, plain, (![X4]:![X5]:![X6]:((~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6)))|m1_subset_1(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_22,negated_conjecture,(m1_subset_1(k3_orders_2(esk1_0,esk4_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_25, plain, (![X5]:![X6]:![X7]:![X8]:(((((((r2_orders_2(X5,X6,X7)|~r2_hidden(X6,k3_orders_2(X5,X8,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(((((r2_hidden(X6,X8)|~r2_hidden(X6,k3_orders_2(X5,X8,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))&((((((~r2_orders_2(X5,X6,X7)|~r2_hidden(X6,X8))|r2_hidden(X6,k3_orders_2(X5,X8,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5))))|~m1_subset_1(X7,u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_orders_2])])])])])])])).
cnf(c_0_26,plain,(r2_hidden(X1,X2)|v2_struct_0(X3)|~m1_orders_2(X2,X3,X4)|~r2_hidden(X5,X2)|~r2_hidden(X1,X4)|~r2_orders_2(X3,X1,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X5,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m1_orders_2(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(m1_subset_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_29, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_30,negated_conjecture,(m1_subset_1(k3_orders_2(esk1_0,esk4_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_31, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|(~v3_orders_2(X4)|(~v4_orders_2(X4)|(~v5_orders_2(X4)|~l1_orders_2(X4)))))|(~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|(~m1_orders_2(X6,X4,X5)|r1_tarski(X6,X5))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t67_orders_2])])])])])])).
cnf(c_0_32,negated_conjecture,(m1_subset_1(k3_orders_2(esk1_0,esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_24]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_33,plain,(v2_struct_0(X1)|r2_hidden(X2,k3_orders_2(X1,X4,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,X4)|~r2_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(r2_hidden(X1,X2)|v2_struct_0(X3)|~r2_orders_2(X3,X1,X4)|~m1_orders_2(X2,X3,X5)|~r2_hidden(X4,X2)|~r2_hidden(X1,X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_35,plain,(v2_struct_0(X1)|r2_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,k3_orders_2(X1,X4,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k3_orders_2(esk1_0,esk4_0,esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_30])).
cnf(c_0_38,plain,(v2_struct_0(X1)|r2_hidden(X2,X4)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,k3_orders_2(X1,X4,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39,plain,(r1_tarski(X1,X2)|v2_struct_0(X3)|~m1_orders_2(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(k3_orders_2(esk1_0,esk3_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_32, c_0_23])).
cnf(c_0_41,plain,(r2_hidden(X1,k3_orders_2(X2,X3,X4))|v2_struct_0(X2)|~r2_orders_2(X2,X1,X4)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_42,negated_conjecture,(r2_hidden(X1,X2)|~r2_orders_2(esk1_0,X1,X3)|~m1_orders_2(X2,esk1_0,esk4_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_43,negated_conjecture,(m1_orders_2(esk3_0,esk1_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_24])).
cnf(c_0_45,plain,(r1_tarski(k3_orders_2(X1,X2,X3),X4)|r2_orders_2(X1,esk5_2(k3_orders_2(X1,X2,X3),X4),X3)|v2_struct_0(X1)|~m1_subset_1(esk5_2(k3_orders_2(X1,X2,X3),X4),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_46,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),X1)|m1_subset_1(esk5_2(k3_orders_2(esk1_0,esk4_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_47,plain,(r1_tarski(k3_orders_2(X1,X2,X3),X4)|r2_hidden(esk5_2(k3_orders_2(X1,X2,X3),X4),X2)|v2_struct_0(X1)|~m1_subset_1(esk5_2(k3_orders_2(X1,X2,X3),X4),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_38, c_0_36])).
cnf(c_0_48,negated_conjecture,(r1_tarski(X1,esk4_0)|~m1_orders_2(X1,esk1_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_49,negated_conjecture,(m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k3_orders_2(esk1_0,esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_40])).
cnf(c_0_50,negated_conjecture,(r2_hidden(X1,k3_orders_2(esk1_0,esk3_0,X2))|~r2_orders_2(esk1_0,X1,X2)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_24]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_51,negated_conjecture,(r2_hidden(X1,esk3_0)|~r2_orders_2(esk1_0,X1,X2)|~r2_hidden(X1,esk4_0)|~r2_hidden(X2,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_52,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),X1)|r2_orders_2(esk1_0,esk5_2(k3_orders_2(esk1_0,esk4_0,esk2_0),X1),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_13]), c_0_23]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),X1)|r2_hidden(esk5_2(k3_orders_2(esk1_0,esk4_0,esk2_0),X1),esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_46]), c_0_13]), c_0_23]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_55,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_56,negated_conjecture,(r1_tarski(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_48, c_0_43])).
cnf(c_0_57,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk3_0,esk2_0),X1)|m1_subset_1(esk5_2(k3_orders_2(esk1_0,esk3_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_36])).
cnf(c_0_58,negated_conjecture,(r2_hidden(X1,k3_orders_2(esk1_0,esk3_0,esk2_0))|~r2_orders_2(esk1_0,X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_50, c_0_23])).
cnf(c_0_59,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),X1)|r2_hidden(esk5_2(k3_orders_2(esk1_0,esk4_0,esk2_0),X1),esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_60,negated_conjecture,(r2_hidden(X1,k3_orders_2(esk1_0,esk4_0,X2))|~r2_orders_2(esk1_0,X1,X2)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_61,negated_conjecture,(r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_62,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk3_0,esk2_0),X1)|r2_hidden(esk5_2(k3_orders_2(esk1_0,esk3_0,esk2_0),X1),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_57]), c_0_24]), c_0_23]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
fof(c_0_63, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_64,plain,(r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_65,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),X1)|r2_hidden(esk5_2(k3_orders_2(esk1_0,esk4_0,esk2_0),X1),k3_orders_2(esk1_0,esk3_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_52]), c_0_59])).
cnf(c_0_66,negated_conjecture,(r2_hidden(X1,k3_orders_2(esk1_0,esk4_0,esk2_0))|~r2_orders_2(esk1_0,X1,esk2_0)|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_60, c_0_23])).
cnf(c_0_67,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk3_0,esk2_0),X1)|r2_orders_2(esk1_0,esk5_2(k3_orders_2(esk1_0,esk3_0,esk2_0),X1),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_57]), c_0_24]), c_0_23]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_68,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk3_0,esk2_0),X1)|r2_hidden(esk5_2(k3_orders_2(esk1_0,esk3_0,esk2_0),X1),esk4_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_69,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_70,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),k3_orders_2(esk1_0,esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_71,negated_conjecture,(k3_orders_2(esk1_0,esk3_0,esk2_0)!=k3_orders_2(esk1_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_72,negated_conjecture,(r1_tarski(k3_orders_2(esk1_0,esk3_0,esk2_0),X1)|r2_hidden(esk5_2(k3_orders_2(esk1_0,esk3_0,esk2_0),X1),k3_orders_2(esk1_0,esk4_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_73,negated_conjecture,(~r1_tarski(k3_orders_2(esk1_0,esk3_0,esk2_0),k3_orders_2(esk1_0,esk4_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
cnf(c_0_74,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_72]), c_0_73]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 56
# Proof object formula steps           : 19
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 79
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 406
# ...of these trivial                  : 1
# ...subsumed                          : 21
# ...remaining for further processing  : 384
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 983
# ...of the previous two non-trivial   : 976
# Contextual simplify-reflections      : 11
# Paramodulations                      : 981
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 358
#    Positive orientable unit clauses  : 48
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 307
# Current number of unprocessed clauses: 619
# ...number of literals in the above   : 1867
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 19389
# Rec. Clause-clause subsumption calls : 12488
# Non-unit clause-clause subsumptions  : 32
# Unit Clause-clause subsumption calls : 809
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1029
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 47593

# -------------------------------------------------
# User time                : 0.052 s
# System time              : 0.000 s
# Total time               : 0.052 s
# Maximum resident set size: 45368 pages
