# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t72_xboole_1, conjecture, (![X1]:![X2]:![X3]:![X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2)), file('out/t72_xboole_1__10_premises__rerun.E_input', t72_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), file('out/t72_xboole_1__10_premises__rerun.E_input', commutativity_k2_xboole_0)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t72_xboole_1__10_premises__rerun.E_input', t12_xboole_1)).
fof(d7_xboole_0, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0)), file('out/t72_xboole_1__10_premises__rerun.E_input', d7_xboole_0)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t72_xboole_1__10_premises__rerun.E_input', t48_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t72_xboole_1__10_premises__rerun.E_input', commutativity_k3_xboole_0)).
fof(t43_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3))), file('out/t72_xboole_1__10_premises__rerun.E_input', t43_xboole_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t72_xboole_1__10_premises__rerun.E_input', t7_xboole_1)).
fof(l32_xboole_1, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('out/t72_xboole_1__10_premises__rerun.E_input', l32_xboole_1)).
fof(t67_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&(r1_tarski(X1,X3)&r1_xboole_0(X2,X3)))=>X1=k1_xboole_0)), file('out/t72_xboole_1__10_premises__rerun.E_input', t67_xboole_1)).
fof(t36_xboole_1, axiom, (![X1]:![X2]:r1_tarski(k4_xboole_0(X1,X2),X1)), file('out/t72_xboole_1__10_premises__rerun.E_input', t36_xboole_1)).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2))), inference(assume_negation,[status(cth)],[t72_xboole_1])).
fof(c_0_12, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k2_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_13, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_14, negated_conjecture, (((k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk4_0)&(r1_xboole_0(esk3_0,esk1_0)&r1_xboole_0(esk4_0,esk2_0)))&esk3_0!=esk2_0)), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
fof(c_0_15, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_xboole_0(X3,X4)|k3_xboole_0(X3,X4)=k1_xboole_0)&(k3_xboole_0(X3,X4)!=k1_xboole_0|r1_xboole_0(X3,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])])])).
fof(c_0_16, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_17, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_18,plain,(k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(~r1_tarski(X4,k2_xboole_0(X5,X6))|r1_tarski(k4_xboole_0(X4,X5),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_21,negated_conjecture,(k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,plain,(r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_26, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_27,plain,(k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_28, plain, (![X3]:![X4]:![X3]:![X4]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|k4_xboole_0(X3,X4)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])])])).
fof(c_0_29, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|(~r1_tarski(X4,X6)|~r1_xboole_0(X5,X6)))|X4=k1_xboole_0)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_xboole_1])])])])).
fof(c_0_30, plain, (![X3]:![X4]:r1_tarski(k4_xboole_0(X3,X4),X3)), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_31,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32,negated_conjecture,(k2_xboole_0(esk2_0,esk1_0)=k2_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_21, c_0_18])).
cnf(c_0_33,plain,(r1_xboole_0(X1,X2)|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_34,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_23]), c_0_23])).
cnf(c_0_35,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_25, c_0_23])).
cnf(c_0_36,negated_conjecture,(r1_xboole_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38,plain,(X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_19, c_0_27])).
cnf(c_0_39,plain,(r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40,plain,(X1=k1_xboole_0|~r1_xboole_0(X2,X3)|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42,negated_conjecture,(r1_tarski(k4_xboole_0(X1,esk2_0),esk1_0)|~r1_tarski(X1,k2_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43,plain,(r1_xboole_0(X1,X2)|k4_xboole_0(X2,k4_xboole_0(X2,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_44,negated_conjecture,(k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk1_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45,negated_conjecture,(r1_tarski(esk2_0,k2_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_37, c_0_32])).
cnf(c_0_46,plain,(X1=X2|k4_xboole_0(X1,X2)!=k1_xboole_0|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_47,plain,(k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48,negated_conjecture,(r1_tarski(k4_xboole_0(esk3_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_37])).
cnf(c_0_49,negated_conjecture,(r1_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50,negated_conjecture,(r1_tarski(k4_xboole_0(esk2_0,esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_31, c_0_45])).
cnf(c_0_51,negated_conjecture,(r1_xboole_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_52,plain,(X1=X2|k4_xboole_0(X1,X2)!=k1_xboole_0|k4_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_39])).
cnf(c_0_53,negated_conjecture,(k4_xboole_0(esk3_0,esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])])).
cnf(c_0_54,negated_conjecture,(k4_xboole_0(esk2_0,esk3_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_50]), c_0_51])])).
cnf(c_0_55,negated_conjecture,(esk3_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_56,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])]), c_0_55]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 34
# Proof object formula steps           : 23
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 15
# Processed clauses                    : 615
# ...of these trivial                  : 26
# ...subsumed                          : 282
# ...remaining for further processing  : 307
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 27
# Backward-rewritten                   : 32
# Generated clauses                    : 1531
# ...of the previous two non-trivial   : 1106
# Contextual simplify-reflections      : 94
# Paramodulations                      : 1531
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 233
#    Positive orientable unit clauses  : 46
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 181
# Current number of unprocessed clauses: 415
# ...number of literals in the above   : 999
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 4023
# Rec. Clause-clause subsumption calls : 3744
# Non-unit clause-clause subsumptions  : 379
# Unit Clause-clause subsumption calls : 176
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 48
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14328

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.008 s
# Total time               : 0.036 s
# Maximum resident set size: 45364 pages
