# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t72_mcart_1, conjecture, (![X1]:![X2]:![X3]:![X4]:~((r2_hidden(X1,k3_zfmisc_1(X2,X3,X4))&![X5]:![X6]:![X7]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&(r2_hidden(X7,X4)&X1=k3_mcart_1(X5,X6,X7)))))))), file('out/t72_mcart_1__6_premises__rerun.E_input', t72_mcart_1)).
fof(d3_zfmisc_1, axiom, (![X1]:![X2]:![X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), file('out/t72_mcart_1__6_premises__rerun.E_input', d3_zfmisc_1)).
fof(t10_mcart_1, axiom, (![X1]:![X2]:![X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3)))), file('out/t72_mcart_1__6_premises__rerun.E_input', t10_mcart_1)).
fof(t31_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), file('out/t72_mcart_1__6_premises__rerun.E_input', t31_mcart_1)).
fof(d4_mcart_1, axiom, (![X1]:![X2]:![X3]:![X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), file('out/t72_mcart_1__6_premises__rerun.E_input', d4_mcart_1)).
fof(l139_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,k2_zfmisc_1(X2,X3))&![X4]:![X5]:~(k4_tarski(X4,X5)=X1)))), file('out/t72_mcart_1__6_premises__rerun.E_input', l139_zfmisc_1)).
fof(t7_mcart_1, axiom, (![X1]:![X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2)), file('out/t72_mcart_1__6_premises__rerun.E_input', t7_mcart_1)).
fof(c_0_7, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:~((r2_hidden(X1,k3_zfmisc_1(X2,X3,X4))&![X5]:![X6]:![X7]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&(r2_hidden(X7,X4)&X1=k3_mcart_1(X5,X6,X7))))))))), inference(assume_negation,[status(cth)],[t72_mcart_1])).
fof(c_0_8, negated_conjecture, (![X12]:![X13]:![X14]:(r2_hidden(esk1_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))&(~r2_hidden(X12,esk2_0)|(~r2_hidden(X13,esk3_0)|(~r2_hidden(X14,esk4_0)|esk1_0!=k3_mcart_1(X12,X13,X14)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:k3_zfmisc_1(X4,X5,X6)=k2_zfmisc_1(k2_zfmisc_1(X4,X5),X6)), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:((r2_hidden(k1_mcart_1(X4),X5)|~r2_hidden(X4,k2_zfmisc_1(X5,X6)))&(r2_hidden(k2_mcart_1(X4),X6)|~r2_hidden(X4,k2_zfmisc_1(X5,X6))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_11,negated_conjecture,(r2_hidden(esk1_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, (![X5]:![X6]:![X7]:![X8]:k4_mcart_1(X5,X6,X7,X8)=k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8)), inference(variable_rename,[status(thm)],[t31_mcart_1])).
fof(c_0_14, plain, (![X5]:![X6]:![X7]:![X8]:k4_mcart_1(X5,X6,X7,X8)=k4_tarski(k3_mcart_1(X5,X6,X7),X8)), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_15, plain, (![X6]:![X7]:![X8]:(~r2_hidden(X6,k2_zfmisc_1(X7,X8))|k4_tarski(esk5_1(X6),esk6_1(X6))=X6)), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l139_zfmisc_1])])])])])])])).
cnf(c_0_16,plain,(r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,negated_conjecture,(r2_hidden(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_18, plain, (![X3]:![X4]:![X3]:![X4]:(k1_mcart_1(k4_tarski(X3,X4))=X3&k2_mcart_1(k4_tarski(X3,X4))=X4)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t7_mcart_1])])])).
cnf(c_0_19,plain,(k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(k4_tarski(esk5_1(X1),esk6_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,negated_conjecture,(r2_hidden(k1_mcart_1(esk1_0),k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,plain,(k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(k4_tarski(k3_mcart_1(X1,X2,X3),X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25,negated_conjecture,(k4_tarski(esk5_1(k1_mcart_1(esk1_0)),esk6_1(k1_mcart_1(esk1_0)))=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26,negated_conjecture,(k4_tarski(esk5_1(esk1_0),esk6_1(esk1_0))=esk1_0), inference(spm,[status(thm)],[c_0_21, c_0_17])).
cnf(c_0_27,negated_conjecture,(esk1_0!=k3_mcart_1(X1,X2,X3)|~r2_hidden(X3,esk4_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28,plain,(k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_23])).
cnf(c_0_29,negated_conjecture,(esk5_1(k1_mcart_1(esk1_0))=k1_mcart_1(k1_mcart_1(esk1_0))), inference(spm,[status(thm)],[c_0_23, c_0_25])).
cnf(c_0_30,negated_conjecture,(esk5_1(esk1_0)=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_23, c_0_26])).
cnf(c_0_31,negated_conjecture,(k4_tarski(k4_tarski(X1,X2),X3)!=esk1_0|~r2_hidden(X3,esk4_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,negated_conjecture,(k4_tarski(k1_mcart_1(k1_mcart_1(esk1_0)),esk6_1(k1_mcart_1(esk1_0)))=k1_mcart_1(esk1_0)), inference(rw,[status(thm)],[c_0_25, c_0_29])).
cnf(c_0_33,negated_conjecture,(r2_hidden(k1_mcart_1(k1_mcart_1(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_16, c_0_22])).
cnf(c_0_34,plain,(k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_35,plain,(r2_hidden(k2_mcart_1(X1),X3)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),esk6_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_26, c_0_30])).
cnf(c_0_37,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),X1)!=esk1_0|~r2_hidden(esk6_1(k1_mcart_1(esk1_0)),esk3_0)|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_38,negated_conjecture,(esk6_1(k1_mcart_1(esk1_0))=k2_mcart_1(k1_mcart_1(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_32])).
cnf(c_0_39,negated_conjecture,(r2_hidden(k2_mcart_1(k1_mcart_1(esk1_0)),esk3_0)), inference(spm,[status(thm)],[c_0_35, c_0_22])).
cnf(c_0_40,negated_conjecture,(esk6_1(esk1_0)=k2_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_34, c_0_36])).
cnf(c_0_41,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),X1)!=esk1_0|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
cnf(c_0_42,negated_conjecture,(k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_36, c_0_40])).
cnf(c_0_43,negated_conjecture,(r2_hidden(k2_mcart_1(esk1_0),esk4_0)), inference(spm,[status(thm)],[c_0_35, c_0_17])).
cnf(c_0_44,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 30
# Proof object formula steps           : 15
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 8
# Processed clauses                    : 34
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 34
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 20
# ...of the previous two non-trivial   : 24
# Contextual simplify-reflections      : 0
# Paramodulations                      : 20
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 19
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 6
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 1
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 6
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1029

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 45476 pages
