# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), file('out/t72_xboole_1__16_premises__rerun.E_input', t4_xboole_0)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t72_xboole_1__16_premises__rerun.E_input', t48_xboole_1)).
fof(t72_xboole_1, conjecture, (![X1]:![X2]:![X3]:![X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2)), file('out/t72_xboole_1__16_premises__rerun.E_input', t72_xboole_1)).
fof(t7_xboole_0, axiom, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1))))), file('out/t72_xboole_1__16_premises__rerun.E_input', t7_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), file('out/t72_xboole_1__16_premises__rerun.E_input', commutativity_k2_xboole_0)).
fof(t47_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), file('out/t72_xboole_1__16_premises__rerun.E_input', t47_xboole_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t72_xboole_1__16_premises__rerun.E_input', t7_xboole_1)).
fof(t4_xboole_1, axiom, (![X1]:![X2]:![X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t72_xboole_1__16_premises__rerun.E_input', t4_xboole_1)).
fof(t43_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3))), file('out/t72_xboole_1__16_premises__rerun.E_input', t43_xboole_1)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t72_xboole_1__16_premises__rerun.E_input', t3_boole)).
fof(idempotence_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X1)=X1), file('out/t72_xboole_1__16_premises__rerun.E_input', idempotence_k2_xboole_0)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t72_xboole_1__16_premises__rerun.E_input', t12_xboole_1)).
fof(t51_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), file('out/t72_xboole_1__16_premises__rerun.E_input', t51_xboole_1)).
fof(t40_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), file('out/t72_xboole_1__16_premises__rerun.E_input', t40_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t72_xboole_1__16_premises__rerun.E_input', commutativity_k3_xboole_0)).
fof(t39_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), file('out/t72_xboole_1__16_premises__rerun.E_input', t39_xboole_1)).
fof(t63_xboole_1, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3))), file('out/t72_xboole_1__16_premises__rerun.E_input', t63_xboole_1)).
fof(c_0_17, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:((r1_xboole_0(X4,X5)|r2_hidden(esk1_2(X4,X5),k3_xboole_0(X4,X5)))&(~r2_hidden(X7,k3_xboole_0(X4,X5))|~r1_xboole_0(X4,X5)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])])).
fof(c_0_18, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_19, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2))), inference(assume_negation,[status(cth)],[t72_xboole_1])).
cnf(c_0_20,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_22, plain, (![X3]:(X3=k1_xboole_0|r2_hidden(esk6_1(X3),X3))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])])])).
fof(c_0_23, negated_conjecture, (((k2_xboole_0(esk2_0,esk3_0)=k2_xboole_0(esk4_0,esk5_0)&(r1_xboole_0(esk4_0,esk2_0)&r1_xboole_0(esk5_0,esk3_0)))&esk4_0!=esk3_0)), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
fof(c_0_24, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k2_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_25, plain, (![X3]:![X4]:k4_xboole_0(X3,k3_xboole_0(X3,X4))=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_26,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,plain,(r2_hidden(esk6_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_29, plain, (![X4]:![X5]:![X6]:k2_xboole_0(k2_xboole_0(X4,X5),X6)=k2_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t4_xboole_1])).
cnf(c_0_30,negated_conjecture,(k2_xboole_0(esk2_0,esk3_0)=k2_xboole_0(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,plain,(k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, (![X4]:![X5]:![X6]:(~r1_tarski(X4,k2_xboole_0(X5,X6))|r1_tarski(k4_xboole_0(X4,X5),X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_33,plain,(k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35,negated_conjecture,(r1_xboole_0(esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_36, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_37,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38,plain,(k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39,negated_conjecture,(k2_xboole_0(esk3_0,esk2_0)=k2_xboole_0(esk4_0,esk5_0)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_40, plain, (![X3]:k2_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_41,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_43,negated_conjecture,(k4_xboole_0(esk5_0,k4_xboole_0(esk5_0,esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_45, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_46,plain,(r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_37, c_0_31])).
fof(c_0_47, plain, (![X3]:![X4]:k2_xboole_0(k3_xboole_0(X3,X4),k4_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t51_xboole_1])).
fof(c_0_48, plain, (![X3]:![X4]:k4_xboole_0(k2_xboole_0(X3,X4),X4)=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_49,negated_conjecture,(k2_xboole_0(esk3_0,k2_xboole_0(esk2_0,X1))=k2_xboole_0(esk4_0,k2_xboole_0(esk5_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_38])).
cnf(c_0_50,plain,(k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51,negated_conjecture,(r1_tarski(k4_xboole_0(X1,esk3_0),esk2_0)|~r1_tarski(X1,k2_xboole_0(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_41, c_0_39])).
cnf(c_0_52,negated_conjecture,(k4_xboole_0(esk5_0,esk3_0)=esk5_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_53,negated_conjecture,(r1_xboole_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_54, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_55,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56,negated_conjecture,(r1_tarski(esk2_0,k2_xboole_0(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_46, c_0_39])).
cnf(c_0_57,negated_conjecture,(r1_tarski(esk3_0,k2_xboole_0(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_37, c_0_39])).
cnf(c_0_58,plain,(k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_59, plain, (![X3]:![X4]:k2_xboole_0(X3,k4_xboole_0(X4,X3))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_60,plain,(k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61,negated_conjecture,(k2_xboole_0(esk4_0,k2_xboole_0(esk2_0,esk5_0))=k2_xboole_0(esk4_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_39]), c_0_31])).
cnf(c_0_62,negated_conjecture,(r1_tarski(esk5_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_46]), c_0_52])).
cnf(c_0_63,negated_conjecture,(k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_53])).
cnf(c_0_64,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65,negated_conjecture,(k2_xboole_0(esk2_0,k2_xboole_0(esk4_0,esk5_0))=k2_xboole_0(esk4_0,esk5_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_66,negated_conjecture,(k2_xboole_0(esk3_0,k2_xboole_0(esk4_0,esk5_0))=k2_xboole_0(esk4_0,esk5_0)), inference(spm,[status(thm)],[c_0_55, c_0_57])).
fof(c_0_67, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_xboole_0(X5,X6))|r1_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_68,plain,(k2_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[c_0_58, c_0_21])).
cnf(c_0_69,plain,(k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70,negated_conjecture,(k4_xboole_0(k2_xboole_0(esk4_0,esk5_0),k2_xboole_0(esk2_0,esk5_0))=k4_xboole_0(esk4_0,k2_xboole_0(esk2_0,esk5_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_71,negated_conjecture,(k2_xboole_0(esk2_0,esk5_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_62]), c_0_31])).
cnf(c_0_72,negated_conjecture,(k4_xboole_0(k2_xboole_0(esk4_0,esk5_0),esk2_0)=k4_xboole_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_60, c_0_39])).
cnf(c_0_73,negated_conjecture,(k4_xboole_0(esk4_0,esk2_0)=esk4_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_63]), c_0_44])).
cnf(c_0_74,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_21]), c_0_21])).
cnf(c_0_75,plain,(r2_hidden(esk1_2(X1,X2),k3_xboole_0(X1,X2))|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_76,negated_conjecture,(k2_xboole_0(esk4_0,k2_xboole_0(esk5_0,k2_xboole_0(esk4_0,esk5_0)))=k2_xboole_0(esk4_0,esk5_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_65]), c_0_66])).
cnf(c_0_77,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X3,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_78,plain,(k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_31]), c_0_69]), c_0_31])).
cnf(c_0_79,negated_conjecture,(k4_xboole_0(esk3_0,esk2_0)=esk4_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_71]), c_0_72]), c_0_71]), c_0_73])).
cnf(c_0_80,plain,(~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_26, c_0_74])).
cnf(c_0_81,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_75, c_0_21])).
cnf(c_0_82,negated_conjecture,(r1_tarski(k2_xboole_0(esk5_0,k2_xboole_0(esk4_0,esk5_0)),k2_xboole_0(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_46, c_0_76])).
cnf(c_0_83,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(k2_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_77, c_0_37])).
cnf(c_0_84,negated_conjecture,(k2_xboole_0(esk4_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_31])).
cnf(c_0_85,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_86,negated_conjecture,(k2_xboole_0(esk5_0,k2_xboole_0(esk4_0,esk5_0))=k2_xboole_0(esk4_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_82]), c_0_38]), c_0_38]), c_0_76])).
cnf(c_0_87,negated_conjecture,(r1_xboole_0(esk4_0,X1)|~r1_xboole_0(esk3_0,X1)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_88,negated_conjecture,(r1_xboole_0(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_85, c_0_35])).
cnf(c_0_89,negated_conjecture,(k2_xboole_0(esk3_0,k2_xboole_0(X1,esk2_0))=k2_xboole_0(esk4_0,k2_xboole_0(esk5_0,X1))), inference(spm,[status(thm)],[c_0_49, c_0_31])).
cnf(c_0_90,negated_conjecture,(k2_xboole_0(esk2_0,k2_xboole_0(esk5_0,k4_xboole_0(esk4_0,k2_xboole_0(esk2_0,esk5_0))))=k2_xboole_0(esk4_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_38]), c_0_38]), c_0_86]), c_0_65])).
cnf(c_0_91,plain,(k2_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(X3,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_31, c_0_38])).
cnf(c_0_92,negated_conjecture,(r1_xboole_0(esk4_0,esk5_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_93,negated_conjecture,(k2_xboole_0(esk4_0,k2_xboole_0(esk3_0,esk5_0))=k2_xboole_0(esk4_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_39]), c_0_66]), c_0_31])).
cnf(c_0_94,negated_conjecture,(k2_xboole_0(esk4_0,esk5_0)=k2_xboole_0(esk4_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_71]), c_0_73]), c_0_31]), c_0_91]), c_0_31]), c_0_71])).
cnf(c_0_95,negated_conjecture,(k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk5_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_92])).
cnf(c_0_96,negated_conjecture,(k4_xboole_0(esk5_0,esk5_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_43, c_0_52])).
cnf(c_0_97,negated_conjecture,(k2_xboole_0(esk4_0,k2_xboole_0(esk3_0,esk5_0))=k2_xboole_0(esk4_0,esk2_0)), inference(rw,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_98,negated_conjecture,(k2_xboole_0(esk4_0,k2_xboole_0(esk3_0,X1))=k2_xboole_0(esk3_0,X1)), inference(spm,[status(thm)],[c_0_38, c_0_84])).
cnf(c_0_99,negated_conjecture,(k4_xboole_0(esk4_0,esk5_0)=esk4_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_95]), c_0_44])).
cnf(c_0_100,negated_conjecture,(k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk5_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_52]), c_0_96])).
cnf(c_0_101,negated_conjecture,(k2_xboole_0(esk3_0,esk5_0)=k2_xboole_0(esk4_0,esk2_0)), inference(rw,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_102,negated_conjecture,(k4_xboole_0(k2_xboole_0(esk4_0,esk2_0),esk5_0)=esk4_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_94]), c_0_99])).
cnf(c_0_103,negated_conjecture,(k4_xboole_0(esk3_0,esk5_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_100]), c_0_44])).
cnf(c_0_104,negated_conjecture,(esk4_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_105,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_101]), c_0_102]), c_0_103]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 71
# Proof object formula steps           : 35
# Proof object conjectures             : 45
# Proof object clause conjectures      : 42
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 20
# Processed clauses                    : 595
# ...of these trivial                  : 48
# ...subsumed                          : 305
# ...remaining for further processing  : 242
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 78
# Generated clauses                    : 2495
# ...of the previous two non-trivial   : 1915
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2495
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 135
#    Positive orientable unit clauses  : 69
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 60
# Current number of unprocessed clauses: 897
# ...number of literals in the above   : 1351
# Current number of archived formulas  : 0
# Current number of archived clauses   : 108
# Clause-clause subsumption calls (NU) : 3919
# Rec. Clause-clause subsumption calls : 2427
# Non-unit clause-clause subsumptions  : 214
# Unit Clause-clause subsumption calls : 43
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 183
# BW rewrite match successes           : 118
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 24805

# -------------------------------------------------
# User time                : 0.036 s
# System time              : 0.004 s
# Total time               : 0.040 s
# Maximum resident set size: 45328 pages
