# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t46_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4))), file('out/t72_enumset1__6_premises__rerun__rerun.E_input', t46_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t72_enumset1__6_premises__rerun__rerun.E_input', t71_enumset1)).
fof(t50_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), file('out/t72_enumset1__6_premises__rerun__rerun.E_input', t50_enumset1)).
fof(t69_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), file('out/t72_enumset1__6_premises__rerun__rerun.E_input', t69_enumset1)).
fof(t72_enumset1, conjecture, (![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X2,X4,X1)), file('out/t72_enumset1__6_premises__rerun__rerun.E_input', t72_enumset1)).
fof(c_0_5, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X5,X6,X7),k1_tarski(X8))), inference(variable_rename,[status(thm)],[t46_enumset1])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8)=k2_enumset1(X5,X6,X7,X8)), inference(variable_rename,[status(thm)],[t50_enumset1])).
fof(c_0_8, plain, (![X5]:![X6]:![X7]:![X8]:k2_enumset1(X5,X6,X7,X8)=k2_enumset1(X6,X8,X5,X7)), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_9,plain,(k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X2,X4,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k2_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k1_tarski(X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_11])).
cnf(c_0_14,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X2,X2,X2,X2,X2,X4,X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_11]), c_0_11])).
cnf(c_0_15,plain,(k2_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X2),k1_tarski(X4))=k6_enumset1(X2,X2,X2,X2,X2,X1,X3,X4)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_16, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X2,X4,X1))), inference(assume_negation,[status(cth)],[t72_enumset1])).
cnf(c_0_17,plain,(k2_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X1,X3),k1_tarski(X4))=k6_enumset1(X1,X1,X1,X1,X1,X3,X2,X4)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18,plain,(k2_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X2),k1_tarski(X4))=k6_enumset1(X1,X1,X1,X1,X1,X4,X2,X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_19, negated_conjecture, (k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0)!=k2_enumset1(esk3_0,esk2_0,esk4_0,esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
cnf(c_0_20,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X1)=k6_enumset1(X1,X1,X1,X1,X1,X3,X3,X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X1,X3)=k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3)), inference(spm,[status(thm)],[c_0_13, c_0_17])).
cnf(c_0_22,negated_conjecture,(k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0)!=k2_enumset1(esk3_0,esk2_0,esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,plain,(k2_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X2,X1),k1_tarski(X3))=k6_enumset1(X1,X1,X1,X1,X1,X3,X2,X2)), inference(spm,[status(thm)],[c_0_18, c_0_20])).
cnf(c_0_24,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X3,X2,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_21]), c_0_17])).
cnf(c_0_25,negated_conjecture,(k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk2_0,esk4_0,esk1_0)!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_11]), c_0_11])).
cnf(c_0_26,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X3)=k6_enumset1(X1,X1,X1,X1,X1,X3,X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_20]), c_0_23])).
cnf(c_0_27,plain,(k2_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X3),k1_tarski(X4))=k6_enumset1(X1,X1,X1,X1,X1,X4,X3,X2)), inference(spm,[status(thm)],[c_0_18, c_0_24])).
cnf(c_0_28,negated_conjecture,(k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk2_0,esk4_0,esk1_0)!=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk2_0)), inference(rw,[status(thm)],[c_0_25, c_0_14])).
cnf(c_0_29,plain,(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X4,X3,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_26]), c_0_27])).
cnf(c_0_30,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 20
# Proof object formula steps           : 11
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 3
# Processed clauses                    : 160
# ...of these trivial                  : 2
# ...subsumed                          : 120
# ...remaining for further processing  : 38
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 3
# Generated clauses                    : 1316
# ...of the previous two non-trivial   : 1180
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1316
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 22
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 846
# ...number of literals in the above   : 846
# Current number of archived formulas  : 0
# Current number of archived clauses   : 13
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 390
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1118
# BW rewrite match successes           : 339
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4030

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.004 s
# Total time               : 0.020 s
# Maximum resident set size: 45260 pages
