# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d20_borsuk_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>(r1_borsuk_1(X1,X2)<=>?[X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))&v3_borsuk_1(X3,X1,X2)))))), file('out/t75_tex_2__2_premises__rerun.E_input', d20_borsuk_1)).
fof(t74_tex_2, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v4_tex_2(X2,X1)&m1_pre_topc(X2,X1)))=>?[X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))&v3_borsuk_1(X3,X1,X2))))), file('out/t75_tex_2__2_premises__rerun.E_input', t74_tex_2)).
fof(t75_tex_2, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v4_tex_2(X2,X1)&m1_pre_topc(X2,X1)))=>r1_borsuk_1(X1,X2)))), file('out/t75_tex_2__2_premises__rerun.E_input', t75_tex_2)).
fof(c_0_3, plain, (![X4]:![X5]:![X7]:((((((v1_funct_1(esk1_2(X4,X5))|~r1_borsuk_1(X4,X5))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((((v1_funct_2(esk1_2(X4,X5),u1_struct_0(X4),u1_struct_0(X5))|~r1_borsuk_1(X4,X5))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&((((v5_pre_topc(esk1_2(X4,X5),X4,X5)|~r1_borsuk_1(X4,X5))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(((m1_subset_1(esk1_2(X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~r1_borsuk_1(X4,X5))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))))&(((v3_borsuk_1(esk1_2(X4,X5),X4,X5)|~r1_borsuk_1(X4,X5))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))&(((((~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X5))|(~v5_pre_topc(X7,X4,X5)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5)))))))|~v3_borsuk_1(X7,X4,X5))|r1_borsuk_1(X4,X5))|(v2_struct_0(X5)|~m1_pre_topc(X5,X4)))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d20_borsuk_1])])])])])])])])).
fof(c_0_4, plain, (![X4]:![X5]:((((v1_funct_1(esk2_2(X4,X5))|(v2_struct_0(X5)|(~v4_tex_2(X5,X4)|~m1_pre_topc(X5,X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v3_tdlat_3(X4)|~l1_pre_topc(X4)))))&(((v1_funct_2(esk2_2(X4,X5),u1_struct_0(X4),u1_struct_0(X5))|(v2_struct_0(X5)|(~v4_tex_2(X5,X4)|~m1_pre_topc(X5,X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v3_tdlat_3(X4)|~l1_pre_topc(X4)))))&(((v5_pre_topc(esk2_2(X4,X5),X4,X5)|(v2_struct_0(X5)|(~v4_tex_2(X5,X4)|~m1_pre_topc(X5,X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v3_tdlat_3(X4)|~l1_pre_topc(X4)))))&((m1_subset_1(esk2_2(X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|(v2_struct_0(X5)|(~v4_tex_2(X5,X4)|~m1_pre_topc(X5,X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v3_tdlat_3(X4)|~l1_pre_topc(X4))))))))&((v3_borsuk_1(esk2_2(X4,X5),X4,X5)|(v2_struct_0(X5)|(~v4_tex_2(X5,X4)|~m1_pre_topc(X5,X4))))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|(~v3_tdlat_3(X4)|~l1_pre_topc(X4))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t74_tex_2])])])])])])])])).
fof(c_0_5, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v4_tex_2(X2,X1)&m1_pre_topc(X2,X1)))=>r1_borsuk_1(X1,X2))))), inference(assume_negation,[status(cth)],[t75_tex_2])).
cnf(c_0_6,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_borsuk_1(X1,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v3_borsuk_1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v5_pre_topc(X3,X1,X2)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk2_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X1)|~v3_tdlat_3(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v4_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v1_funct_1(esk2_2(X1,X2))|~l1_pre_topc(X1)|~v3_tdlat_3(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v4_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v1_funct_2(esk2_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~l1_pre_topc(X1)|~v3_tdlat_3(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v4_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(esk2_2(X1,X2),X1,X2)|~l1_pre_topc(X1)|~v3_tdlat_3(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v4_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v3_borsuk_1(esk2_2(X1,X2),X1,X2)|~l1_pre_topc(X1)|~v3_tdlat_3(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v4_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_12, negated_conjecture, (((~v2_struct_0(esk3_0)&(v2_pre_topc(esk3_0)&(v3_tdlat_3(esk3_0)&l1_pre_topc(esk3_0))))&((~v2_struct_0(esk4_0)&(v4_tex_2(esk4_0,esk3_0)&m1_pre_topc(esk4_0,esk3_0)))&~r1_borsuk_1(esk3_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_13,plain,(r1_borsuk_1(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~v4_tex_2(X2,X1)|~v3_tdlat_3(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11])).
cnf(c_0_14,negated_conjecture,(v4_tex_2(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15,negated_conjecture,(v3_tdlat_3(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(m1_pre_topc(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,negated_conjecture,(v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(~r1_borsuk_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18])]), c_0_19]), c_0_20]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 16
# Proof object formula steps           : 7
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 39
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 39
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 3
# ...of the previous two non-trivial   : 1
# Contextual simplify-reflections      : 4
# Paramodulations                      : 3
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 20
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 51
# Rec. Clause-clause subsumption calls : 4
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2052

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45308 pages
