# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t70_enumset1)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t95_xboole_1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t71_enumset1)).
fof(t75_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3)))))))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t75_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X1)=X1), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', idempotence_k3_xboole_0)).
fof(idempotence_k2_xboole_0, axiom, (![X1]:![X2]:k2_xboole_0(X1,X1)=X1), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', idempotence_k2_xboole_0)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t92_xboole_1, axiom, (![X1]:k5_xboole_0(X1,X1)=k1_xboole_0), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t92_xboole_1)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t91_xboole_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t7_xboole_1)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t5_boole)).
fof(l45_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k2_tarski(X2,X3))<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3)))))))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', l45_zfmisc_1)).
fof(t12_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2)), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t12_xboole_1)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t3_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', d3_tarski)).
fof(t66_xboole_1, axiom, (![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1))))), file('out/t75_zfmisc_1__19_premises__rerun__rerun.E_input', t66_xboole_1)).
fof(c_0_17, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_18, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_19, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_20,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_22, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_23, negated_conjecture, (~(![X1]:![X2]:![X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))))), inference(assume_negation,[status(cth)],[t75_zfmisc_1])).
fof(c_0_24, plain, (![X3]:k3_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_25,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, (![X3]:k2_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_29, negated_conjecture, ((((esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk2_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&(esk1_0!=k2_tarski(esk2_0,esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0))))&(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0|(esk1_0=k1_xboole_0|(esk1_0=k1_tarski(esk2_0)|(esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])])).
fof(c_0_30, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_31,plain,(k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
fof(c_0_33, plain, (![X2]:k5_xboole_0(X2,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t92_xboole_1])).
cnf(c_0_34,plain,(k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35,negated_conjecture,(esk1_0=k2_tarski(esk2_0,esk3_0)|esk1_0=k1_tarski(esk3_0)|esk1_0=k1_tarski(esk2_0)|esk1_0=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_37, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_38,plain,(k5_xboole_0(k5_xboole_0(X1,X1),k3_tarski(k2_enumset1(X1,X1,X1,X1)))=X1), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_26]), c_0_27])).
cnf(c_0_41,negated_conjecture,(esk1_0=k1_xboole_0|esk1_0=k1_tarski(esk2_0)|esk1_0=k1_tarski(esk3_0)|esk1_0=k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)))))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_21]), c_0_21]), c_0_36]), c_0_27]), c_0_27]), c_0_32])).
cnf(c_0_42,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_43,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
fof(c_0_44, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_45,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))))=k1_xboole_0|k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_39]), c_0_43])).
fof(c_0_47, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_48,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_tarski(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_49, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:((~r1_tarski(X4,k2_tarski(X5,X6))|(X4=k1_xboole_0|(X4=k1_tarski(X5)|(X4=k1_tarski(X6)|X4=k2_tarski(X5,X6)))))&((X4!=k1_xboole_0|r1_tarski(X4,k2_tarski(X5,X6)))&((X4!=k1_tarski(X5)|r1_tarski(X4,k2_tarski(X5,X6)))&((X4!=k1_tarski(X6)|r1_tarski(X4,k2_tarski(X5,X6)))&(X4!=k2_tarski(X5,X6)|r1_tarski(X4,k2_tarski(X5,X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l45_zfmisc_1])])])])])])).
cnf(c_0_50,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))=k1_xboole_0|k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53,negated_conjecture,(esk1_0!=k1_tarski(esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_21]), c_0_36]), c_0_27]), c_0_32])).
fof(c_0_54, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k2_xboole_0(X3,X4)=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_55,plain,(r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56,plain,(X1=k2_tarski(X2,X3)|X1=k1_tarski(X3)|X1=k1_tarski(X2)|X1=k1_xboole_0|~r1_tarski(X1,k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57,plain,(r1_tarski(X1,k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_26]), c_0_27])).
cnf(c_0_58,negated_conjecture,(k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)))=k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)|k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_51]), c_0_52])).
cnf(c_0_59,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_60,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k1_tarski(esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_53, c_0_42])).
cnf(c_0_61,plain,(k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62,plain,(r1_tarski(X1,k2_enumset1(X2,X2,X2,X3))|X1!=k1_tarski(X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_21]), c_0_27])).
cnf(c_0_63,plain,(X1=k1_xboole_0|X1=k1_tarski(X3)|X1=k1_tarski(X2)|X1=k2_enumset1(X2,X2,X2,X3)|~r1_tarski(X1,k2_enumset1(X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_21]), c_0_21]), c_0_27]), c_0_27])).
cnf(c_0_64,negated_conjecture,(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_xboole_0=esk1_0|r1_tarski(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65,negated_conjecture,(esk1_0!=k1_tarski(esk2_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_21]), c_0_36]), c_0_27]), c_0_32])).
cnf(c_0_66,plain,(r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_67,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k1_tarski(esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_60, c_0_46])).
cnf(c_0_68,plain,(k3_tarski(k2_enumset1(X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_26]), c_0_27])).
cnf(c_0_69,plain,(r1_tarski(k1_tarski(X1),k2_enumset1(X2,X2,X2,X1))), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_70,negated_conjecture,(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_71,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_72,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k2_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_73,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k1_tarski(esk2_0)!=esk1_0), inference(rw,[status(thm)],[c_0_65, c_0_42])).
cnf(c_0_74,plain,(r1_tarski(X1,k2_enumset1(X2,X2,X2,X3))|X1!=k1_tarski(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_21]), c_0_27])).
cnf(c_0_75,negated_conjecture,(k1_tarski(esk3_0)!=esk1_0|~r1_tarski(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_39])])).
cnf(c_0_76,negated_conjecture,(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0|r1_tarski(esk1_0,k2_enumset1(X1,X1,X1,esk3_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
fof(c_0_77, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk4_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk4_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_78, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_79,negated_conjecture,(esk1_0!=k1_xboole_0|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_21]), c_0_36]), c_0_27]), c_0_32])).
cnf(c_0_80,negated_conjecture,(esk1_0!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_21]), c_0_21]), c_0_36]), c_0_27]), c_0_27]), c_0_32])).
cnf(c_0_81,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k1_tarski(esk2_0)!=esk1_0), inference(rw,[status(thm)],[c_0_73, c_0_46])).
cnf(c_0_82,plain,(r1_tarski(k1_tarski(X1),k2_enumset1(X1,X1,X1,X2))), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_83,negated_conjecture,(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_70])).
cnf(c_0_84,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_85,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
fof(c_0_86, plain, (![X2]:![X2]:((r1_xboole_0(X2,X2)|X2!=k1_xboole_0)&(X2=k1_xboole_0|~r1_xboole_0(X2,X2)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])])])).
cnf(c_0_87,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k1_xboole_0!=esk1_0), inference(rw,[status(thm)],[c_0_79, c_0_42])).
cnf(c_0_88,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_80, c_0_42])).
cnf(c_0_89,negated_conjecture,(k1_tarski(esk2_0)!=esk1_0|~r1_tarski(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_68]), c_0_39])])).
cnf(c_0_90,negated_conjecture,(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0|r1_tarski(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,X1))), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_91,plain,(r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92,plain,(r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_93,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k1_xboole_0!=esk1_0), inference(rw,[status(thm)],[c_0_87, c_0_46])).
cnf(c_0_94,negated_conjecture,(k5_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_88, c_0_46])).
cnf(c_0_95,negated_conjecture,(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_83])).
cnf(c_0_96,plain,(r1_tarski(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_91, c_0_85])).
cnf(c_0_97,plain,(r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_92])).
cnf(c_0_98,negated_conjecture,(k1_xboole_0!=esk1_0|~r1_tarski(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_68]), c_0_39])])).
cnf(c_0_99,negated_conjecture,(k1_xboole_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_40]), c_0_39])])).
cnf(c_0_100,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_101,negated_conjecture,(~r1_tarski(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_99])])).
cnf(c_0_102,plain,(r1_tarski(esk1_0,X1)), inference(rw,[status(thm)],[c_0_100, c_0_99])).
cnf(c_0_103,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 104
# Proof object clause steps            : 69
# Proof object formula steps           : 35
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 75
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 25
# Processed clauses                    : 129
# ...of these trivial                  : 5
# ...subsumed                          : 14
# ...remaining for further processing  : 110
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 32
# Generated clauses                    : 269
# ...of the previous two non-trivial   : 172
# Contextual simplify-reflections      : 5
# Paramodulations                      : 264
# Factorizations                       : 0
# Equation resolutions                 : 5
# Current number of processed clauses  : 38
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 19
# Current number of unprocessed clauses: 56
# ...number of literals in the above   : 85
# Current number of archived formulas  : 0
# Current number of archived clauses   : 72
# Clause-clause subsumption calls (NU) : 368
# Rec. Clause-clause subsumption calls : 264
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 34
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 48
# BW rewrite match successes           : 28
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4303

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 43040 pages
