# No SInE strategy applied
# Auto-Mode selected heuristic G_E___110_C45_F1_PI_SE_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t75_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3)))))))), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t75_zfmisc_1)).
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t95_xboole_1)).
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t91_xboole_1, axiom, (![X1]:![X2]:![X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t91_xboole_1)).
fof(t92_xboole_1, axiom, (![X1]:k5_xboole_0(X1,X1)=k1_xboole_0), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t92_xboole_1)).
fof(t100_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t100_xboole_1)).
fof(t5_boole, axiom, (![X1]:k5_xboole_0(X1,k1_xboole_0)=X1), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t5_boole)).
fof(t39_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t39_xboole_1)).
fof(t45_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1)))), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t45_xboole_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t7_xboole_1)).
fof(t42_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(r1_tarski(X1,k2_tarski(X2,X3))<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3)))))))), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', t42_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X1)=X1), file('out/t75_zfmisc_1__18_premises__rerun__rerun.E_input', idempotence_k3_xboole_0)).
fof(c_0_12, negated_conjecture, (~(![X1]:![X2]:![X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))))), inference(assume_negation,[status(cth)],[t75_zfmisc_1])).
fof(c_0_13, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
fof(c_0_14, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:k5_xboole_0(k5_xboole_0(X4,X5),X6)=k5_xboole_0(X4,k5_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t91_xboole_1])).
fof(c_0_16, plain, (![X2]:k5_xboole_0(X2,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_17, negated_conjecture, ((((esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk2_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&(esk1_0!=k2_tarski(esk2_0,esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0))))&(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0|(esk1_0=k1_xboole_0|(esk1_0=k1_tarski(esk2_0)|(esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])).
fof(c_0_18, plain, (![X3]:![X4]:k4_xboole_0(X3,X4)=k5_xboole_0(X3,k3_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_19,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, (![X2]:k5_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_24,negated_conjecture,(esk1_0=k2_tarski(esk2_0,esk3_0)|esk1_0=k1_tarski(esk3_0)|esk1_0=k1_tarski(esk2_0)|esk1_0=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_27, plain, (![X3]:![X4]:k2_xboole_0(X3,k4_xboole_0(X4,X3))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_28,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=k5_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29,plain,(k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,negated_conjecture,(esk1_0=k1_xboole_0|esk1_0=k1_tarski(esk2_0)|esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0)),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0)))))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
fof(c_0_31, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|X4=k2_xboole_0(X3,k4_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
cnf(c_0_32,plain,(k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_tarski(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_34, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_35,plain,(k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_22]), c_0_29])).
cnf(c_0_36,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))))=k1_xboole_0|k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[c_0_30, c_0_21])).
cnf(c_0_37,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_38,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_39,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k2_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40,plain,(X1=k2_xboole_0(X2,k4_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41,plain,(k3_tarski(k2_tarski(X1,k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X1),k3_tarski(k2_tarski(X2,X1))))))=k3_tarski(k2_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_20]), c_0_20]), c_0_25]), c_0_26])).
fof(c_0_42, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:((~r1_tarski(X4,k2_tarski(X5,X6))|(X4=k1_xboole_0|(X4=k1_tarski(X5)|(X4=k1_tarski(X6)|X4=k2_tarski(X5,X6)))))&((X4!=k1_xboole_0|r1_tarski(X4,k2_tarski(X5,X6)))&((X4!=k1_tarski(X5)|r1_tarski(X4,k2_tarski(X5,X6)))&((X4!=k1_tarski(X6)|r1_tarski(X4,k2_tarski(X5,X6)))&(X4!=k2_tarski(X5,X6)|r1_tarski(X4,k2_tarski(X5,X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_zfmisc_1])])])])])])).
cnf(c_0_43,negated_conjecture,(esk1_0!=k1_tarski(esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0)),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_25]), c_0_26])).
cnf(c_0_44,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45,plain,(k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[c_0_28, c_0_35])).
cnf(c_0_46,negated_conjecture,(k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))=k1_xboole_0|k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_28]), c_0_35])).
cnf(c_0_47,negated_conjecture,(esk1_0!=k1_xboole_0|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0)),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_25]), c_0_26])).
cnf(c_0_48,negated_conjecture,(esk1_0!=k1_tarski(esk2_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0)),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_25]), c_0_26])).
cnf(c_0_49,negated_conjecture,(esk1_0!=k2_tarski(esk2_0,esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0)),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_25]), c_0_26])).
cnf(c_0_50,plain,(X1=k3_tarski(k2_tarski(X2,k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_tarski(X1,X2))))))|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_20]), c_0_25]), c_0_26])).
cnf(c_0_51,plain,(k3_tarski(k2_tarski(X1,k5_xboole_0(X2,k5_xboole_0(X2,k5_xboole_0(X1,k3_tarski(k2_tarski(X2,X1)))))))=k3_tarski(k2_tarski(X1,X2))), inference(rw,[status(thm)],[c_0_41, c_0_21])).
cnf(c_0_52,plain,(r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_53, plain, (![X3]:k3_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_54,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))))!=k1_xboole_0|k1_tarski(esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_43, c_0_21])).
cnf(c_0_55,plain,(r1_tarski(X1,k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_44, c_0_20])).
cnf(c_0_56,negated_conjecture,(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0)))=k2_tarski(esk2_0,esk3_0)|k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_29])).
cnf(c_0_57,plain,(X1=k2_tarski(X2,X3)|X1=k1_tarski(X3)|X1=k1_tarski(X2)|X1=k1_xboole_0|~r1_tarski(X1,k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_58,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))))!=k1_xboole_0|k1_xboole_0!=esk1_0), inference(rw,[status(thm)],[c_0_47, c_0_21])).
cnf(c_0_59,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))))!=k1_xboole_0|k1_tarski(esk2_0)!=esk1_0), inference(rw,[status(thm)],[c_0_48, c_0_21])).
cnf(c_0_60,negated_conjecture,(k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))))!=k1_xboole_0|k2_tarski(esk2_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_49, c_0_21])).
cnf(c_0_61,plain,(k3_tarski(k2_tarski(X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_21]), c_0_51])).
cnf(c_0_62,plain,(r1_tarski(k1_tarski(X1),k2_tarski(X2,X1))), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_63,plain,(k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64,plain,(r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_65,negated_conjecture,(k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))!=k1_xboole_0|k1_tarski(esk3_0)!=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_28]), c_0_35])).
cnf(c_0_66,negated_conjecture,(k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_67,negated_conjecture,(k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))!=k1_xboole_0|k1_xboole_0!=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_28]), c_0_35])).
cnf(c_0_68,negated_conjecture,(k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))!=k1_xboole_0|k1_tarski(esk2_0)!=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_28]), c_0_35])).
cnf(c_0_69,negated_conjecture,(k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))!=k1_xboole_0|k2_tarski(esk2_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_28]), c_0_35])).
cnf(c_0_70,plain,(k3_tarski(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)))=k2_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_71,plain,(k5_xboole_0(k5_xboole_0(X1,X1),k3_tarski(k2_tarski(X1,X1)))=X1), inference(rw,[status(thm)],[c_0_63, c_0_26])).
cnf(c_0_72,plain,(r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_73,negated_conjecture,(k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))!=k1_xboole_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_68]), c_0_69])).
cnf(c_0_74,negated_conjecture,(k3_tarski(k2_tarski(esk1_0,k2_tarski(X1,esk3_0)))=k2_tarski(X1,esk3_0)|k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(spm,[status(thm)],[c_0_70, c_0_66])).
cnf(c_0_75,plain,(k5_xboole_0(X1,k5_xboole_0(X1,k3_tarski(k2_tarski(X1,X1))))=X1), inference(rw,[status(thm)],[c_0_71, c_0_21])).
cnf(c_0_76,plain,(k3_tarski(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)))=k2_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_61, c_0_72])).
cnf(c_0_77,negated_conjecture,(k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_22])])).
cnf(c_0_78,plain,(k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(k3_tarski(k2_tarski(X1,X1)),X2)))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_75]), c_0_21])).
cnf(c_0_79,negated_conjecture,(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,X1)))=k2_tarski(esk2_0,X1)|k2_tarski(esk2_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_80,plain,(k5_xboole_0(X1,k3_tarski(k2_tarski(X1,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_22]), c_0_29]), c_0_22])).
cnf(c_0_81,plain,(r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_82,negated_conjecture,(k2_tarski(esk2_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_79]), c_0_22])])).
cnf(c_0_83,plain,(k3_tarski(k2_tarski(X1,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_80]), c_0_29])).
cnf(c_0_84,plain,(r1_tarski(k1_xboole_0,k2_tarski(X1,X2))), inference(er,[status(thm)],[c_0_81])).
cnf(c_0_85,negated_conjecture,(k1_xboole_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_82]), c_0_83]), c_0_22])])).
cnf(c_0_86,plain,(k3_tarski(k2_tarski(k1_xboole_0,k2_tarski(X1,X2)))=k2_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_61, c_0_84])).
cnf(c_0_87,negated_conjecture,(k5_xboole_0(k2_tarski(esk2_0,esk3_0),k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0))))!=esk1_0), inference(rw,[status(thm)],[c_0_73, c_0_85])).
cnf(c_0_88,plain,(k3_tarski(k2_tarski(esk1_0,k2_tarski(X1,X2)))=k2_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_86, c_0_85])).
cnf(c_0_89,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_88]), c_0_22]), c_0_85])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 65
# Proof object formula steps           : 25
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 66
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 17
# Processed clauses                    : 971
# ...of these trivial                  : 562
# ...subsumed                          : 244
# ...remaining for further processing  : 165
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 25
# Backward-rewritten                   : 37
# Generated clauses                    : 27207
# ...of the previous two non-trivial   : 23554
# Contextual simplify-reflections      : 12
# Paramodulations                      : 27203
# Factorizations                       : 0
# Equation resolutions                 : 4
# Current number of processed clauses  : 82
#    Positive orientable unit clauses  : 77
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 1
# Current number of unprocessed clauses: 17896
# ...number of literals in the above   : 17900
# Current number of archived formulas  : 0
# Current number of archived clauses   : 82
# Clause-clause subsumption calls (NU) : 430
# Rec. Clause-clause subsumption calls : 100
# Non-unit clause-clause subsumptions  : 73
# Unit Clause-clause subsumption calls : 40
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1638
# BW rewrite match successes           : 85
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 39876058

# -------------------------------------------------
# User time                : 6.152 s
# System time              : 0.124 s
# Total time               : 6.276 s
# Maximum resident set size: 89344 pages
