# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t95_xboole_1, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t95_xboole_1)).
fof(l51_zfmisc_1, axiom, (![X1]:![X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', l51_zfmisc_1)).
fof(t51_xboole_1, axiom, (![X1]:![X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t51_xboole_1)).
fof(t46_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t46_xboole_1)).
fof(t1_boole, axiom, (![X1]:k2_xboole_0(X1,k1_xboole_0)=X1), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t1_boole)).
fof(t40_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t40_xboole_1)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(t52_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(X1,k4_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X3))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t52_xboole_1)).
fof(t41_xboole_1, axiom, (![X1]:![X2]:![X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t41_xboole_1)).
fof(t45_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1)))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t45_xboole_1)).
fof(t7_xboole_1, axiom, (![X1]:![X2]:r1_tarski(X1,k2_xboole_0(X1,X2))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t7_xboole_1)).
fof(t88_xboole_1, axiom, (![X1]:![X2]:(r1_xboole_0(X1,X2)=>k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1)), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t88_xboole_1)).
fof(t81_xboole_1, axiom, (![X1]:![X2]:![X3]:(r1_xboole_0(X1,k4_xboole_0(X2,X3))=>r1_xboole_0(X2,k4_xboole_0(X1,X3)))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t81_xboole_1)).
fof(l31_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)<=>(~(r2_hidden(X1,X3))&(r2_hidden(X2,X3)|X1=X2)))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', l31_zfmisc_1)).
fof(t79_xboole_1, axiom, (![X1]:![X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2)), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t79_xboole_1)).
fof(t73_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3)))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t73_zfmisc_1)).
fof(t75_zfmisc_1, conjecture, (![X1]:![X2]:![X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3)))))))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t75_zfmisc_1)).
fof(t37_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t37_zfmisc_1)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', d10_xboole_0)).
fof(t65_zfmisc_1, axiom, (![X1]:![X2]:![X3]:~((~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0)&(~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1))&(~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X2))&~(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2))))))), file('out/t75_zfmisc_1__34_premises__rerun__rerun.E_input', t65_zfmisc_1)).
fof(c_0_20, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k5_xboole_0(k5_xboole_0(X3,X4),k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t95_xboole_1])).
fof(c_0_21, plain, (![X3]:![X4]:k3_tarski(k2_tarski(X3,X4))=k2_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_22, plain, (![X3]:![X4]:k2_xboole_0(k3_xboole_0(X3,X4),k4_xboole_0(X3,X4))=X3), inference(variable_rename,[status(thm)],[t51_xboole_1])).
cnf(c_0_23,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_25, plain, (![X3]:![X4]:k4_xboole_0(X3,k2_xboole_0(X3,X4))=k1_xboole_0), inference(variable_rename,[status(thm)],[t46_xboole_1])).
fof(c_0_26, plain, (![X2]:k2_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_27, plain, (![X3]:![X4]:k4_xboole_0(k2_xboole_0(X3,X4),X4)=k4_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_28,plain,(k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_30, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_31, plain, (![X4]:![X5]:![X6]:k4_xboole_0(X4,k4_xboole_0(X5,X6))=k2_xboole_0(k4_xboole_0(X4,X5),k3_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[t52_xboole_1])).
cnf(c_0_32,plain,(k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,plain,(k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, (![X4]:![X5]:![X6]:k4_xboole_0(k4_xboole_0(X4,X5),X6)=k4_xboole_0(X4,k2_xboole_0(X5,X6))), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_35, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|X4=k2_xboole_0(X3,k4_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
fof(c_0_36, plain, (![X3]:![X4]:r1_tarski(X3,k2_xboole_0(X3,X4))), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_37,plain,(k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38,plain,(k3_tarski(k2_tarski(k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_tarski(X1,X2))),k4_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_24]), c_0_29])).
cnf(c_0_39,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41,plain,(k4_xboole_0(X1,k3_tarski(k2_tarski(X1,X2)))=k1_xboole_0), inference(rw,[status(thm)],[c_0_32, c_0_24])).
cnf(c_0_42,plain,(k3_tarski(k2_tarski(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_33, c_0_24])).
cnf(c_0_43,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44,plain,(X1=k2_xboole_0(X2,k4_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45,plain,(r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46,plain,(k4_xboole_0(k3_tarski(k2_tarski(X1,X2)),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_37, c_0_24])).
fof(c_0_47, plain, (![X3]:![X4]:(~r1_xboole_0(X3,X4)|k4_xboole_0(k2_xboole_0(X3,X4),X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t88_xboole_1])])).
fof(c_0_48, plain, (![X4]:![X5]:![X6]:(~r1_xboole_0(X4,k4_xboole_0(X5,X6))|r1_xboole_0(X5,k4_xboole_0(X4,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t81_xboole_1])])).
cnf(c_0_49,plain,(k3_tarski(k2_tarski(k4_xboole_0(X1,X2),k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_tarski(X1,X2)))))=X1), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X3))=k3_tarski(k2_tarski(k4_xboole_0(X1,X2),k5_xboole_0(k5_xboole_0(X1,X3),k3_tarski(k2_tarski(X1,X3)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_24]), c_0_29])).
cnf(c_0_51,plain,(k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_52, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(((~r2_hidden(X4,X6)|k4_xboole_0(k2_tarski(X4,X5),X6)!=k1_tarski(X4))&((r2_hidden(X5,X6)|X4=X5)|k4_xboole_0(k2_tarski(X4,X5),X6)!=k1_tarski(X4)))&(((~r2_hidden(X5,X6)|r2_hidden(X4,X6))|k4_xboole_0(k2_tarski(X4,X5),X6)=k1_tarski(X4))&((X4!=X5|r2_hidden(X4,X6))|k4_xboole_0(k2_tarski(X4,X5),X6)=k1_tarski(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l31_zfmisc_1])])])])])])).
cnf(c_0_53,plain,(k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k3_tarski(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[c_0_43, c_0_24])).
cnf(c_0_54,plain,(X1=k3_tarski(k2_tarski(X2,k4_xboole_0(X1,X2)))|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[c_0_44, c_0_24])).
cnf(c_0_55,plain,(r1_tarski(X1,k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_45, c_0_24])).
cnf(c_0_56,plain,(k4_xboole_0(k3_tarski(k2_tarski(X1,X2)),X1)=k4_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_39])).
cnf(c_0_57,plain,(k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58,plain,(r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X2,k4_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
fof(c_0_60, plain, (![X3]:![X4]:r1_xboole_0(k4_xboole_0(X3,X4),X4)), inference(variable_rename,[status(thm)],[t79_xboole_1])).
cnf(c_0_61,plain,(k4_xboole_0(k2_tarski(X1,X2),X3)!=k1_tarski(X1)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62,plain,(k4_xboole_0(X1,k3_tarski(k2_tarski(k3_tarski(k2_tarski(X2,X3)),X4)))=k4_xboole_0(X1,k3_tarski(k2_tarski(X2,k3_tarski(k2_tarski(X3,X4)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_53]), c_0_53]), c_0_53])).
cnf(c_0_63,plain,(k3_tarski(k2_tarski(X1,k4_xboole_0(X2,X1)))=k3_tarski(k2_tarski(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])).
fof(c_0_64, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(((r2_hidden(X4,X6)|k4_xboole_0(k2_tarski(X4,X5),X6)!=k1_xboole_0)&(r2_hidden(X5,X6)|k4_xboole_0(k2_tarski(X4,X5),X6)!=k1_xboole_0))&((~r2_hidden(X4,X6)|~r2_hidden(X5,X6))|k4_xboole_0(k2_tarski(X4,X5),X6)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_zfmisc_1])])])])])).
cnf(c_0_65,plain,(k4_xboole_0(k3_tarski(k2_tarski(X1,X2)),X2)=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_57, c_0_24])).
cnf(c_0_66,plain,(r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_59])).
cnf(c_0_67,plain,(r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_68, negated_conjecture, (~(![X1]:![X2]:![X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))))), inference(assume_negation,[status(cth)],[t75_zfmisc_1])).
cnf(c_0_69,plain,(k4_xboole_0(k2_tarski(X1,X2),k3_tarski(k2_tarski(X3,k3_tarski(k2_tarski(X4,X5)))))!=k1_tarski(X1)|~r2_hidden(X1,k3_tarski(k2_tarski(k3_tarski(k2_tarski(X3,X4)),X5)))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_70,plain,(k3_tarski(k2_tarski(X1,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_51]), c_0_42])).
cnf(c_0_71,plain,(r2_hidden(X2,X3)|k4_xboole_0(k2_tarski(X1,X2),X3)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_72, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_zfmisc_1])])])])).
cnf(c_0_73,plain,(k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_65, c_0_46])).
cnf(c_0_74,plain,(r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
fof(c_0_75, negated_conjecture, ((((esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk2_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&(esk1_0!=k2_tarski(esk2_0,esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0))))&(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0|(esk1_0=k1_xboole_0|(esk1_0=k1_tarski(esk2_0)|(esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_68])])])])])).
cnf(c_0_76,plain,(k4_xboole_0(k2_tarski(X1,X2),k3_tarski(k2_tarski(X3,X4)))!=k1_tarski(X1)|~r2_hidden(X1,k3_tarski(k2_tarski(X4,k3_tarski(k2_tarski(X3,X4)))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_39])).
cnf(c_0_77,plain,(r2_hidden(X1,k3_tarski(k2_tarski(k2_tarski(X2,X1),X3)))), inference(spm,[status(thm)],[c_0_71, c_0_41])).
fof(c_0_78, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_79,plain,(k3_tarski(k2_tarski(k1_xboole_0,X1))=X1), inference(spm,[status(thm)],[c_0_42, c_0_39])).
cnf(c_0_80,plain,(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)|r2_hidden(X1,X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_81,plain,(k1_tarski(X1)!=k1_xboole_0|~r2_hidden(X1,k3_tarski(k2_tarski(k2_tarski(X1,X2),X3)))), inference(spm,[status(thm)],[c_0_61, c_0_41])).
cnf(c_0_82,plain,(r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_83,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_84,negated_conjecture,(esk1_0=k2_tarski(esk2_0,esk3_0)|esk1_0=k1_tarski(esk3_0)|esk1_0=k1_tarski(esk2_0)|esk1_0=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_85,plain,(k4_xboole_0(k2_tarski(X1,X2),k3_tarski(k2_tarski(X3,k2_tarski(X4,X1))))!=k1_tarski(X1)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_86,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_87,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_55, c_0_79])).
cnf(c_0_88,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_89,plain,(k4_xboole_0(k2_tarski(X1,X1),X2)=k1_tarski(X1)|r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_80])).
cnf(c_0_90,plain,(k1_tarski(X1)!=k1_xboole_0|~r2_hidden(X1,k3_tarski(k2_tarski(X2,k2_tarski(X1,X3))))), inference(spm,[status(thm)],[c_0_81, c_0_39])).
cnf(c_0_91,plain,(r2_hidden(X1,k3_tarski(k2_tarski(k1_tarski(X1),X2)))), inference(spm,[status(thm)],[c_0_82, c_0_55])).
fof(c_0_92, plain, (![X4]:![X5]:![X6]:(k4_xboole_0(k2_tarski(X4,X5),X6)=k1_xboole_0|(k4_xboole_0(k2_tarski(X4,X5),X6)=k1_tarski(X4)|(k4_xboole_0(k2_tarski(X4,X5),X6)=k1_tarski(X5)|k4_xboole_0(k2_tarski(X4,X5),X6)=k2_tarski(X4,X5))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
cnf(c_0_93,plain,(k4_xboole_0(k3_tarski(k2_tarski(X1,X2)),k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_63]), c_0_83])).
cnf(c_0_94,negated_conjecture,(k3_tarski(k2_tarski(esk1_0,k2_tarski(esk2_0,esk3_0)))=k2_tarski(esk2_0,esk3_0)|k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|esk1_0=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_84]), c_0_42]), c_0_39])).
cnf(c_0_95,plain,(k4_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X1))!=k1_tarski(X1)), inference(spm,[status(thm)],[c_0_85, c_0_79])).
cnf(c_0_96,plain,(X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_97,plain,(k4_xboole_0(k2_tarski(X1,X1),X2)=k1_tarski(X1)|r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_98,plain,(k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_99,plain,(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)|k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X2)|k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)|k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_100,negated_conjecture,(k4_xboole_0(k2_tarski(esk2_0,esk3_0),k4_xboole_0(k2_tarski(esk2_0,esk3_0),esk1_0))=esk1_0|k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_101,plain,(k4_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X2))!=k1_tarski(X2)), inference(spm,[status(thm)],[c_0_95, c_0_39])).
cnf(c_0_102,plain,(k2_tarski(X1,X1)=k1_tarski(X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_59]), c_0_98])).
cnf(c_0_103,plain,(k4_xboole_0(k2_tarski(X1,X2),k2_tarski(X1,X3))!=k1_tarski(X1)), inference(spm,[status(thm)],[c_0_95, c_0_39])).
cnf(c_0_104,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k2_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_105,negated_conjecture,(k2_tarski(esk2_0,esk3_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|k1_tarski(esk2_0)=esk1_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_106,plain,(k4_xboole_0(k1_tarski(X1),k2_tarski(X2,X1))!=k1_tarski(X1)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_107,plain,(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)|k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_99, c_0_102])).
cnf(c_0_108,plain,(k4_xboole_0(k2_tarski(X1,X2),k2_tarski(X2,X3))!=k1_tarski(X2)), inference(spm,[status(thm)],[c_0_103, c_0_39])).
cnf(c_0_109,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_tarski(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_110,negated_conjecture,(k1_tarski(esk2_0)=esk1_0|k1_tarski(esk3_0)=esk1_0|esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_51])])).
cnf(c_0_111,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_112,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0|esk1_0!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_113,plain,(k4_xboole_0(k1_tarski(X1),k2_tarski(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_114,plain,(k4_xboole_0(k1_tarski(X1),k2_tarski(X1,X2))!=k1_tarski(X1)), inference(spm,[status(thm)],[c_0_108, c_0_102])).
cnf(c_0_115,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_111]), c_0_112])).
cnf(c_0_116,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(X1,esk3_0))=k1_xboole_0|k1_tarski(esk2_0)=esk1_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_113, c_0_110])).
cnf(c_0_117,plain,(k4_xboole_0(k1_tarski(X1),k2_tarski(X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_114, c_0_107])).
cnf(c_0_118,negated_conjecture,(k1_tarski(esk2_0)=esk1_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_119,negated_conjecture,(k4_xboole_0(esk1_0,k2_tarski(esk2_0,X1))=k1_xboole_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_120,negated_conjecture,(esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_119])).
cnf(c_0_121,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_51]), c_0_41])).
cnf(c_0_122,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_120]), c_0_121])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 123
# Proof object clause steps            : 82
# Proof object formula steps           : 41
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 30
# Processed clauses                    : 1788
# ...of these trivial                  : 80
# ...subsumed                          : 1311
# ...remaining for further processing  : 397
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 95
# Backward-rewritten                   : 79
# Generated clauses                    : 15940
# ...of the previous two non-trivial   : 13950
# Contextual simplify-reflections      : 323
# Paramodulations                      : 15907
# Factorizations                       : 26
# Equation resolutions                 : 7
# Current number of processed clauses  : 220
#    Positive orientable unit clauses  : 69
#    Positive unorientable unit clauses: 10
#    Negative unit clauses             : 37
#    Non-unit-clauses                  : 104
# Current number of unprocessed clauses: 5110
# ...number of literals in the above   : 11198
# Current number of archived formulas  : 0
# Current number of archived clauses   : 176
# Clause-clause subsumption calls (NU) : 14699
# Rec. Clause-clause subsumption calls : 4715
# Non-unit clause-clause subsumptions  : 1049
# Unit Clause-clause subsumption calls : 1526
# Rewrite failures with RHS unbound    : 15
# BW rewrite match attempts            : 1007
# BW rewrite match successes           : 147
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 307074

# -------------------------------------------------
# User time                : 0.232 s
# System time              : 0.008 s
# Total time               : 0.240 s
# Maximum resident set size: 43024 pages
