# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t76_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>(v25_waybel_0(X1)<=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r2_yellow_0(X1,X2))))), file('out/t76_waybel_0__2_premises__rerun.E_input', t76_waybel_0)).
fof(d40_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>(v25_waybel_0(X1)<=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))))), file('out/t76_waybel_0__2_premises__rerun.E_input', d40_waybel_0)).
fof(t16_yellow_0, axiom, (![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('out/t76_waybel_0__2_premises__rerun.E_input', t16_yellow_0)).
fof(c_0_3, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>(v25_waybel_0(X1)<=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r2_yellow_0(X1,X2)))))), inference(assume_negation,[status(cth)],[t76_waybel_0])).
fof(c_0_4, negated_conjecture, (![X5]:((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&((((~v1_xboole_0(esk2_0)|~v25_waybel_0(esk1_0))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v25_waybel_0(esk1_0)))&(~r2_yellow_0(esk1_0,esk2_0)|~v25_waybel_0(esk1_0)))&(v25_waybel_0(esk1_0)|((v1_xboole_0(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_yellow_0(esk1_0,X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])])])).
fof(c_0_5, plain, (![X5]:![X6]:![X8]:![X10]:(((((m1_subset_1(esk3_2(X5,X6),u1_struct_0(X5))|(v1_xboole_0(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))))|~v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5))))&((((r1_lattice3(X5,X6,esk3_2(X5,X6))|(v1_xboole_0(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))))|~v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5))))&((((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X6,X8)|r1_orders_2(X5,X8,esk3_2(X5,X6))))|(v1_xboole_0(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))))|~v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5))))))&((((~v1_xboole_0(esk4_1(X5))|v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5))))&((m1_subset_1(esk4_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5)))))&(((((m1_subset_1(esk5_2(X5,X10),u1_struct_0(X5))|~r1_lattice3(X5,esk4_1(X5),X10))|~m1_subset_1(X10,u1_struct_0(X5)))|v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5))))&(((((r1_lattice3(X5,esk4_1(X5),esk5_2(X5,X10))|~r1_lattice3(X5,esk4_1(X5),X10))|~m1_subset_1(X10,u1_struct_0(X5)))|v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5))))&((((~r1_orders_2(X5,esk5_2(X5,X10),X10)|~r1_lattice3(X5,esk4_1(X5),X10))|~m1_subset_1(X10,u1_struct_0(X5)))|v25_waybel_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|~l1_orders_2(X5))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d40_waybel_0])])])])])])])])).
cnf(c_0_6,negated_conjecture,(r2_yellow_0(esk1_0,X1)|v1_xboole_0(X1)|v25_waybel_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v25_waybel_0(X1)|m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_11, plain, (![X5]:![X6]:![X8]:![X6]:![X9]:((((m1_subset_1(esk6_2(X5,X6),u1_struct_0(X5))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((r1_lattice3(X5,X6,esk6_2(X5,X6))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((~m1_subset_1(X8,u1_struct_0(X5))|(~r1_lattice3(X5,X6,X8)|r1_orders_2(X5,X8,esk6_2(X5,X6))))|~r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))))&(((((m1_subset_1(esk7_3(X5,X6,X9),u1_struct_0(X5))|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&(((((r1_lattice3(X5,X6,esk7_3(X5,X6,X9))|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5)))&((((~r1_orders_2(X5,esk7_3(X5,X6,X9),X9)|~r1_lattice3(X5,X6,X9))|~m1_subset_1(X9,u1_struct_0(X5)))|r2_yellow_0(X5,X6))|(~v5_orders_2(X5)|~l1_orders_2(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])])).
cnf(c_0_12,plain,(v2_struct_0(X1)|v25_waybel_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(esk4_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13,negated_conjecture,(r2_yellow_0(esk1_0,esk4_1(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v25_waybel_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_14,plain,(m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,negated_conjecture,(r2_yellow_0(esk1_0,esk4_1(esk1_0))|v25_waybel_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_16,negated_conjecture,(v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,plain,(r1_lattice3(X1,X2,esk6_2(X1,X2))|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,plain,(v2_struct_0(X1)|v25_waybel_0(X1)|m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,esk4_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19,negated_conjecture,(m1_subset_1(esk6_2(esk1_0,esk4_1(esk1_0)),u1_struct_0(esk1_0))|v25_waybel_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_8]), c_0_16])])).
cnf(c_0_20,negated_conjecture,(r1_lattice3(esk1_0,esk4_1(esk1_0),esk6_2(esk1_0,esk4_1(esk1_0)))|v25_waybel_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_15]), c_0_8]), c_0_16])])).
cnf(c_0_21,plain,(r1_orders_2(X1,X3,esk6_2(X1,X2))|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22,negated_conjecture,(m1_subset_1(esk5_2(esk1_0,esk6_2(esk1_0,esk4_1(esk1_0))),u1_struct_0(esk1_0))|v25_waybel_0(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_8]), c_0_9])]), c_0_10]), c_0_20])).
cnf(c_0_23,plain,(v2_struct_0(X1)|v25_waybel_0(X1)|r1_lattice3(X1,esk4_1(X1),esk5_2(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,esk4_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_24,plain,(v2_struct_0(X1)|v25_waybel_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,esk4_1(X1),X2)|~r1_orders_2(X1,esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25,negated_conjecture,(r1_orders_2(esk1_0,esk5_2(esk1_0,esk6_2(esk1_0,esk4_1(esk1_0))),esk6_2(esk1_0,X1))|v25_waybel_0(esk1_0)|~r1_lattice3(esk1_0,X1,esk5_2(esk1_0,esk6_2(esk1_0,esk4_1(esk1_0))))|~r2_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_8]), c_0_16])])).
cnf(c_0_26,negated_conjecture,(r1_lattice3(esk1_0,esk4_1(esk1_0),esk5_2(esk1_0,esk6_2(esk1_0,esk4_1(esk1_0))))|v25_waybel_0(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_19]), c_0_8]), c_0_9])]), c_0_10]), c_0_20])).
cnf(c_0_27,negated_conjecture,(v25_waybel_0(esk1_0)|~r1_orders_2(esk1_0,esk5_2(esk1_0,esk6_2(esk1_0,esk4_1(esk1_0))),esk6_2(esk1_0,esk4_1(esk1_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_19]), c_0_8]), c_0_9])]), c_0_10]), c_0_20])).
cnf(c_0_28,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v25_waybel_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_29,negated_conjecture,(v25_waybel_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_15]), c_0_26]), c_0_27])).
cnf(c_0_30,negated_conjecture,(~v25_waybel_0(esk1_0)|~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_31,plain,(v2_struct_0(X1)|v1_xboole_0(X2)|m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v25_waybel_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_32,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29])])).
cnf(c_0_33,negated_conjecture,(~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_29])])).
cnf(c_0_34,plain,(r2_yellow_0(X1,X2)|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35,negated_conjecture,(m1_subset_1(esk3_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_29]), c_0_8]), c_0_9])]), c_0_33]), c_0_10])).
cnf(c_0_36,plain,(v2_struct_0(X1)|v1_xboole_0(X2)|r1_lattice3(X1,X2,esk3_2(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v25_waybel_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_37,negated_conjecture,(~v25_waybel_0(esk1_0)|~r2_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_38,negated_conjecture,(r2_yellow_0(esk1_0,X1)|m1_subset_1(esk7_3(esk1_0,X1,esk3_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,X1,esk3_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_8]), c_0_16])])).
cnf(c_0_39,negated_conjecture,(r1_lattice3(esk1_0,esk2_0,esk3_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_32]), c_0_29]), c_0_8]), c_0_9])]), c_0_33]), c_0_10])).
cnf(c_0_40,negated_conjecture,(~r2_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_29])])).
cnf(c_0_41,plain,(r2_yellow_0(X1,X2)|r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42,plain,(r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_43,plain,(v2_struct_0(X1)|v1_xboole_0(X2)|r1_orders_2(X1,X3,esk3_2(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v25_waybel_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk7_3(esk1_0,esk2_0,esk3_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_45,negated_conjecture,(r1_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,esk3_2(esk1_0,esk2_0)))|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,X1,esk3_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_35]), c_0_8]), c_0_16])])).
cnf(c_0_46,negated_conjecture,(r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,esk7_3(esk1_0,X1,esk3_2(esk1_0,esk2_0)),esk3_2(esk1_0,esk2_0))|~r1_lattice3(esk1_0,X1,esk3_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_35]), c_0_8]), c_0_16])])).
cnf(c_0_47,negated_conjecture,(r1_orders_2(esk1_0,esk7_3(esk1_0,esk2_0,esk3_2(esk1_0,esk2_0)),esk3_2(esk1_0,X1))|v1_xboole_0(X1)|~r1_lattice3(esk1_0,X1,esk7_3(esk1_0,esk2_0,esk3_2(esk1_0,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_29]), c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_48,negated_conjecture,(r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk3_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_39]), c_0_40])).
cnf(c_0_49,negated_conjecture,(~r1_orders_2(esk1_0,esk7_3(esk1_0,esk2_0,esk3_2(esk1_0,esk2_0)),esk3_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_39]), c_0_40])).
cnf(c_0_50,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_32]), c_0_48])]), c_0_49]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 44
# Proof object formula steps           : 7
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 73
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 87
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 83
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 21
# Generated clauses                    : 64
# ...of the previous two non-trivial   : 53
# Contextual simplify-reflections      : 12
# Paramodulations                      : 64
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 6
# Current number of archived formulas  : 0
# Current number of archived clauses   : 46
# Clause-clause subsumption calls (NU) : 292
# Rec. Clause-clause subsumption calls : 126
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 26
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4715

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 45316 pages
