# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C45_F1_PI_AE_Q7_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t74_relat_1, axiom, (![X1]:![X2]:![X3]:![X4]:(v1_relat_1(X4)=>(r2_hidden(k4_tarski(X1,X2),k5_relat_1(k6_relat_1(X3),X4))<=>(r2_hidden(X1,X3)&r2_hidden(k4_tarski(X1,X2),X4))))), file('out/t76_relat_1__5_premises__rerun.E_input', t74_relat_1)).
fof(d3_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3]:![X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2))))), file('out/t76_relat_1__5_premises__rerun.E_input', d3_relat_1)).
fof(t76_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(k5_relat_1(X2,k6_relat_1(X1)),X2)&r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2)))), file('out/t76_relat_1__5_premises__rerun.E_input', t76_relat_1)).
fof(t75_relat_1, axiom, (![X1]:![X2]:![X3]:![X4]:(v1_relat_1(X4)=>(r2_hidden(k4_tarski(X1,X2),k5_relat_1(X4,k6_relat_1(X3)))<=>(r2_hidden(X2,X3)&r2_hidden(k4_tarski(X1,X2),X4))))), file('out/t76_relat_1__5_premises__rerun.E_input', t75_relat_1)).
fof(dt_k5_relat_1, axiom, (![X1]:![X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2)))), file('out/t76_relat_1__5_premises__rerun.E_input', dt_k5_relat_1)).
fof(dt_k6_relat_1, axiom, (![X1]:v1_relat_1(k6_relat_1(X1))), file('out/t76_relat_1__5_premises__rerun.E_input', dt_k6_relat_1)).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:((((r2_hidden(X5,X7)|~r2_hidden(k4_tarski(X5,X6),k5_relat_1(k6_relat_1(X7),X8)))|~v1_relat_1(X8))&((r2_hidden(k4_tarski(X5,X6),X8)|~r2_hidden(k4_tarski(X5,X6),k5_relat_1(k6_relat_1(X7),X8)))|~v1_relat_1(X8)))&(((~r2_hidden(X5,X7)|~r2_hidden(k4_tarski(X5,X6),X8))|r2_hidden(k4_tarski(X5,X6),k5_relat_1(k6_relat_1(X7),X8)))|~v1_relat_1(X8)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t74_relat_1])])])).
fof(c_0_7, plain, (![X5]:![X6]:![X7]:![X8]:![X6]:(((~r1_tarski(X5,X6)|(~r2_hidden(k4_tarski(X7,X8),X5)|r2_hidden(k4_tarski(X7,X8),X6)))|~v1_relat_1(X5))&(((r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X5)|r1_tarski(X5,X6))|~v1_relat_1(X5))&((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|r1_tarski(X5,X6))|~v1_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])])).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(k5_relat_1(X2,k6_relat_1(X1)),X2)&r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2))))), inference(assume_negation,[status(cth)],[t76_relat_1])).
cnf(c_0_9,plain,(r2_hidden(k4_tarski(X2,X3),X1)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X3),k5_relat_1(k6_relat_1(X4),X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(r1_tarski(X1,X2)|r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X8]:((((r2_hidden(X6,X7)|~r2_hidden(k4_tarski(X5,X6),k5_relat_1(X8,k6_relat_1(X7))))|~v1_relat_1(X8))&((r2_hidden(k4_tarski(X5,X6),X8)|~r2_hidden(k4_tarski(X5,X6),k5_relat_1(X8,k6_relat_1(X7))))|~v1_relat_1(X8)))&(((~r2_hidden(X6,X7)|~r2_hidden(k4_tarski(X5,X6),X8))|r2_hidden(k4_tarski(X5,X6),k5_relat_1(X8,k6_relat_1(X7))))|~v1_relat_1(X8)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t75_relat_1])])])).
fof(c_0_12, negated_conjecture, ((v1_relat_1(esk4_0)&(~r1_tarski(k5_relat_1(esk4_0,k6_relat_1(esk3_0)),esk4_0)|~r1_tarski(k5_relat_1(k6_relat_1(esk3_0),esk4_0),esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_13,plain,(r1_tarski(X1,X2)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14,plain,(r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X3)|r2_hidden(k4_tarski(esk1_2(k5_relat_1(k6_relat_1(X1),X2),X3),esk2_2(k5_relat_1(k6_relat_1(X1),X2),X3)),X2)|~v1_relat_1(k5_relat_1(k6_relat_1(X1),X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_15,plain,(r2_hidden(k4_tarski(X2,X3),X1)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X3),k5_relat_1(X1,k6_relat_1(X4)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(~r1_tarski(k5_relat_1(k6_relat_1(esk3_0),esk4_0),esk4_0)|~r1_tarski(k5_relat_1(esk4_0,k6_relat_1(esk3_0)),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2)|~v1_relat_1(k5_relat_1(k6_relat_1(X1),X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18,negated_conjecture,(v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(r1_tarski(k5_relat_1(X1,k6_relat_1(X2)),X3)|r2_hidden(k4_tarski(esk1_2(k5_relat_1(X1,k6_relat_1(X2)),X3),esk2_2(k5_relat_1(X1,k6_relat_1(X2)),X3)),X1)|~v1_relat_1(k5_relat_1(X1,k6_relat_1(X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_15, c_0_10])).
cnf(c_0_20,negated_conjecture,(~r1_tarski(k5_relat_1(esk4_0,k6_relat_1(esk3_0)),esk4_0)|~v1_relat_1(k5_relat_1(k6_relat_1(esk3_0),esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
cnf(c_0_21,plain,(r1_tarski(k5_relat_1(X1,k6_relat_1(X2)),X1)|~v1_relat_1(k5_relat_1(X1,k6_relat_1(X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_13, c_0_19])).
fof(c_0_22, plain, (![X3]:![X4]:((~v1_relat_1(X3)|~v1_relat_1(X4))|v1_relat_1(k5_relat_1(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_23, plain, (![X2]:v1_relat_1(k6_relat_1(X2))), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_24,negated_conjecture,(~v1_relat_1(k5_relat_1(k6_relat_1(esk3_0),esk4_0))|~v1_relat_1(k5_relat_1(esk4_0,k6_relat_1(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_18])])).
cnf(c_0_25,plain,(v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,plain,(v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27,negated_conjecture,(~v1_relat_1(k5_relat_1(esk4_0,k6_relat_1(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_18]), c_0_26])])).
cnf(c_0_28,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_25]), c_0_26]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 16
# Proof object formula steps           : 13
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 48
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 63
# ...of the previous two non-trivial   : 57
# Contextual simplify-reflections      : 0
# Paramodulations                      : 63
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 35
# ...number of literals in the above   : 166
# Current number of archived formulas  : 0
# Current number of archived clauses   : 13
# Clause-clause subsumption calls (NU) : 136
# Rec. Clause-clause subsumption calls : 69
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2765

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.008 s
# Total time               : 0.016 s
# Maximum resident set size: 45300 pages
