# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_xboole_1, axiom, (![X1]:![X2]:![X3]:k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), file('out/t78_relat_1__3_premises__rerun.E_input', t23_xboole_1)).
fof(t67_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k5_relat_1(X2,X1)=k3_xboole_0(X2,k2_zfmisc_1(X1,k10_xtuple_0(X2))))), file('out/t78_relat_1__3_premises__rerun.E_input', t67_relat_1)).
fof(t97_zfmisc_1, axiom, (![X1]:![X2]:![X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2)))), file('out/t78_relat_1__3_premises__rerun.E_input', t97_zfmisc_1)).
fof(t78_relat_1, conjecture, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k5_relat_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k5_relat_1(X3,X1),k5_relat_1(X3,X2)))), file('out/t78_relat_1__3_premises__rerun.E_input', t78_relat_1)).
fof(c_0_4, plain, (![X4]:![X5]:![X6]:k3_xboole_0(X4,k2_xboole_0(X5,X6))=k2_xboole_0(k3_xboole_0(X4,X5),k3_xboole_0(X4,X6))), inference(variable_rename,[status(thm)],[t23_xboole_1])).
fof(c_0_5, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k5_relat_1(X4,X3)=k3_xboole_0(X4,k2_zfmisc_1(X3,k10_xtuple_0(X4))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_relat_1])])).
cnf(c_0_6,plain,(k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(k5_relat_1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,k10_xtuple_0(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:![X6]:(k2_zfmisc_1(k2_xboole_0(X4,X5),X6)=k2_xboole_0(k2_zfmisc_1(X4,X6),k2_zfmisc_1(X5,X6))&k2_zfmisc_1(X6,k2_xboole_0(X4,X5))=k2_xboole_0(k2_zfmisc_1(X6,X4),k2_zfmisc_1(X6,X5)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[t97_zfmisc_1])])])).
cnf(c_0_9,plain,(k3_xboole_0(X1,k2_xboole_0(X2,k2_zfmisc_1(X3,k10_xtuple_0(X1))))=k2_xboole_0(k3_xboole_0(X1,X2),k5_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10,plain,(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:![X3]:(v1_relat_1(X3)=>k5_relat_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k5_relat_1(X3,X1),k5_relat_1(X3,X2))))), inference(assume_negation,[status(cth)],[t78_relat_1])).
cnf(c_0_12,plain,(k2_xboole_0(k3_xboole_0(X1,k2_zfmisc_1(X2,k10_xtuple_0(X1))),k5_relat_1(X1,X3))=k3_xboole_0(X1,k2_zfmisc_1(k2_xboole_0(X2,X3),k10_xtuple_0(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
fof(c_0_13, negated_conjecture, ((v1_relat_1(esk3_0)&k5_relat_1(esk3_0,k2_xboole_0(esk1_0,esk2_0))!=k2_xboole_0(k5_relat_1(esk3_0,esk1_0),k5_relat_1(esk3_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_14,plain,(k3_xboole_0(X1,k2_zfmisc_1(k2_xboole_0(X2,X3),k10_xtuple_0(X1)))=k2_xboole_0(k5_relat_1(X1,X2),k5_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_12, c_0_7])).
cnf(c_0_15,negated_conjecture,(k5_relat_1(esk3_0,k2_xboole_0(esk1_0,esk2_0))!=k2_xboole_0(k5_relat_1(esk3_0,esk1_0),k5_relat_1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16,plain,(k2_xboole_0(k5_relat_1(X1,X2),k5_relat_1(X1,X3))=k5_relat_1(X1,k2_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_7, c_0_14])).
cnf(c_0_17,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 19
# Proof object clause steps            : 10
# Proof object formula steps           : 9
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 6
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 6
# Processed clauses                    : 66
# ...of these trivial                  : 3
# ...subsumed                          : 26
# ...remaining for further processing  : 37
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 1101
# ...of the previous two non-trivial   : 1045
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1101
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 31
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 15
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 7
# Current number of unprocessed clauses: 991
# ...number of literals in the above   : 1144
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6
# Clause-clause subsumption calls (NU) : 13
# Rec. Clause-clause subsumption calls : 13
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 34
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 31
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15116

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.004 s
# Total time               : 0.024 s
# Maximum resident set size: 42916 pages
