# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t79_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((m1_pre_topc(X4,X3)&m1_pre_topc(X5,X4))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>r2_funct_2(u1_struct_0(X5),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,X5,k3_tmap_1(X1,X2,X3,X4,X6)),k3_tmap_1(X1,X2,X3,X5,X6)))))))))), file('out/t79_tmap_1__9_premises__rerun.E_input', t79_tmap_1)).
fof(dt_m1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2)))), file('out/t79_tmap_1__9_premises__rerun.E_input', dt_m1_pre_topc)).
fof(cc1_pre_topc, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2)))), file('out/t79_tmap_1__9_premises__rerun.E_input', cc1_pre_topc)).
fof(d5_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:(m1_pre_topc(X3,X1)=>![X4]:(m1_pre_topc(X4,X1)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>(m1_pre_topc(X4,X3)=>k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))))))))), file('out/t79_tmap_1__9_premises__rerun.E_input', d5_tmap_1)).
fof(t7_tsep_1, axiom, (![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X2)=>m1_pre_topc(X3,X1))))), file('out/t79_tmap_1__9_premises__rerun.E_input', t7_tsep_1)).
fof(d4_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_pre_topc(X4,X1)=>k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4))))))), file('out/t79_tmap_1__9_premises__rerun.E_input', d4_tmap_1)).
fof(dt_k3_tmap_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&(m1_pre_topc(X3,X1)&(m1_pre_topc(X4,X1)&(v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))))))=>(v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))), file('out/t79_tmap_1__9_premises__rerun.E_input', dt_k3_tmap_1)).
fof(t2_tsep_1, axiom, (![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1))), file('out/t79_tmap_1__9_premises__rerun.E_input', t2_tsep_1)).
fof(t77_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>((r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),X6,k2_tmap_1(X1,X2,X5,X3))&m1_pre_topc(X4,X3))=>r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,X4,X6),k2_tmap_1(X1,X2,X5,X4)))))))))), file('out/t79_tmap_1__9_premises__rerun.E_input', t77_tmap_1)).
fof(t74_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),X4,k3_tmap_1(X1,X2,X3,X3,X4))))))), file('out/t79_tmap_1__9_premises__rerun.E_input', t74_tmap_1)).
fof(c_0_10, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((m1_pre_topc(X4,X3)&m1_pre_topc(X5,X4))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>r2_funct_2(u1_struct_0(X5),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,X5,k3_tmap_1(X1,X2,X3,X4,X6)),k3_tmap_1(X1,X2,X3,X5,X6))))))))))), inference(assume_negation,[status(cth)],[t79_tmap_1])).
fof(c_0_11, plain, (![X3]:![X4]:(~l1_pre_topc(X3)|(~m1_pre_topc(X4,X3)|l1_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])])])).
fof(c_0_12, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&((~v2_struct_0(esk5_0)&m1_pre_topc(esk5_0,esk1_0))&((m1_pre_topc(esk4_0,esk3_0)&m1_pre_topc(esk5_0,esk4_0))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))))&~r2_funct_2(u1_struct_0(esk5_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)),k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,esk6_0)))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])).
fof(c_0_13, plain, (![X3]:![X4]:((~v2_pre_topc(X3)|~l1_pre_topc(X3))|(~m1_pre_topc(X4,X3)|v2_pre_topc(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])])])).
fof(c_0_14, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))|((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|(~m1_pre_topc(X8,X6)|(~m1_pre_topc(X9,X6)|((~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X8),u1_struct_0(X7))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X7))))))|(~m1_pre_topc(X9,X8)|k3_tmap_1(X6,X7,X8,X9,X10)=k2_partfun1(u1_struct_0(X8),u1_struct_0(X7),X10,u1_struct_0(X9))))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tmap_1])])])])])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:((~v2_pre_topc(X4)|~l1_pre_topc(X4))|(~m1_pre_topc(X5,X4)|(~m1_pre_topc(X6,X5)|m1_pre_topc(X6,X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tsep_1])])])])])).
fof(c_0_16, plain, (![X5]:![X6]:![X7]:![X8]:((v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))|((~v1_funct_1(X7)|(~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6))))))|(~m1_pre_topc(X8,X5)|k2_tmap_1(X5,X6,X7,X8)=k2_partfun1(u1_struct_0(X5),u1_struct_0(X6),X7,u1_struct_0(X8))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tmap_1])])])])])])).
cnf(c_0_17,plain,(l1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(v2_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_22, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:((v1_funct_1(k3_tmap_1(X6,X7,X8,X9,X10))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))|((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|(~m1_pre_topc(X8,X6)|(~m1_pre_topc(X9,X6)|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X8),u1_struct_0(X7))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X7)))))))))))&((v1_funct_2(k3_tmap_1(X6,X7,X8,X9,X10),u1_struct_0(X9),u1_struct_0(X7))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))|((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|(~m1_pre_topc(X8,X6)|(~m1_pre_topc(X9,X6)|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X8),u1_struct_0(X7))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X7)))))))))))&(m1_subset_1(k3_tmap_1(X6,X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X7))))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))|((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|(~m1_pre_topc(X8,X6)|(~m1_pre_topc(X9,X6)|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X8),u1_struct_0(X7))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X7)))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_tmap_1])])])])).
cnf(c_0_23,plain,(k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X4,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,plain,(m1_pre_topc(X1,X2)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26,negated_conjecture,(v1_funct_2(esk6_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27,negated_conjecture,(m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_28,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_29,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30,negated_conjecture,(l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_31,negated_conjecture,(v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32,negated_conjecture,(v2_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_18]), c_0_19]), c_0_21])])).
cnf(c_0_33,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_35, plain, (![X2]:(~l1_pre_topc(X2)|m1_pre_topc(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_36,plain,(v2_struct_0(X3)|v2_struct_0(X5)|v1_funct_1(k3_tmap_1(X5,X3,X2,X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X2,X5)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X5)|~v2_pre_topc(X5)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_37,plain,(k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_38,negated_conjecture,(k2_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk2_0),esk6_0,u1_struct_0(X1))=k2_tmap_1(esk3_0,esk2_0,esk6_0,X1)|~m1_pre_topc(X1,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34])).
cnf(c_0_39,negated_conjecture,(m1_pre_topc(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_40,plain,(m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41,plain,(v2_struct_0(X3)|v2_struct_0(X5)|v1_funct_2(k3_tmap_1(X5,X3,X2,X4,X1),u1_struct_0(X4),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X2,X5)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X5)|~v2_pre_topc(X5)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_42,negated_conjecture,(v1_funct_1(k3_tmap_1(X1,esk2_0,esk3_0,X2,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_31])]), c_0_33])).
cnf(c_0_43,plain,(v2_struct_0(X3)|v2_struct_0(X5)|m1_subset_1(k3_tmap_1(X5,X3,X2,X4,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X2,X5)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X5)|~v2_pre_topc(X5)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_44,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk3_0,X2,esk6_0)=k2_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk2_0),esk6_0,u1_struct_0(X2))|v2_struct_0(X1)|~m1_pre_topc(X2,esk3_0)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_31])]), c_0_33])).
cnf(c_0_45,negated_conjecture,(k2_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk2_0),esk6_0,u1_struct_0(esk4_0))=k2_tmap_1(esk3_0,esk2_0,esk6_0,esk4_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46,negated_conjecture,(m1_pre_topc(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_40, c_0_30])).
cnf(c_0_47,negated_conjecture,(v1_funct_2(k3_tmap_1(X1,esk2_0,esk3_0,X2,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_31])]), c_0_33])).
cnf(c_0_48,negated_conjecture,(v1_funct_1(k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0))|~m1_pre_topc(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_39]), c_0_34]), c_0_30]), c_0_32])])).
cnf(c_0_49,negated_conjecture,(m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50,negated_conjecture,(m1_subset_1(k3_tmap_1(X1,esk2_0,esk3_0,X2,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_31])]), c_0_33])).
cnf(c_0_51,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk6_0)=k2_tmap_1(esk3_0,esk2_0,esk6_0,esk4_0)|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_39]), c_0_45])).
cnf(c_0_52,negated_conjecture,(k2_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk2_0),esk6_0,u1_struct_0(esk3_0))=k2_tmap_1(esk3_0,esk2_0,esk6_0,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_46])).
cnf(c_0_53,negated_conjecture,(v1_funct_2(k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_39]), c_0_30]), c_0_32])]), c_0_34]), c_0_46])])).
cnf(c_0_54,negated_conjecture,(v1_funct_1(k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_46])])).
cnf(c_0_55,negated_conjecture,(l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_49]), c_0_19])])).
cnf(c_0_56,negated_conjecture,(v2_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_49]), c_0_19]), c_0_21])])).
cnf(c_0_57,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_58,negated_conjecture,(m1_subset_1(k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_39]), c_0_46]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_59,negated_conjecture,(k2_tmap_1(esk3_0,esk2_0,esk6_0,esk4_0)=k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_46]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_60,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk6_0)=k2_tmap_1(esk3_0,esk2_0,esk6_0,esk3_0)|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_46]), c_0_52])).
cnf(c_0_61,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_62,negated_conjecture,(k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0),u1_struct_0(X1))=k2_tmap_1(esk4_0,esk2_0,k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0),X1)|~m1_pre_topc(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_53]), c_0_54]), c_0_29]), c_0_55]), c_0_31]), c_0_56])]), c_0_33]), c_0_57]), c_0_58])])).
cnf(c_0_63,negated_conjecture,(m1_pre_topc(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_64,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk6_0)=k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0)|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_51, c_0_59])).
cnf(c_0_65,negated_conjecture,(k2_tmap_1(esk3_0,esk2_0,esk6_0,esk3_0)=k3_tmap_1(esk3_0,esk2_0,esk3_0,esk3_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_46]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_66,negated_conjecture,(v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_49]), c_0_18]), c_0_19]), c_0_21])]), c_0_61])).
cnf(c_0_67,negated_conjecture,(v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_49]), c_0_18]), c_0_19]), c_0_21])]), c_0_61])).
cnf(c_0_68,negated_conjecture,(m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_49]), c_0_18]), c_0_19]), c_0_21])]), c_0_61])).
cnf(c_0_69,negated_conjecture,(k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0),u1_struct_0(esk5_0))=k2_tmap_1(esk4_0,esk2_0,k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0),esk5_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_70,negated_conjecture,(k3_tmap_1(esk3_0,esk2_0,esk3_0,esk4_0,esk6_0)=k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_18]), c_0_19]), c_0_21])]), c_0_61])).
fof(c_0_71, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|((v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))|((v2_struct_0(X9)|~m1_pre_topc(X9,X7))|((v2_struct_0(X10)|~m1_pre_topc(X10,X7))|((~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X7),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X8))))))|((~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))|((~r2_funct_2(u1_struct_0(X9),u1_struct_0(X8),X12,k2_tmap_1(X7,X8,X11,X9))|~m1_pre_topc(X10,X9))|r2_funct_2(u1_struct_0(X10),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,X10,X12),k2_tmap_1(X7,X8,X11,X10)))))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t77_tmap_1])])])])])])).
cnf(c_0_72,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk6_0)=k3_tmap_1(esk3_0,esk2_0,esk3_0,esk3_0,esk6_0)|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_60, c_0_65])).
fof(c_0_73, plain, (![X5]:![X6]:![X7]:![X8]:((v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))|((v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))|((v2_struct_0(X7)|~m1_pre_topc(X7,X5))|((~v1_funct_1(X8)|(~v1_funct_2(X8,u1_struct_0(X7),u1_struct_0(X6))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X7),u1_struct_0(X6))))))|r2_funct_2(u1_struct_0(X7),u1_struct_0(X6),X8,k3_tmap_1(X5,X6,X7,X7,X8))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t74_tmap_1])])])])])])).
cnf(c_0_74,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk4_0,X2,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),u1_struct_0(X2))|v2_struct_0(X1)|~m1_pre_topc(X2,esk4_0)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_66]), c_0_67]), c_0_29]), c_0_31])]), c_0_33]), c_0_68])])).
cnf(c_0_75,negated_conjecture,(k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),u1_struct_0(esk5_0))=k2_tmap_1(esk4_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_70]), c_0_70])).
cnf(c_0_76,negated_conjecture,(m1_pre_topc(X1,esk3_0)|~m1_pre_topc(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_39]), c_0_30]), c_0_32])])).
cnf(c_0_77,plain,(r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k3_tmap_1(X3,X2,X4,X1,X5),k2_tmap_1(X3,X2,X6,X1))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X1,X4)|~r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),X5,k2_tmap_1(X3,X2,X6,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78,negated_conjecture,(k3_tmap_1(esk3_0,esk2_0,esk3_0,esk3_0,esk6_0)=k3_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_18]), c_0_19]), c_0_21])]), c_0_61])).
cnf(c_0_79,plain,(r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,k3_tmap_1(X4,X2,X1,X1,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X4)|~v2_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_80,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=k2_tmap_1(esk4_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk5_0)|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_63]), c_0_75])).
cnf(c_0_81,negated_conjecture,(m1_pre_topc(esk4_0,esk4_0)), inference(spm,[status(thm)],[c_0_40, c_0_55])).
cnf(c_0_82,negated_conjecture,(m1_pre_topc(esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_63])).
cnf(c_0_83,plain,(r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k3_tmap_1(X3,X2,X4,X1,X5),k2_tmap_1(X3,X2,X6,X1))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),X5,k2_tmap_1(X3,X2,X6,X4))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_77, c_0_24])).
cnf(c_0_84,negated_conjecture,(k2_tmap_1(esk3_0,esk2_0,esk6_0,esk3_0)=k3_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[c_0_65, c_0_78])).
cnf(c_0_85,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),esk6_0,k3_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_31])]), c_0_33]), c_0_34])).
cnf(c_0_86,negated_conjecture,(k2_tmap_1(esk4_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk5_0)=k3_tmap_1(esk4_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_87,negated_conjecture,(k2_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk2_0),esk6_0,u1_struct_0(esk5_0))=k2_tmap_1(esk3_0,esk2_0,esk6_0,esk5_0)), inference(spm,[status(thm)],[c_0_38, c_0_82])).
cnf(c_0_88,negated_conjecture,(r2_funct_2(u1_struct_0(X1),u1_struct_0(esk2_0),k3_tmap_1(esk3_0,esk2_0,esk3_0,X1,X2),k2_tmap_1(esk3_0,esk2_0,esk6_0,X1))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),X2,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~m1_pre_topc(X1,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_27]), c_0_26]), c_0_28]), c_0_46]), c_0_30]), c_0_29]), c_0_32]), c_0_31])]), c_0_34]), c_0_33])).
cnf(c_0_89,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),esk6_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_18]), c_0_19]), c_0_21])]), c_0_61])).
cnf(c_0_90,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=k3_tmap_1(esk4_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_80, c_0_86])).
cnf(c_0_91,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk3_0,esk5_0,esk6_0)=k2_tmap_1(esk3_0,esk2_0,esk6_0,esk5_0)|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_82]), c_0_87])).
cnf(c_0_92,negated_conjecture,(k2_tmap_1(esk3_0,esk2_0,esk6_0,esk4_0)=k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(rw,[status(thm)],[c_0_59, c_0_70])).
cnf(c_0_93,negated_conjecture,(r2_funct_2(u1_struct_0(X1),u1_struct_0(esk2_0),k3_tmap_1(esk3_0,esk2_0,esk3_0,X1,esk6_0),k2_tmap_1(esk3_0,esk2_0,esk6_0,X1))|v2_struct_0(X1)|~m1_pre_topc(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_27]), c_0_26]), c_0_28])])).
cnf(c_0_94,negated_conjecture,(k3_tmap_1(esk4_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=k3_tmap_1(esk3_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_39]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_95,negated_conjecture,(k2_tmap_1(esk3_0,esk2_0,esk6_0,esk5_0)=k3_tmap_1(esk3_0,esk2_0,esk3_0,esk5_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_46]), c_0_30]), c_0_32])]), c_0_34])).
cnf(c_0_96,negated_conjecture,(r2_funct_2(u1_struct_0(X1),u1_struct_0(esk2_0),k3_tmap_1(esk3_0,esk2_0,esk4_0,X1,X2),k2_tmap_1(esk3_0,esk2_0,esk6_0,X1))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),X2,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~m1_pre_topc(X1,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_92]), c_0_27]), c_0_26]), c_0_28]), c_0_39]), c_0_30]), c_0_29]), c_0_32]), c_0_31])]), c_0_57]), c_0_34]), c_0_33])).
cnf(c_0_97,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_39]), c_0_70]), c_0_59]), c_0_70]), c_0_57])).
cnf(c_0_98,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=k3_tmap_1(esk3_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_90, c_0_94])).
cnf(c_0_99,negated_conjecture,(k3_tmap_1(X1,esk2_0,esk3_0,esk5_0,esk6_0)=k3_tmap_1(esk3_0,esk2_0,esk3_0,esk5_0,esk6_0)|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_91, c_0_95])).
cnf(c_0_100,negated_conjecture,(r2_funct_2(u1_struct_0(X1),u1_struct_0(esk2_0),k3_tmap_1(esk3_0,esk2_0,esk4_0,X1,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)),k2_tmap_1(esk3_0,esk2_0,esk6_0,X1))|v2_struct_0(X1)|~m1_pre_topc(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_68]), c_0_66]), c_0_67])])).
cnf(c_0_101,negated_conjecture,(k3_tmap_1(esk3_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_49]), c_0_19]), c_0_21])]), c_0_61])).
cnf(c_0_102,negated_conjecture,(k3_tmap_1(esk3_0,esk2_0,esk3_0,esk5_0,esk6_0)=k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,esk6_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_18]), c_0_19]), c_0_21])]), c_0_61])).
cnf(c_0_103,negated_conjecture,(~r2_funct_2(u1_struct_0(esk5_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)),k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_104,negated_conjecture,(~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_105,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_63]), c_0_101]), c_0_95]), c_0_102]), c_0_103]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 85
# Proof object formula steps           : 21
# Proof object conjectures             : 75
# Proof object clause conjectures      : 72
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 197
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 473
# ...of these trivial                  : 10
# ...subsumed                          : 7
# ...remaining for further processing  : 456
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 138
# Generated clauses                    : 664
# ...of the previous two non-trivial   : 698
# Contextual simplify-reflections      : 2
# Paramodulations                      : 664
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 289
#    Positive orientable unit clauses  : 156
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 127
# Current number of unprocessed clauses: 127
# ...number of literals in the above   : 793
# Current number of archived formulas  : 0
# Current number of archived clauses   : 167
# Clause-clause subsumption calls (NU) : 7169
# Rec. Clause-clause subsumption calls : 721
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 3523
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1310
# BW rewrite match successes           : 27
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 43704

# -------------------------------------------------
# User time                : 0.048 s
# System time              : 0.000 s
# Total time               : 0.048 s
# Maximum resident set size: 45320 pages
