# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t7_subset_1__5_premises__rerun.E_input', d5_xboole_0)).
fof(t7_subset_1, conjecture, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,X2)<=>(r2_hidden(X5,X3)&~(r2_hidden(X5,X4)))))=>X2=k7_subset_1(X1,X3,X4)))))), file('out/t7_subset_1__5_premises__rerun.E_input', t7_subset_1)).
fof(t2_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3))))), file('out/t7_subset_1__5_premises__rerun.E_input', t2_subset_1)).
fof(dt_k7_subset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k7_subset_1(X1,X2,X3),k1_zfmisc_1(X1)))), file('out/t7_subset_1__5_premises__rerun.E_input', dt_k7_subset_1)).
fof(redefinition_k7_subset_1, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3))), file('out/t7_subset_1__5_premises__rerun.E_input', redefinition_k7_subset_1)).
fof(d10_xboole_0, axiom, (![X1]:![X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('out/t7_subset_1__5_premises__rerun.E_input', d10_xboole_0)).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
cnf(c_0_7,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_8,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|~r2_hidden(esk1_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X1)|r2_hidden(esk1_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(r2_hidden(X4,X2)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,X2)<=>(r2_hidden(X5,X3)&~(r2_hidden(X5,X4)))))=>X2=k7_subset_1(X1,X3,X4))))))), inference(assume_negation,[status(cth)],[t7_subset_1])).
cnf(c_0_12,plain,(~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(X1=k4_xboole_0(X2,X2)|r2_hidden(esk1_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_10])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:((((m1_subset_1(esk6_3(X5,X6,X7),X5)|r1_tarski(X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(X5)))|~m1_subset_1(X6,k1_zfmisc_1(X5)))&((((r2_hidden(esk6_3(X5,X6,X7),X6)|r1_tarski(X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(X5)))|~m1_subset_1(X6,k1_zfmisc_1(X5)))&(((~r2_hidden(esk6_3(X5,X6,X7),X7)|r1_tarski(X6,X7))|~m1_subset_1(X7,k1_zfmisc_1(X5)))|~m1_subset_1(X6,k1_zfmisc_1(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset_1])])])])])])])).
fof(c_0_16, negated_conjecture, (![X10]:(m1_subset_1(esk3_0,k1_zfmisc_1(esk2_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(esk2_0))&(m1_subset_1(esk5_0,k1_zfmisc_1(esk2_0))&(((((r2_hidden(X10,esk4_0)|~r2_hidden(X10,esk3_0))|~m1_subset_1(X10,esk2_0))&((~r2_hidden(X10,esk5_0)|~r2_hidden(X10,esk3_0))|~m1_subset_1(X10,esk2_0)))&(((~r2_hidden(X10,esk4_0)|r2_hidden(X10,esk5_0))|r2_hidden(X10,esk3_0))|~m1_subset_1(X10,esk2_0)))&esk3_0!=k7_subset_1(esk2_0,esk4_0,esk5_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])])])).
cnf(c_0_17,plain,(k4_xboole_0(X1,X2)=k4_xboole_0(X3,X3)|~r2_hidden(esk1_3(X3,X3,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18,plain,(k4_xboole_0(X1,X2)=k4_xboole_0(X3,X3)|r2_hidden(esk1_3(X3,X3,k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_14, c_0_13])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X5,k1_zfmisc_1(X4))|m1_subset_1(k7_subset_1(X4,X5,X6),k1_zfmisc_1(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_subset_1])])])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X5,k1_zfmisc_1(X4))|k7_subset_1(X4,X5,X6)=k4_xboole_0(X5,X6))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])])])).
cnf(c_0_21,plain,(r1_tarski(X1,X3)|r2_hidden(esk6_3(X2,X1,X3),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,negated_conjecture,(r2_hidden(X1,esk4_0)|~m1_subset_1(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(r1_tarski(X1,X3)|m1_subset_1(esk6_3(X2,X1,X3),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24,plain,(k4_xboole_0(X1,X1)=k4_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25,plain,(m1_subset_1(k7_subset_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~m1_subset_1(k4_xboole_0(X1,X2),k1_zfmisc_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(X4))|~r2_hidden(esk6_3(X4,k4_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_12, c_0_21])).
cnf(c_0_28,negated_conjecture,(r1_tarski(X1,X2)|r2_hidden(esk6_3(esk2_0,X1,X2),esk4_0)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~r2_hidden(esk6_3(esk2_0,X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29,plain,(~r2_hidden(X1,k4_xboole_0(X2,X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_12, c_0_24])).
cnf(c_0_30,plain,(m1_subset_1(k4_xboole_0(X1,X2),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31,negated_conjecture,(r1_tarski(k4_xboole_0(X1,esk4_0),X2)|~m1_subset_1(k4_xboole_0(X1,esk4_0),k1_zfmisc_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~r2_hidden(esk6_3(esk2_0,k4_xboole_0(X1,esk4_0),X2),esk3_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|r2_hidden(esk6_3(X4,k4_xboole_0(X1,X2),X3),X1)|~m1_subset_1(k4_xboole_0(X1,X2),k1_zfmisc_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(X4))), inference(spm,[status(thm)],[c_0_14, c_0_21])).
cnf(c_0_33,plain,(~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_24]), c_0_29])).
cnf(c_0_34,plain,(m1_subset_1(k4_xboole_0(X1,X1),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_30, c_0_24])).
cnf(c_0_35,plain,(r2_hidden(X4,X1)|r2_hidden(X4,X3)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_36, plain, (![X3]:![X4]:![X3]:![X4]:(((r1_tarski(X3,X4)|X3!=X4)&(r1_tarski(X4,X3)|X3!=X4))&((~r1_tarski(X3,X4)|~r1_tarski(X4,X3))|X3=X4))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_37,negated_conjecture,(r1_tarski(k4_xboole_0(esk3_0,esk4_0),X1)|~m1_subset_1(k4_xboole_0(esk3_0,esk4_0),k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39,plain,(r1_tarski(k4_xboole_0(X1,X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_21]), c_0_34])).
cnf(c_0_40,negated_conjecture,(~m1_subset_1(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_41,plain,(r1_tarski(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(esk6_3(X2,X1,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_42,plain,(r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_43,negated_conjecture,(r2_hidden(X1,esk3_0)|r2_hidden(X1,esk5_0)|~m1_subset_1(X1,esk2_0)|~r2_hidden(X1,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_44,plain,(X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45,negated_conjecture,(r1_tarski(k4_xboole_0(esk3_0,esk4_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_30]), c_0_38])])).
cnf(c_0_46,plain,(r1_tarski(k4_xboole_0(X1,X1),k4_xboole_0(X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(X4))), inference(spm,[status(thm)],[c_0_39, c_0_30])).
cnf(c_0_47,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_48,negated_conjecture,(r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~r2_hidden(esk6_3(esk2_0,X1,X2),esk3_0)|~r2_hidden(esk6_3(esk2_0,X1,X2),esk5_0)), inference(spm,[status(thm)],[c_0_40, c_0_23])).
cnf(c_0_49,plain,(r1_tarski(X1,k4_xboole_0(X2,X3))|r2_hidden(esk6_3(X4,X1,k4_xboole_0(X2,X3)),X3)|~m1_subset_1(k4_xboole_0(X2,X3),k1_zfmisc_1(X4))|~m1_subset_1(X1,k1_zfmisc_1(X4))|~r2_hidden(esk6_3(X4,X1,k4_xboole_0(X2,X3)),X2)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50,negated_conjecture,(r1_tarski(X1,X2)|r2_hidden(esk6_3(esk2_0,X1,X2),esk5_0)|r2_hidden(esk6_3(esk2_0,X1,X2),esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~r2_hidden(esk6_3(esk2_0,X1,X2),esk4_0)), inference(spm,[status(thm)],[c_0_43, c_0_23])).
cnf(c_0_51,negated_conjecture,(X1=k4_xboole_0(esk3_0,esk4_0)|~r1_tarski(X1,k4_xboole_0(esk3_0,esk4_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_52,negated_conjecture,(r1_tarski(k4_xboole_0(X1,X1),k4_xboole_0(esk3_0,X2))), inference(spm,[status(thm)],[c_0_46, c_0_38])).
cnf(c_0_53,negated_conjecture,(m1_subset_1(k4_xboole_0(X1,X1),k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_34, c_0_47])).
cnf(c_0_54,negated_conjecture,(r1_tarski(esk5_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~r2_hidden(esk6_3(esk2_0,esk5_0,X1),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_21]), c_0_47])])).
cnf(c_0_55,plain,(r1_tarski(X1,k4_xboole_0(X1,X2))|r2_hidden(esk6_3(X3,X1,k4_xboole_0(X1,X2)),X2)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_21]), c_0_30])).
cnf(c_0_56,negated_conjecture,(r1_tarski(k4_xboole_0(X1,esk5_0),X2)|r2_hidden(esk6_3(esk2_0,k4_xboole_0(X1,esk5_0),X2),esk3_0)|~m1_subset_1(k4_xboole_0(X1,esk5_0),k1_zfmisc_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~r2_hidden(esk6_3(esk2_0,k4_xboole_0(X1,esk5_0),X2),esk4_0)), inference(spm,[status(thm)],[c_0_27, c_0_50])).
cnf(c_0_57,plain,(r1_tarski(k4_xboole_0(k4_xboole_0(X1,X2),X3),X4)|r2_hidden(esk6_3(X5,k4_xboole_0(k4_xboole_0(X1,X2),X3),X4),X1)|~m1_subset_1(k4_xboole_0(k4_xboole_0(X1,X2),X3),k1_zfmisc_1(X5))|~m1_subset_1(X4,k1_zfmisc_1(X5))), inference(spm,[status(thm)],[c_0_14, c_0_32])).
cnf(c_0_58,negated_conjecture,(k4_xboole_0(X1,X1)=k4_xboole_0(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])])).
cnf(c_0_59,plain,(X1=k4_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),X3)|~r2_hidden(esk1_3(X2,X3,X1),X2)|~r2_hidden(esk1_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_60,plain,(k4_xboole_0(X1,X2)=X1|r2_hidden(esk1_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_9])).
cnf(c_0_61,negated_conjecture,(r1_tarski(esk5_0,k4_xboole_0(esk5_0,esk3_0))|~m1_subset_1(k4_xboole_0(esk5_0,esk3_0),k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_47])])).
cnf(c_0_62,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_32]), c_0_30])).
cnf(c_0_63,plain,(X1=k4_xboole_0(X2,k4_xboole_0(X3,X4))|r2_hidden(esk1_3(X2,k4_xboole_0(X3,X4),X1),X4)|r2_hidden(esk1_3(X2,k4_xboole_0(X3,X4),X1),X1)|~r2_hidden(esk1_3(X2,k4_xboole_0(X3,X4),X1),X3)), inference(spm,[status(thm)],[c_0_8, c_0_42])).
cnf(c_0_64,negated_conjecture,(r1_tarski(k4_xboole_0(k4_xboole_0(esk4_0,X1),esk5_0),X2)|r2_hidden(esk6_3(esk2_0,k4_xboole_0(k4_xboole_0(esk4_0,X1),esk5_0),X2),esk3_0)|~m1_subset_1(k4_xboole_0(k4_xboole_0(esk4_0,X1),esk5_0),k1_zfmisc_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65,negated_conjecture,(~r2_hidden(X1,k4_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_33, c_0_58])).
cnf(c_0_66,plain,(k4_xboole_0(X1,X2)=X1|r2_hidden(esk1_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_60])).
cnf(c_0_67,negated_conjecture,(r1_tarski(esk5_0,k4_xboole_0(esk5_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_30]), c_0_47])])).
cnf(c_0_68,negated_conjecture,(r1_tarski(k4_xboole_0(esk5_0,X1),esk5_0)), inference(spm,[status(thm)],[c_0_62, c_0_47])).
cnf(c_0_69,negated_conjecture,(r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_58]), c_0_33])).
cnf(c_0_70,plain,(X1=k4_xboole_0(X2,k4_xboole_0(X2,X3))|r2_hidden(esk1_3(X2,k4_xboole_0(X2,X3),X1),X1)|r2_hidden(esk1_3(X2,k4_xboole_0(X2,X3),X1),X3)), inference(spm,[status(thm)],[c_0_63, c_0_9])).
cnf(c_0_71,negated_conjecture,(r1_tarski(k4_xboole_0(k4_xboole_0(esk4_0,X1),esk5_0),esk3_0)|~m1_subset_1(k4_xboole_0(k4_xboole_0(esk4_0,X1),esk5_0),k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_64]), c_0_38])])).
cnf(c_0_72,negated_conjecture,(k4_xboole_0(X1,k4_xboole_0(esk3_0,esk4_0))=X1), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X3))=X1|~r2_hidden(esk1_3(X1,k4_xboole_0(X2,X3),X1),X3)), inference(spm,[status(thm)],[c_0_12, c_0_66])).
cnf(c_0_74,negated_conjecture,(k4_xboole_0(esk5_0,esk3_0)=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_67]), c_0_68])])).
cnf(c_0_75,negated_conjecture,(X1=k4_xboole_0(esk4_0,X2)|r2_hidden(esk1_3(esk4_0,X2,X1),X2)|~r2_hidden(esk1_3(esk4_0,X2,X1),esk3_0)|~r2_hidden(esk1_3(esk4_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_59, c_0_69])).
cnf(c_0_76,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|r2_hidden(esk1_3(X1,k4_xboole_0(X1,X2),X2),X2)), inference(ef,[status(thm)],[c_0_70])).
cnf(c_0_77,negated_conjecture,(r1_tarski(k4_xboole_0(esk4_0,esk5_0),esk3_0)|~m1_subset_1(k4_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_79,negated_conjecture,(esk3_0!=k7_subset_1(esk2_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_80,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_73, c_0_60])).
cnf(c_0_81,plain,(r1_tarski(k4_xboole_0(X1,k4_xboole_0(X2,X3)),X4)|r2_hidden(esk6_3(X5,k4_xboole_0(X1,k4_xboole_0(X2,X3)),X4),X3)|~m1_subset_1(k4_xboole_0(X1,k4_xboole_0(X2,X3)),k1_zfmisc_1(X5))|~m1_subset_1(X4,k1_zfmisc_1(X5))|~r2_hidden(esk6_3(X5,k4_xboole_0(X1,k4_xboole_0(X2,X3)),X4),X2)), inference(spm,[status(thm)],[c_0_27, c_0_42])).
cnf(c_0_82,negated_conjecture,(~r2_hidden(X1,esk5_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_12, c_0_74])).
cnf(c_0_83,plain,(r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))|r2_hidden(esk6_3(X4,k4_xboole_0(X1,X2),k4_xboole_0(X1,X3)),X3)|~m1_subset_1(k4_xboole_0(X1,X3),k1_zfmisc_1(X4))|~m1_subset_1(k4_xboole_0(X1,X2),k1_zfmisc_1(X4))), inference(spm,[status(thm)],[c_0_49, c_0_32])).
cnf(c_0_84,negated_conjecture,(k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk3_0))=esk3_0|r2_hidden(esk1_3(esk4_0,k4_xboole_0(esk4_0,esk3_0),esk3_0),k4_xboole_0(esk4_0,esk3_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_76])).
cnf(c_0_85,negated_conjecture,(r1_tarski(k4_xboole_0(esk4_0,esk5_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_30]), c_0_78])])).
cnf(c_0_86,negated_conjecture,(k4_xboole_0(esk4_0,esk5_0)!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_26]), c_0_78])])).
cnf(c_0_87,negated_conjecture,(k4_xboole_0(esk3_0,esk5_0)=esk3_0), inference(spm,[status(thm)],[c_0_80, c_0_74])).
cnf(c_0_88,plain,(r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)|r2_hidden(esk6_3(X4,k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3),X2)|~m1_subset_1(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k1_zfmisc_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(X4))), inference(spm,[status(thm)],[c_0_81, c_0_32])).
cnf(c_0_89,plain,(k4_xboole_0(X1,k4_xboole_0(X2,X2))=X1), inference(spm,[status(thm)],[c_0_33, c_0_66])).
cnf(c_0_90,negated_conjecture,(r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,esk5_0))|~m1_subset_1(k4_xboole_0(X1,esk5_0),k1_zfmisc_1(X3))|~m1_subset_1(k4_xboole_0(X1,X2),k1_zfmisc_1(X3))|~r2_hidden(esk6_3(X3,k4_xboole_0(X1,X2),k4_xboole_0(X1,esk5_0)),esk3_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_91,negated_conjecture,(k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk3_0))=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_84]), c_0_76])).
cnf(c_0_92,negated_conjecture,(~r1_tarski(esk3_0,k4_xboole_0(esk4_0,esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_85]), c_0_86])).
cnf(c_0_93,negated_conjecture,(r1_tarski(esk3_0,X1)|r2_hidden(esk6_3(esk2_0,esk3_0,X1),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~r2_hidden(esk6_3(esk2_0,esk3_0,X1),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_87]), c_0_38])])).
cnf(c_0_94,negated_conjecture,(r1_tarski(esk3_0,X1)|r2_hidden(esk6_3(X2,esk3_0,X1),esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_58]), c_0_89]), c_0_89]), c_0_89])).
cnf(c_0_95,negated_conjecture,(~m1_subset_1(k4_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(X1))|~m1_subset_1(esk3_0,k1_zfmisc_1(X1))|~r2_hidden(esk6_3(X1,esk3_0,k4_xboole_0(esk4_0,esk5_0)),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92])).
cnf(c_0_96,negated_conjecture,(r1_tarski(esk3_0,X1)|r2_hidden(esk6_3(esk2_0,esk3_0,X1),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_38])])).
cnf(c_0_97,negated_conjecture,(~m1_subset_1(k4_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_38])]), c_0_92])).
cnf(c_0_98,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_30]), c_0_78])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 99
# Proof object clause steps            : 86
# Proof object formula steps           : 13
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 64
# Proof object simplifying inferences  : 43
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 14755
# ...of these trivial                  : 96
# ...subsumed                          : 13744
# ...remaining for further processing  : 915
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 73
# Backward-rewritten                   : 24
# Generated clauses                    : 60483
# ...of the previous two non-trivial   : 53032
# Contextual simplify-reflections      : 1734
# Paramodulations                      : 60410
# Factorizations                       : 68
# Equation resolutions                 : 5
# Current number of processed clauses  : 793
#    Positive orientable unit clauses  : 92
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 693
# Current number of unprocessed clauses: 34772
# ...number of literals in the above   : 105291
# Current number of archived formulas  : 0
# Current number of archived clauses   : 117
# Clause-clause subsumption calls (NU) : 585672
# Rec. Clause-clause subsumption calls : 483891
# Non-unit clause-clause subsumptions  : 10627
# Unit Clause-clause subsumption calls : 3734
# Rewrite failures with RHS unbound    : 123
# BW rewrite match attempts            : 525
# BW rewrite match successes           : 37
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1088458

# -------------------------------------------------
# User time                : 1.132 s
# System time              : 0.044 s
# Total time               : 1.176 s
# Maximum resident set size: 43088 pages
