# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t85_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3)))))), file('out/t85_tmap_1__10_premises__rerun.E_input', t85_tmap_1)).
fof(t53_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X2))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))))=>(![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>(r2_hidden(X6,u1_struct_0(X3))=>k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X4,X6)=k1_funct_1(X5,X6)))=>r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5)))))))), file('out/t85_tmap_1__10_premises__rerun.E_input', t53_tmap_1)).
fof(d6_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>k4_tmap_1(X1,X2)=k2_tmap_1(X1,X1,k3_struct_0(X1),X2)))), file('out/t85_tmap_1__10_premises__rerun.E_input', d6_tmap_1)).
fof(dt_k3_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>(v1_funct_1(k3_struct_0(X1))&(v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(k3_struct_0(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))))), file('out/t85_tmap_1__10_premises__rerun.E_input', dt_k3_struct_0)).
fof(dt_l1_pre_topc, axiom, (![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1))), file('out/t85_tmap_1__10_premises__rerun.E_input', dt_l1_pre_topc)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t85_tmap_1__10_premises__rerun.E_input', t2_subset)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t85_tmap_1__10_premises__rerun.E_input', fc2_struct_0)).
fof(t18_funct_1, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>k1_funct_1(k4_relat_1(X2),X1)=X1)), file('out/t85_tmap_1__10_premises__rerun.E_input', t18_funct_1)).
fof(d4_struct_0, axiom, (![X1]:(l1_struct_0(X1)=>k3_struct_0(X1)=k6_partfun1(u1_struct_0(X1)))), file('out/t85_tmap_1__10_premises__rerun.E_input', d4_struct_0)).
fof(redefinition_k6_partfun1, axiom, (![X1]:k6_partfun1(X1)=k4_relat_1(X1)), file('out/t85_tmap_1__10_premises__rerun.E_input', redefinition_k6_partfun1)).
fof(redefinition_k3_funct_2, axiom, (![X1]:![X2]:![X3]:![X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4))), file('out/t85_tmap_1__10_premises__rerun.E_input', redefinition_k3_funct_2)).
fof(c_0_11, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3))))))), inference(assume_negation,[status(cth)],[t85_tmap_1])).
fof(c_0_12, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:(((((((m1_subset_1(esk4_5(X7,X8,X9,X10,X11),u1_struct_0(X8))|r2_funct_2(u1_struct_0(X9),u1_struct_0(X7),k2_tmap_1(X8,X7,X10,X9),X11))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X7))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X7)))))))|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X8),u1_struct_0(X7))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X7)))))))|(v2_struct_0(X9)|~m1_pre_topc(X9,X8)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((((r2_hidden(esk4_5(X7,X8,X9,X10,X11),u1_struct_0(X9))|r2_funct_2(u1_struct_0(X9),u1_struct_0(X7),k2_tmap_1(X8,X7,X10,X9),X11))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X7))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X7)))))))|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X8),u1_struct_0(X7))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X7)))))))|(v2_struct_0(X9)|~m1_pre_topc(X9,X8)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((k3_funct_2(u1_struct_0(X8),u1_struct_0(X7),X10,esk4_5(X7,X8,X9,X10,X11))!=k1_funct_1(X11,esk4_5(X7,X8,X9,X10,X11))|r2_funct_2(u1_struct_0(X9),u1_struct_0(X7),k2_tmap_1(X8,X7,X10,X9),X11))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X7))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X7)))))))|(~v1_funct_1(X10)|(~v1_funct_2(X10,u1_struct_0(X8),u1_struct_0(X7))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),u1_struct_0(X7)))))))|(v2_struct_0(X9)|~m1_pre_topc(X9,X8)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t53_tmap_1])])])])])])])])).
fof(c_0_13, negated_conjecture, (![X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&((~m1_subset_1(X8,u1_struct_0(esk1_0))|(~r2_hidden(X8,u1_struct_0(esk2_0))|X8=k1_funct_1(esk3_0,X8)))&~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])])).
fof(c_0_14, plain, (![X3]:![X4]:((v2_struct_0(X3)|(~v2_pre_topc(X3)|~l1_pre_topc(X3)))|(~m1_pre_topc(X4,X3)|k4_tmap_1(X3,X4)=k2_tmap_1(X3,X3,k3_struct_0(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_tmap_1])])])])])])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5)|m1_subset_1(esk4_5(X1,X2,X3,X4,X5),u1_struct_0(X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X5)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,negated_conjecture,(m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,plain,(k4_tmap_1(X1,X2)=k2_tmap_1(X1,X1,k3_struct_0(X1),X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,negated_conjecture,(r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),k2_tmap_1(esk1_0,X1,X2,esk2_0),X3)|m1_subset_1(esk4_5(X1,esk1_0,esk2_0,X2,X3),u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])]), c_0_19]), c_0_20])).
cnf(c_0_23,negated_conjecture,(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_26, plain, (![X2]:((v1_funct_1(k3_struct_0(X2))|~l1_struct_0(X2))&((v1_funct_2(k3_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2))|~l1_struct_0(X2))&(m1_subset_1(k3_struct_0(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|~l1_struct_0(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_struct_0])])])).
cnf(c_0_27,negated_conjecture,(~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_28,negated_conjecture,(k4_tmap_1(esk1_0,esk2_0)=k2_tmap_1(esk1_0,esk1_0,k3_struct_0(esk1_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_16]), c_0_17]), c_0_18])]), c_0_20])).
cnf(c_0_29,negated_conjecture,(r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk1_0,esk1_0,X1,esk2_0),esk3_0)|m1_subset_1(esk4_5(esk1_0,esk1_0,esk2_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_17]), c_0_18])]), c_0_20])).
cnf(c_0_30,plain,(v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31,negated_conjecture,(~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk1_0,esk1_0,k3_struct_0(esk1_0),esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,plain,(v1_funct_1(k3_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,plain,(m1_subset_1(k3_struct_0(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, (![X2]:(~l1_pre_topc(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_35, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_36,negated_conjecture,(m1_subset_1(esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0),u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_37,plain,(l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_38, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_39,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40,negated_conjecture,(m1_subset_1(esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_17])])).
cnf(c_0_41,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_43, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|k1_funct_1(k4_relat_1(X4),X3)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_funct_1])])).
cnf(c_0_44,negated_conjecture,(r2_hidden(esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0),u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_20])).
fof(c_0_45, plain, (![X2]:(~l1_struct_0(X2)|k3_struct_0(X2)=k6_partfun1(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_struct_0])])).
fof(c_0_46, plain, (![X2]:k6_partfun1(X2)=k4_relat_1(X2)), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
cnf(c_0_47,plain,(k1_funct_1(k4_relat_1(X1),X2)=X2|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48,negated_conjecture,(r2_hidden(esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_37]), c_0_17])])).
cnf(c_0_49,plain,(k3_struct_0(X1)=k6_partfun1(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_50,plain,(k6_partfun1(X1)=k4_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_51, plain, (![X5]:![X6]:![X7]:![X8]:((v1_xboole_0(X5)|((~v1_funct_1(X7)|(~v1_funct_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X8,X5)))|k3_funct_2(X5,X6,X7,X8)=k1_funct_1(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_52,negated_conjecture,(k1_funct_1(k4_relat_1(u1_struct_0(esk1_0)),esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0))=esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53,plain,(k3_struct_0(X1)=k4_relat_1(u1_struct_0(X1))|~l1_struct_0(X1)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5)|r2_hidden(esk4_5(X1,X2,X3,X4,X5),u1_struct_0(X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X5)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_55,plain,(k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56,negated_conjecture,(k1_funct_1(k3_struct_0(esk1_0),esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0))=esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0)|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_57,negated_conjecture,(r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),k2_tmap_1(esk1_0,X1,X2,esk2_0),X3)|r2_hidden(esk4_5(X1,esk1_0,esk2_0,X2,X3),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_16]), c_0_17]), c_0_18])]), c_0_19]), c_0_20])).
cnf(c_0_58,plain,(k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),k3_struct_0(X1),X2)=k1_funct_1(k3_struct_0(X1),X2)|v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_30]), c_0_32]), c_0_33])).
cnf(c_0_59,negated_conjecture,(k1_funct_1(k3_struct_0(esk1_0),esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0))=esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_37]), c_0_17])])).
cnf(c_0_60,negated_conjecture,(r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk1_0,esk1_0,X1,esk2_0),esk3_0)|r2_hidden(esk4_5(esk1_0,esk1_0,esk2_0,X1,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_23]), c_0_24]), c_0_25]), c_0_17]), c_0_18])]), c_0_20])).
cnf(c_0_61,negated_conjecture,(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k3_struct_0(esk1_0),esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0))=esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_40]), c_0_59])).
cnf(c_0_62,negated_conjecture,(r2_hidden(esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0),u1_struct_0(esk2_0))|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_63,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X5)|k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X4,esk4_5(X1,X2,X3,X4,X5))!=k1_funct_1(X5,esk4_5(X1,X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_64,negated_conjecture,(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k3_struct_0(esk1_0),esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0))=esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_37]), c_0_17])])).
cnf(c_0_65,negated_conjecture,(X1=k1_funct_1(esk3_0,X1)|~r2_hidden(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_66,negated_conjecture,(r2_hidden(esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_37]), c_0_17])])).
cnf(c_0_67,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|k1_funct_1(esk3_0,esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0))!=esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0)|~m1_subset_1(k3_struct_0(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(k3_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(k3_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_24]), c_0_23]), c_0_25]), c_0_16]), c_0_17]), c_0_18])]), c_0_31]), c_0_19]), c_0_20])).
cnf(c_0_68,negated_conjecture,(k1_funct_1(esk3_0,esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0))=esk4_5(esk1_0,esk1_0,esk2_0,k3_struct_0(esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_40])])).
cnf(c_0_69,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(k3_struct_0(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(k3_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(k3_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_70,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_30]), c_0_32]), c_0_33])).
cnf(c_0_71,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_37]), c_0_17])])).
cnf(c_0_72,negated_conjecture,(~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_71]), c_0_20])).
cnf(c_0_73,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_37]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 74
# Proof object clause steps            : 51
# Proof object formula steps           : 23
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 69
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 23
# Processed clauses                    : 75
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 75
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 10
# Generated clauses                    : 30
# ...of the previous two non-trivial   : 31
# Contextual simplify-reflections      : 9
# Paramodulations                      : 30
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 11
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 753
# Rec. Clause-clause subsumption calls : 96
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4129

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45240 pages
