# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t86_funct_1, conjecture, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>![X4]:((v1_relat_1(X4)&v1_funct_1(X4))=>((X1=k10_xtuple_0(X2)&(k9_xtuple_0(X3)=X1&(k9_xtuple_0(X4)=X1&k3_relat_1(X2,X3)=k3_relat_1(X2,X4))))=>X3=X4))))), file('out/t86_funct_1__3_premises__rerun.E_input', t86_funct_1)).
fof(t2_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((k9_xtuple_0(X1)=k9_xtuple_0(X2)&![X3]:(r2_hidden(X3,k9_xtuple_0(X1))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3)))=>X1=X2)))), file('out/t86_funct_1__3_premises__rerun.E_input', t2_funct_1)).
fof(d3_funct_1, axiom, (![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X1))&X3=k1_funct_1(X1,X4)))))), file('out/t86_funct_1__3_premises__rerun.E_input', d3_funct_1)).
fof(t13_funct_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k9_xtuple_0(X2))=>k1_funct_1(k3_relat_1(X2,X3),X1)=k1_funct_1(X3,k1_funct_1(X2,X1)))))), file('out/t86_funct_1__3_premises__rerun.E_input', t13_funct_1)).
fof(c_0_4, negated_conjecture, (~(![X1]:![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>![X4]:((v1_relat_1(X4)&v1_funct_1(X4))=>((X1=k10_xtuple_0(X2)&(k9_xtuple_0(X3)=X1&(k9_xtuple_0(X4)=X1&k3_relat_1(X2,X3)=k3_relat_1(X2,X4))))=>X3=X4)))))), inference(assume_negation,[status(cth)],[t86_funct_1])).
fof(c_0_5, plain, (![X4]:![X5]:(((((r2_hidden(esk5_2(X4,X5),k9_xtuple_0(X4))|k9_xtuple_0(X4)!=k9_xtuple_0(X5))|X4=X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))|(~v1_relat_1(X4)|~v1_funct_1(X4)))&((((k1_funct_1(X4,esk5_2(X4,X5))!=k1_funct_1(X5,esk5_2(X4,X5))|k9_xtuple_0(X4)!=k9_xtuple_0(X5))|X4=X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))|(~v1_relat_1(X4)|~v1_funct_1(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_funct_1])])])])])])])).
fof(c_0_6, negated_conjecture, (((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((v1_relat_1(esk4_0)&v1_funct_1(esk4_0))&((esk1_0=k10_xtuple_0(esk2_0)&(k9_xtuple_0(esk3_0)=esk1_0&(k9_xtuple_0(esk4_0)=esk1_0&k3_relat_1(esk2_0,esk3_0)=k3_relat_1(esk2_0,esk4_0))))&esk3_0!=esk4_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_7,plain,(X1=X2|r2_hidden(esk5_2(X1,X2),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|k9_xtuple_0(X1)!=k9_xtuple_0(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,negated_conjecture,(k9_xtuple_0(esk4_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:![X6]:![X11]:((((((r2_hidden(esk6_3(X5,X6,X7),k9_xtuple_0(X5))|~r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((X7=k1_funct_1(X5,esk6_3(X5,X6,X7))|~r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(X9,k9_xtuple_0(X5))|X7!=k1_funct_1(X5,X9))|r2_hidden(X7,X6))|X6!=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((((~r2_hidden(esk7_2(X5,X6),X6)|(~r2_hidden(X11,k9_xtuple_0(X5))|esk7_2(X5,X6)!=k1_funct_1(X5,X11)))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&((((r2_hidden(esk8_2(X5,X6),k9_xtuple_0(X5))|r2_hidden(esk7_2(X5,X6),X6))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(((esk7_2(X5,X6)=k1_funct_1(X5,esk8_2(X5,X6))|r2_hidden(esk7_2(X5,X6),X6))|X6=k10_xtuple_0(X5))|(~v1_relat_1(X5)|~v1_funct_1(X5))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_1])])])])])])])).
cnf(c_0_12,negated_conjecture,(X1=esk4_0|r2_hidden(esk5_2(X1,esk4_0),k9_xtuple_0(X1))|k9_xtuple_0(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_13,negated_conjecture,(k9_xtuple_0(esk3_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,negated_conjecture,(v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,negated_conjecture,(esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17,plain,(r2_hidden(esk6_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(r2_hidden(esk5_2(esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_19,plain,(X3=k1_funct_1(X1,esk6_3(X1,X2,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:((~v1_relat_1(X5)|~v1_funct_1(X5))|((~v1_relat_1(X6)|~v1_funct_1(X6))|(~r2_hidden(X4,k9_xtuple_0(X5))|k1_funct_1(k3_relat_1(X5,X6),X4)=k1_funct_1(X6,k1_funct_1(X5,X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_funct_1])])])])])).
cnf(c_0_21,negated_conjecture,(r2_hidden(esk6_3(X1,esk1_0,esk5_2(esk3_0,esk4_0)),k9_xtuple_0(X1))|k10_xtuple_0(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,negated_conjecture,(esk1_0=k10_xtuple_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23,negated_conjecture,(v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25,negated_conjecture,(k1_funct_1(X1,esk6_3(X1,esk1_0,esk5_2(esk3_0,esk4_0)))=esk5_2(esk3_0,esk4_0)|k10_xtuple_0(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_18])).
cnf(c_0_26,plain,(k1_funct_1(k3_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|~r2_hidden(X3,k9_xtuple_0(X1))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,negated_conjecture,(r2_hidden(esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)),k9_xtuple_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_28,negated_conjecture,(k1_funct_1(esk2_0,esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)))=esk5_2(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_29,negated_conjecture,(k1_funct_1(k3_relat_1(esk2_0,X1),esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)))=k1_funct_1(X1,esk5_2(esk3_0,esk4_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_23]), c_0_24])]), c_0_28])).
cnf(c_0_30,negated_conjecture,(k3_relat_1(esk2_0,esk3_0)=k3_relat_1(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31,plain,(X1=X2|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|k9_xtuple_0(X1)!=k9_xtuple_0(X2)|k1_funct_1(X1,esk5_2(X1,X2))!=k1_funct_1(X2,esk5_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_32,negated_conjecture,(k1_funct_1(k3_relat_1(esk2_0,esk3_0),esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)))=k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_8]), c_0_30]), c_0_10])])).
cnf(c_0_33,negated_conjecture,(X1=esk4_0|k1_funct_1(X1,esk5_2(X1,esk4_0))!=k1_funct_1(esk4_0,esk5_2(X1,esk4_0))|k9_xtuple_0(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_34,negated_conjecture,(k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=k1_funct_1(esk3_0,esk5_2(esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_14]), c_0_15])]), c_0_32])).
cnf(c_0_35,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_13]), c_0_14]), c_0_15])]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 27
# Proof object formula steps           : 9
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 145
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 139
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 24
# Generated clauses                    : 189
# ...of the previous two non-trivial   : 204
# Contextual simplify-reflections      : 0
# Paramodulations                      : 155
# Factorizations                       : 0
# Equation resolutions                 : 34
# Current number of processed clauses  : 95
#    Positive orientable unit clauses  : 38
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 56
# Current number of unprocessed clauses: 78
# ...number of literals in the above   : 288
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 569
# Rec. Clause-clause subsumption calls : 383
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 131
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6170

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45268 pages
