# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.004 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t88_mcart_1, conjecture, (![X1]:(v1_relat_1(X1)=>![X2]:(r2_hidden(X2,X1)=>r2_hidden(k2_xtuple_0(X2),k9_relat_1(X1,k1_xtuple_0(X2)))))), file('out/t88_mcart_1__32_premises__rerun__rerun.E_input', t88_mcart_1)).
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t88_mcart_1__32_premises__rerun__rerun.E_input', t7_boole)).
fof(t87_mcart_1, axiom, (![X1]:((~(v1_xboole_0(X1))&v1_relat_1(X1))=>![X2]:k9_relat_1(X1,X2)=a_2_0_mcart_1(X1,X2))), file('out/t88_mcart_1__32_premises__rerun__rerun.E_input', t87_mcart_1)).
fof(fraenkel_a_2_0_mcart_1, axiom, (![X1]:![X2]:![X3]:((~(v1_xboole_0(X2))&v1_relat_1(X2))=>(r2_hidden(X1,a_2_0_mcart_1(X2,X3))<=>?[X4]:(m1_subset_1(X4,X2)&(X1=k2_xtuple_0(X4)&k1_xtuple_0(X4)=X3))))), file('out/t88_mcart_1__32_premises__rerun__rerun.E_input', fraenkel_a_2_0_mcart_1)).
fof(t1_subset, axiom, (![X1]:![X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2))), file('out/t88_mcart_1__32_premises__rerun__rerun.E_input', t1_subset)).
fof(c_0_5, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>![X2]:(r2_hidden(X2,X1)=>r2_hidden(k2_xtuple_0(X2),k9_relat_1(X1,k1_xtuple_0(X2))))))), inference(assume_negation,[status(cth)],[t88_mcart_1])).
fof(c_0_6, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_7, negated_conjecture, ((v1_relat_1(esk1_0)&(r2_hidden(esk2_0,esk1_0)&~r2_hidden(k2_xtuple_0(esk2_0),k9_relat_1(esk1_0,k1_xtuple_0(esk2_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
fof(c_0_8, plain, (![X3]:![X4]:((v1_xboole_0(X3)|~v1_relat_1(X3))|k9_relat_1(X3,X4)=a_2_0_mcart_1(X3,X4))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t87_mcart_1])])])])])])).
cnf(c_0_9,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(r2_hidden(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X7]:![X9]:((((m1_subset_1(esk3_3(X5,X6,X7),X6)|~r2_hidden(X5,a_2_0_mcart_1(X6,X7)))|(v1_xboole_0(X6)|~v1_relat_1(X6)))&(((X5=k2_xtuple_0(esk3_3(X5,X6,X7))|~r2_hidden(X5,a_2_0_mcart_1(X6,X7)))|(v1_xboole_0(X6)|~v1_relat_1(X6)))&((k1_xtuple_0(esk3_3(X5,X6,X7))=X7|~r2_hidden(X5,a_2_0_mcart_1(X6,X7)))|(v1_xboole_0(X6)|~v1_relat_1(X6)))))&(((~m1_subset_1(X9,X6)|(X5!=k2_xtuple_0(X9)|k1_xtuple_0(X9)!=X7))|r2_hidden(X5,a_2_0_mcart_1(X6,X7)))|(v1_xboole_0(X6)|~v1_relat_1(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_mcart_1])])])])])])])])).
cnf(c_0_12,negated_conjecture,(~r2_hidden(k2_xtuple_0(esk2_0),k9_relat_1(esk1_0,k1_xtuple_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,plain,(k9_relat_1(X1,X2)=a_2_0_mcart_1(X1,X2)|v1_xboole_0(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,negated_conjecture,(v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,negated_conjecture,(~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_16,plain,(v1_xboole_0(X1)|r2_hidden(X2,a_2_0_mcart_1(X1,X3))|~v1_relat_1(X1)|k1_xtuple_0(X4)!=X3|X2!=k2_xtuple_0(X4)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(~r2_hidden(k2_xtuple_0(esk2_0),a_2_0_mcart_1(esk1_0,k1_xtuple_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_18,plain,(v1_xboole_0(X1)|r2_hidden(k2_xtuple_0(X2),a_2_0_mcart_1(X1,k1_xtuple_0(X2)))|~m1_subset_1(X2,X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_16])])).
fof(c_0_19, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|m1_subset_1(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_20,negated_conjecture,(~m1_subset_1(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_14])]), c_0_15])).
cnf(c_0_21,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_10])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 12
# Proof object formula steps           : 11
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 10
# Processed clauses                    : 24
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 24
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 8
# ...of the previous two non-trivial   : 7
# Contextual simplify-reflections      : 0
# Paramodulations                      : 7
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 7
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 18
# Current number of archived formulas  : 0
# Current number of archived clauses   : 10
# Clause-clause subsumption calls (NU) : 46
# Rec. Clause-clause subsumption calls : 22
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 860

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 45100 pages
