# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_funct_1, conjecture, (![X1]:![X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k9_xtuple_0(X3)&r1_tarski(k10_xtuple_0(X3),X1))))))), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', t8_funct_1)).
fof(t5_funct_1, axiom, (![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k9_xtuple_0(X3)=X1&k10_xtuple_0(X3)=k1_tarski(X2))))), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', t5_funct_1)).
fof(t42_relat_1, axiom, (![X1]:(v1_relat_1(X1)=>(k9_xtuple_0(X1)=k1_xboole_0<=>k10_xtuple_0(X1)=k1_xboole_0))), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', t42_relat_1)).
fof(t2_xboole_1, axiom, (![X1]:r1_tarski(k1_xboole_0,X1)), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', t2_xboole_1)).
fof(l1_zfmisc_1, axiom, (![X1]:![X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2))), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', l1_zfmisc_1)).
fof(s3_funct_1__e7_15__funct_1, axiom, (![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k9_xtuple_0(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=1)))), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', s3_funct_1__e7_15__funct_1)).
fof(t1_xregular, axiom, (![X1]:(~(v1_xboole_0(X1))=>?[X2]:(r2_hidden(X2,X1)&r1_xboole_0(X2,X1)))), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', t1_xregular)).
fof(l13_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t8_funct_1__8_premises__rerun__rerun.E_input', l13_xboole_0)).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k9_xtuple_0(X3)&r1_tarski(k10_xtuple_0(X3),X1)))))))), inference(assume_negation,[status(cth)],[t8_funct_1])).
fof(c_0_9, negated_conjecture, (![X6]:((esk1_0!=k1_xboole_0|esk2_0=k1_xboole_0)&((~v1_relat_1(X6)|~v1_funct_1(X6))|(esk2_0!=k9_xtuple_0(X6)|~r1_tarski(k10_xtuple_0(X6),esk1_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])])).
fof(c_0_10, plain, (![X4]:![X5]:(((v1_relat_1(esk3_2(X4,X5))|X4=k1_xboole_0)&(v1_funct_1(esk3_2(X4,X5))|X4=k1_xboole_0))&((k9_xtuple_0(esk3_2(X4,X5))=X4|X4=k1_xboole_0)&(k10_xtuple_0(esk3_2(X4,X5))=k1_tarski(X5)|X4=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_funct_1])])])])])])])])).
fof(c_0_11, plain, (![X2]:(((k9_xtuple_0(X2)!=k1_xboole_0|k10_xtuple_0(X2)=k1_xboole_0)|~v1_relat_1(X2))&((k10_xtuple_0(X2)!=k1_xboole_0|k9_xtuple_0(X2)=k1_xboole_0)|~v1_relat_1(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t42_relat_1])])])).
fof(c_0_12, plain, (![X2]:r1_tarski(k1_xboole_0,X2)), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_13,negated_conjecture,(~r1_tarski(k10_xtuple_0(X1),esk1_0)|esk2_0!=k9_xtuple_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(X1=k1_xboole_0|k10_xtuple_0(esk3_2(X1,X2))=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(X1=k1_xboole_0|v1_relat_1(esk3_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(X1=k1_xboole_0|v1_funct_1(esk3_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, plain, (![X3]:![X4]:![X3]:![X4]:((~r1_tarski(k1_tarski(X3),X4)|r2_hidden(X3,X4))&(~r2_hidden(X3,X4)|r1_tarski(k1_tarski(X3),X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])])])).
cnf(c_0_18,plain,(k10_xtuple_0(X1)=k1_xboole_0|~v1_relat_1(X1)|k9_xtuple_0(X1)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,(r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_20, plain, (![X4]:![X6]:((v1_relat_1(esk4_1(X4))&v1_funct_1(esk4_1(X4)))&(k9_xtuple_0(esk4_1(X4))=X4&(~r2_hidden(X6,X4)|k1_funct_1(esk4_1(X4),X6)=1)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e7_15__funct_1])])])])])])).
cnf(c_0_21,negated_conjecture,(X1=k1_xboole_0|k9_xtuple_0(esk3_2(X1,X2))!=esk2_0|~r1_tarski(k1_tarski(X2),esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])).
cnf(c_0_22,plain,(r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,negated_conjecture,(k9_xtuple_0(X1)!=esk2_0|k9_xtuple_0(X1)!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_18]), c_0_19])])).
cnf(c_0_24,plain,(v1_funct_1(esk4_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(v1_relat_1(esk4_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,plain,(k9_xtuple_0(esk4_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27,negated_conjecture,(X1=k1_xboole_0|k9_xtuple_0(esk3_2(X1,X2))!=esk2_0|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28,plain,(X1=k1_xboole_0|k9_xtuple_0(esk3_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29,negated_conjecture,(esk2_0!=k1_xboole_0), inference(er,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26]), c_0_26])])).
fof(c_0_30, plain, (![X3]:((r2_hidden(esk5_1(X3),X3)|v1_xboole_0(X3))&(r1_xboole_0(esk5_1(X3),X3)|v1_xboole_0(X3)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_xregular])])])])])])])).
fof(c_0_31, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_32,negated_conjecture,(~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28])]), c_0_29])).
cnf(c_0_33,plain,(v1_xboole_0(X1)|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_34,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_35,negated_conjecture,(v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,negated_conjecture,(esk2_0=k1_xboole_0|esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37,negated_conjecture,(esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_37])]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 22
# Proof object formula steps           : 17
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 47
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 19
# Contextual simplify-reflections      : 4
# Paramodulations                      : 16
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 15
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 18
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 51
# Rec. Clause-clause subsumption calls : 43
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 25
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1191

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45244 pages
