# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_subset_1, conjecture, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>r2_hidden(X4,X3)))=>X2=X3)))), file('out/t8_subset_1__4_premises__rerun.E_input', t8_subset_1)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t8_subset_1__4_premises__rerun.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t8_subset_1__4_premises__rerun.E_input', d1_xboole_0)).
fof(l3_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1)))), file('out/t8_subset_1__4_premises__rerun.E_input', l3_subset_1)).
fof(t2_tarski, axiom, (![X1]:![X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2)), file('out/t8_subset_1__4_premises__rerun.E_input', t2_tarski)).
fof(c_0_5, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>r2_hidden(X4,X3)))=>X2=X3))))), inference(assume_negation,[status(cth)],[t8_subset_1])).
fof(c_0_6, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
fof(c_0_7, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk4_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
fof(c_0_8, negated_conjecture, (![X8]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((((~r2_hidden(X8,esk2_0)|r2_hidden(X8,esk3_0))|~m1_subset_1(X8,esk1_0))&((~r2_hidden(X8,esk3_0)|r2_hidden(X8,esk2_0))|~m1_subset_1(X8,esk1_0)))&esk2_0!=esk3_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X5,k1_zfmisc_1(X4))|(~r2_hidden(X6,X5)|r2_hidden(X6,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])])])).
cnf(c_0_10,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(v1_xboole_0(X1)|r2_hidden(X2,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_18,negated_conjecture,(v1_xboole_0(esk3_0)|~v1_xboole_0(k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19,negated_conjecture,(v1_xboole_0(k1_zfmisc_1(esk1_0))|r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_14, c_0_13])).
cnf(c_0_20,negated_conjecture,(v1_xboole_0(esk2_0)|~v1_xboole_0(k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_12, c_0_15])).
cnf(c_0_21,negated_conjecture,(v1_xboole_0(k1_zfmisc_1(esk1_0))|r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22,plain,(r2_hidden(X1,X2)|~r2_hidden(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,negated_conjecture,(v1_xboole_0(esk3_0)|r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_24, plain, (![X4]:![X5]:(((~r2_hidden(esk5_2(X4,X5),X4)|~r2_hidden(esk5_2(X4,X5),X5))|X4=X5)&((r2_hidden(esk5_2(X4,X5),X4)|r2_hidden(esk5_2(X4,X5),X5))|X4=X5))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])])])).
cnf(c_0_25,negated_conjecture,(r2_hidden(X1,esk3_0)|~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_26,negated_conjecture,(v1_xboole_0(esk2_0)|r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_11])).
cnf(c_0_28,plain,(X1=X2|r2_hidden(esk5_2(X1,X2),X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29,negated_conjecture,(r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_25, c_0_17])).
cnf(c_0_30,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_26]), c_0_11])).
cnf(c_0_31,negated_conjecture,(X1=esk3_0|r2_hidden(esk5_2(X1,esk3_0),esk1_0)|r2_hidden(esk5_2(X1,esk3_0),X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,negated_conjecture,(esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_33,negated_conjecture,(esk2_0=X1|r2_hidden(esk5_2(esk2_0,X1),esk3_0)|r2_hidden(esk5_2(esk2_0,X1),X1)|~r2_hidden(esk5_2(esk2_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_34,negated_conjecture,(r2_hidden(esk5_2(esk2_0,esk3_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_35,plain,(X1=X2|~r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36,negated_conjecture,(r2_hidden(esk5_2(esk2_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_32])).
cnf(c_0_37,negated_conjecture,(r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_38,negated_conjecture,(~r2_hidden(esk5_2(esk2_0,esk3_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_32])).
cnf(c_0_39,negated_conjecture,(r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_17])).
cnf(c_0_40,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_36]), c_0_34])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 41
# Proof object clause steps            : 30
# Proof object formula steps           : 11
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 14
# Processed clauses                    : 156
# ...of these trivial                  : 2
# ...subsumed                          : 53
# ...remaining for further processing  : 101
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 8
# Generated clauses                    : 200
# ...of the previous two non-trivial   : 182
# Contextual simplify-reflections      : 25
# Paramodulations                      : 188
# Factorizations                       : 6
# Equation resolutions                 : 0
# Current number of processed clauses  : 66
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 40
# ...number of literals in the above   : 168
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 400
# Rec. Clause-clause subsumption calls : 280
# Non-unit clause-clause subsumptions  : 59
# Unit Clause-clause subsumption calls : 42
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3320

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 42956 pages
