# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_subset_1, conjecture, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>r2_hidden(X4,X3)))=>X2=X3)))), file('out/t8_subset_1__7_premises__rerun.E_input', t8_subset_1)).
fof(d2_subset_1, axiom, (![X1]:![X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2))))), file('out/t8_subset_1__7_premises__rerun.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, (![X1]:~(v1_xboole_0(k1_zfmisc_1(X1)))), file('out/t8_subset_1__7_premises__rerun.E_input', fc1_subset_1)).
fof(d1_zfmisc_1, axiom, (![X1]:![X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1)))), file('out/t8_subset_1__7_premises__rerun.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t8_subset_1__7_premises__rerun.E_input', d3_tarski)).
fof(d1_xboole_0, axiom, (![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1)))), file('out/t8_subset_1__7_premises__rerun.E_input', d1_xboole_0)).
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t8_subset_1__7_premises__rerun.E_input', d4_xboole_0)).
fof(t28_xboole_1, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1)), file('out/t8_subset_1__7_premises__rerun.E_input', t28_xboole_1)).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>r2_hidden(X4,X3)))=>X2=X3))))), inference(assume_negation,[status(cth)],[t8_subset_1])).
fof(c_0_9, plain, (![X3]:![X4]:![X4]:![X3]:![X4]:![X4]:((((~m1_subset_1(X4,X3)|r2_hidden(X4,X3))|v1_xboole_0(X3))&((~r2_hidden(X4,X3)|m1_subset_1(X4,X3))|v1_xboole_0(X3)))&(((~m1_subset_1(X4,X3)|v1_xboole_0(X4))|~v1_xboole_0(X3))&((~v1_xboole_0(X4)|m1_subset_1(X4,X3))|~v1_xboole_0(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])])])).
fof(c_0_10, negated_conjecture, (![X8]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((((~r2_hidden(X8,esk2_0)|r2_hidden(X8,esk3_0))|~m1_subset_1(X8,esk1_0))&((~r2_hidden(X8,esk3_0)|r2_hidden(X8,esk2_0))|~m1_subset_1(X8,esk1_0)))&esk2_0!=esk3_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])])])).
fof(c_0_11, plain, (![X2]:~v1_xboole_0(k1_zfmisc_1(X2))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:![X6]:![X4]:![X5]:((((~r2_hidden(X6,X5)|r1_tarski(X6,X4))|X5!=k1_zfmisc_1(X4))&((~r1_tarski(X6,X4)|r2_hidden(X6,X5))|X5!=k1_zfmisc_1(X4)))&(((~r2_hidden(esk4_2(X4,X5),X5)|~r1_tarski(esk4_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))&((r2_hidden(esk4_2(X4,X5),X5)|r1_tarski(esk4_2(X4,X5),X4))|X5=k1_zfmisc_1(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])])).
cnf(c_0_13,plain,(v1_xboole_0(X1)|r2_hidden(X2,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(r1_tarski(X3,X2)|X1!=k1_zfmisc_1(X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk5_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk5_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_19,negated_conjecture,(r1_tarski(esk3_0,X1)|k1_zfmisc_1(esk1_0)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_20, plain, (![X3]:![X4]:![X3]:((~v1_xboole_0(X3)|~r2_hidden(X4,X3))&(r2_hidden(esk6_1(X3),X3)|v1_xboole_0(X3)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])])).
cnf(c_0_21,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22,negated_conjecture,(r1_tarski(esk3_0,esk1_0)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_24,plain,(v1_xboole_0(X1)|m1_subset_1(X2,X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25,plain,(~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27,plain,(r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28,negated_conjecture,(r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_23]), c_0_15])).
cnf(c_0_29,plain,(m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30,negated_conjecture,(r1_tarski(esk3_0,X1)|r2_hidden(esk5_2(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_31, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk7_3(X5,X6,X7),X7)|(~r2_hidden(esk7_3(X5,X6,X7),X5)|~r2_hidden(esk7_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk7_3(X5,X6,X7),X5)|r2_hidden(esk7_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk7_3(X5,X6,X7),X6)|r2_hidden(esk7_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
cnf(c_0_32,negated_conjecture,(r1_tarski(esk2_0,X1)|k1_zfmisc_1(esk1_0)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_16, c_0_28])).
cnf(c_0_33,negated_conjecture,(r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_34,negated_conjecture,(r1_tarski(esk3_0,X1)|m1_subset_1(esk5_2(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk7_3(X2,X3,X1),X1)|r2_hidden(esk7_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36,negated_conjecture,(r1_tarski(esk2_0,esk1_0)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_37,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk7_3(X2,X3,X1),X1)|r2_hidden(esk7_3(X2,X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_38, plain, (![X3]:![X4]:(~r1_tarski(X3,X4)|k3_xboole_0(X3,X4)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_39,plain,(r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_40,negated_conjecture,(r1_tarski(esk3_0,X1)|r2_hidden(esk5_2(esk3_0,X1),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_27])).
cnf(c_0_41,plain,(r2_hidden(X4,X1)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42,plain,(k3_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_35])).
cnf(c_0_43,negated_conjecture,(r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_21, c_0_36])).
cnf(c_0_44,negated_conjecture,(X1=k3_xboole_0(esk3_0,X2)|r2_hidden(esk7_3(esk3_0,X2,X1),esk1_0)|r2_hidden(esk7_3(esk3_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_26, c_0_37])).
cnf(c_0_45,plain,(k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46,negated_conjecture,(r1_tarski(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_47,plain,(k3_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X3)|X3!=k3_xboole_0(X2,X4)|~r2_hidden(esk7_3(X1,X2,X2),X4)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48,negated_conjecture,(k3_xboole_0(esk3_0,X1)=esk2_0|r2_hidden(esk7_3(esk3_0,X1,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49,negated_conjecture,(k3_xboole_0(esk3_0,esk2_0)=esk3_0), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50,negated_conjecture,(k3_xboole_0(esk2_0,esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_45, c_0_36])).
cnf(c_0_51,negated_conjecture,(esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_52,plain,(r2_hidden(X4,X3)|X1!=k3_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_53,negated_conjecture,(r2_hidden(esk7_3(esk3_0,esk2_0,esk2_0),X1)|X1!=esk2_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_54,negated_conjecture,(r2_hidden(X1,esk1_0)|X2!=esk2_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_50])).
cnf(c_0_55,negated_conjecture,(r2_hidden(esk7_3(esk3_0,esk2_0,esk2_0),esk2_0)), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_56,negated_conjecture,(r2_hidden(esk7_3(esk3_0,esk2_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_57,plain,(X1=k3_xboole_0(X2,X3)|~r2_hidden(esk7_3(X2,X3,X1),X3)|~r2_hidden(esk7_3(X2,X3,X1),X2)|~r2_hidden(esk7_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_58,negated_conjecture,(r2_hidden(X1,esk3_0)|~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_59,negated_conjecture,(m1_subset_1(esk7_3(esk3_0,esk2_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_56])).
cnf(c_0_60,negated_conjecture,(~r2_hidden(esk7_3(esk3_0,esk2_0,esk2_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_55]), c_0_49]), c_0_55])]), c_0_51])).
cnf(c_0_61,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_55])]), c_0_60]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 45
# Proof object formula steps           : 17
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 4391
# ...of these trivial                  : 65
# ...subsumed                          : 2513
# ...remaining for further processing  : 1813
# Other redundant clauses eliminated   : 216
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 25
# Backward-rewritten                   : 363
# Generated clauses                    : 16964
# ...of the previous two non-trivial   : 16252
# Contextual simplify-reflections      : 756
# Paramodulations                      : 16180
# Factorizations                       : 56
# Equation resolutions                 : 609
# Current number of processed clauses  : 1306
#    Positive orientable unit clauses  : 158
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 1138
# Current number of unprocessed clauses: 8064
# ...number of literals in the above   : 29368
# Current number of archived formulas  : 0
# Current number of archived clauses   : 507
# Clause-clause subsumption calls (NU) : 397978
# Rec. Clause-clause subsumption calls : 274848
# Non-unit clause-clause subsumptions  : 3261
# Unit Clause-clause subsumption calls : 13260
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 874
# BW rewrite match successes           : 31
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 261245

# -------------------------------------------------
# User time                : 0.276 s
# System time              : 0.004 s
# Total time               : 0.280 s
# Maximum resident set size: 42956 pages
