# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_9, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>r1_waybel_0(X1,X2,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)))))), file('out/t8_waybel_9__9_premises__rerun.E_input', t8_waybel_9)).
fof(dt_k4_yellow_6, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&l1_waybel_0(X2,X1))=>m1_subset_1(k4_yellow_6(X1,X2),u1_struct_0(X1)))), file('out/t8_waybel_9__9_premises__rerun.E_input', dt_k4_yellow_6)).
fof(rc4_waybel_0, axiom, (![X1]:(l1_struct_0(X1)=>?[X2]:(l1_waybel_0(X2,X1)&v6_waybel_0(X2,X1)))), file('out/t8_waybel_9__9_premises__rerun.E_input', rc4_waybel_0)).
fof(dt_l1_waybel_0, axiom, (![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2)))), file('out/t8_waybel_9__9_premises__rerun.E_input', dt_l1_waybel_0)).
fof(d11_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))))), file('out/t8_waybel_9__9_premises__rerun.E_input', d11_waybel_0)).
fof(dt_l1_orders_2, axiom, (![X1]:(l1_orders_2(X1)=>l1_struct_0(X1))), file('out/t8_waybel_9__9_premises__rerun.E_input', dt_l1_orders_2)).
fof(t189_funct_2, axiom, (![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>r2_hidden(k3_funct_2(X1,X2,X4,X3),k2_relset_1(X1,X2,X4))))))), file('out/t8_waybel_9__9_premises__rerun.E_input', t189_funct_2)).
fof(dt_u1_waybel_0, axiom, (![X1]:![X2]:((l1_struct_0(X1)&l1_waybel_0(X2,X1))=>(v1_funct_1(u1_waybel_0(X1,X2))&(v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1)))))))), file('out/t8_waybel_9__9_premises__rerun.E_input', dt_u1_waybel_0)).
fof(d8_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>k2_waybel_0(X1,X2,X3)=k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),X3))))), file('out/t8_waybel_9__9_premises__rerun.E_input', d8_waybel_0)).
fof(fc2_struct_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('out/t8_waybel_9__9_premises__rerun.E_input', fc2_struct_0)).
fof(c_0_10, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>r1_waybel_0(X1,X2,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2))))))), inference(assume_negation,[status(cth)],[t8_waybel_9])).
fof(c_0_11, plain, (![X3]:![X4]:(((v2_struct_0(X3)|~l1_struct_0(X3))|~l1_waybel_0(X4,X3))|m1_subset_1(k4_yellow_6(X3,X4),u1_struct_0(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_yellow_6])])])).
fof(c_0_12, plain, (![X3]:((l1_waybel_0(esk3_1(X3),X3)|~l1_struct_0(X3))&(v6_waybel_0(esk3_1(X3),X3)|~l1_struct_0(X3)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rc4_waybel_0])])])])])])).
fof(c_0_13, plain, (![X3]:![X4]:(~l1_struct_0(X3)|(~l1_waybel_0(X4,X3)|l1_orders_2(X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])])])).
fof(c_0_14, negated_conjecture, (((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))&~r1_waybel_0(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])).
fof(c_0_15, plain, (![X6]:![X7]:![X8]:![X10]:![X8]:![X11]:(((((m1_subset_1(esk4_3(X6,X7,X8),u1_struct_0(X7))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~m1_subset_1(X10,u1_struct_0(X7))|(~r1_orders_2(X7,esk4_3(X6,X7,X8),X10)|r2_hidden(k2_waybel_0(X6,X7,X10),X8)))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))&(((((m1_subset_1(esk5_4(X6,X7,X8,X11),u1_struct_0(X7))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X11,esk5_4(X6,X7,X8,X11))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~r2_hidden(k2_waybel_0(X6,X7,esk5_4(X6,X7,X8,X11)),X8)|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_waybel_0])])])])])])])])).
cnf(c_0_16,plain,(m1_subset_1(k4_yellow_6(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(l1_waybel_0(esk3_1(X1),X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, (![X2]:(~l1_orders_2(X2)|l1_struct_0(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_19,plain,(l1_orders_2(X1)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,negated_conjecture,(l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_22, plain, (![X5]:![X6]:![X7]:![X8]:(v1_xboole_0(X5)|(v1_xboole_0(X6)|(~m1_subset_1(X7,X5)|((~v1_funct_1(X8)|(~v1_funct_2(X8,X5,X6)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))))|r2_hidden(k3_funct_2(X5,X6,X8,X7),k2_relset_1(X5,X6,X8))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t189_funct_2])])])])])])).
fof(c_0_23, plain, (![X3]:![X4]:((v1_funct_1(u1_waybel_0(X3,X4))|(~l1_struct_0(X3)|~l1_waybel_0(X4,X3)))&((v1_funct_2(u1_waybel_0(X3,X4),u1_struct_0(X4),u1_struct_0(X3))|(~l1_struct_0(X3)|~l1_waybel_0(X4,X3)))&(m1_subset_1(u1_waybel_0(X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|(~l1_struct_0(X3)|~l1_waybel_0(X4,X3)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_waybel_0])])])).
cnf(c_0_24,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(m1_subset_1(k4_yellow_6(X1,esk3_1(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_26,plain,(l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,negated_conjecture,(l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_28,plain,(r2_hidden(k3_funct_2(X1,X2,X3,X4),k2_relset_1(X1,X2,X3))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(m1_subset_1(u1_waybel_0(X2,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(v1_funct_1(u1_waybel_0(X2,X1))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,plain,(v1_funct_2(u1_waybel_0(X2,X1),u1_struct_0(X1),u1_struct_0(X2))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(m1_subset_1(esk5_4(X1,X2,X3,k4_yellow_6(X2,esk3_1(X2))),u1_struct_0(X2))|r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,negated_conjecture,(l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_36, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|~l1_struct_0(X4))|((v2_struct_0(X5)|~l1_waybel_0(X5,X4))|(~m1_subset_1(X6,u1_struct_0(X5))|k2_waybel_0(X4,X5,X6)=k3_funct_2(u1_struct_0(X5),u1_struct_0(X4),u1_waybel_0(X4,X5),X6))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_waybel_0])])])])])])).
cnf(c_0_37,plain,(v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X2))|r2_hidden(k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),u1_waybel_0(X2,X1),X3),k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),u1_waybel_0(X2,X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31])).
cnf(c_0_38,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk2_0,X1,k4_yellow_6(esk2_0,esk3_1(esk2_0))),u1_struct_0(esk2_0))|r1_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_20]), c_0_21]), c_0_33])]), c_0_34]), c_0_35])).
cnf(c_0_39,plain,(k2_waybel_0(X1,X2,X3)=k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_40,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(X1))|r2_hidden(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),u1_waybel_0(X1,esk2_0),esk5_4(esk1_0,esk2_0,X2,k4_yellow_6(esk2_0,esk3_1(esk2_0)))),k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(X1),u1_waybel_0(X1,esk2_0)))|r1_waybel_0(esk1_0,esk2_0,X2)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_41,negated_conjecture,(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),u1_waybel_0(X1,esk2_0),esk5_4(esk1_0,esk2_0,X2,k4_yellow_6(esk2_0,esk3_1(esk2_0))))=k2_waybel_0(X1,esk2_0,esk5_4(esk1_0,esk2_0,X2,k4_yellow_6(esk2_0,esk3_1(esk2_0))))|r1_waybel_0(esk1_0,esk2_0,X2)|v2_struct_0(X1)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_38]), c_0_35])).
cnf(c_0_42,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(k2_waybel_0(X1,X2,esk5_4(X1,X2,X3,X4)),X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_43,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0),esk5_4(esk1_0,esk2_0,X1,k4_yellow_6(esk2_0,esk3_1(esk2_0)))),k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|r1_waybel_0(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_20]), c_0_21])])).
cnf(c_0_44,negated_conjecture,(k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0),esk5_4(esk1_0,esk2_0,X1,k4_yellow_6(esk2_0,esk3_1(esk2_0))))=k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,X1,k4_yellow_6(esk2_0,esk3_1(esk2_0))))|r1_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_20]), c_0_21])]), c_0_34])).
fof(c_0_45, plain, (![X2]:((v2_struct_0(X2)|~l1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_46,plain,(r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X1,X2,esk5_4(X1,X2,X3,k4_yellow_6(X2,esk3_1(X2)))),X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_42, c_0_25])).
cnf(c_0_47,negated_conjecture,(v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,X1,k4_yellow_6(esk2_0,esk3_1(esk2_0)))),k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))|r1_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48,negated_conjecture,(~r1_waybel_0(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49,plain,(v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_50,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_20]), c_0_21]), c_0_33])]), c_0_48]), c_0_34]), c_0_35])).
cnf(c_0_51,negated_conjecture,(v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_33])]), c_0_35])).
cnf(c_0_52,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_51]), c_0_21])]), c_0_34]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 32
# Proof object formula steps           : 21
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 101
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 101
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 2
# Generated clauses                    : 78
# ...of the previous two non-trivial   : 77
# Contextual simplify-reflections      : 4
# Paramodulations                      : 78
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 76
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 66
# Current number of unprocessed clauses: 18
# ...number of literals in the above   : 135
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 2414
# Rec. Clause-clause subsumption calls : 95
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7229

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 45364 pages
