# No SInE strategy applied
# Auto-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_0, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:![X4]:(r1_tarski(X3,X4)=>((r1_waybel_0(X1,X2,X3)=>r1_waybel_0(X1,X2,X4))&(r2_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X4))))))), file('out/t8_waybel_0__3_premises__rerun.E_input', t8_waybel_0)).
fof(d12_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))))), file('out/t8_waybel_0__3_premises__rerun.E_input', d12_waybel_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t8_waybel_0__3_premises__rerun.E_input', d3_tarski)).
fof(d11_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))))), file('out/t8_waybel_0__3_premises__rerun.E_input', d11_waybel_0)).
fof(c_0_4, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:![X4]:(r1_tarski(X3,X4)=>((r1_waybel_0(X1,X2,X3)=>r1_waybel_0(X1,X2,X4))&(r2_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X4)))))))), inference(assume_negation,[status(cth)],[t8_waybel_0])).
fof(c_0_5, plain, (![X6]:![X7]:![X8]:![X9]:![X8]:![X12]:((((((m1_subset_1(esk5_4(X6,X7,X8,X9),u1_struct_0(X7))|~m1_subset_1(X9,u1_struct_0(X7)))|~r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X9,esk5_4(X6,X7,X8,X9))|~m1_subset_1(X9,u1_struct_0(X7)))|~r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((r2_hidden(k2_waybel_0(X6,X7,esk5_4(X6,X7,X8,X9)),X8)|~m1_subset_1(X9,u1_struct_0(X7)))|~r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))))&((((m1_subset_1(esk6_3(X6,X7,X8),u1_struct_0(X7))|r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~m1_subset_1(X12,u1_struct_0(X7))|(~r1_orders_2(X7,esk6_3(X6,X7,X8),X12)|~r2_hidden(k2_waybel_0(X6,X7,X12),X8)))|r2_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])])).
fof(c_0_6, negated_conjecture, (((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))&(r1_tarski(esk3_0,esk4_0)&(((r2_waybel_0(esk1_0,esk2_0,esk3_0)|r1_waybel_0(esk1_0,esk2_0,esk3_0))&(~r2_waybel_0(esk1_0,esk2_0,esk4_0)|r1_waybel_0(esk1_0,esk2_0,esk3_0)))&((r2_waybel_0(esk1_0,esk2_0,esk3_0)|~r1_waybel_0(esk1_0,esk2_0,esk4_0))&(~r2_waybel_0(esk1_0,esk2_0,esk4_0)|~r1_waybel_0(esk1_0,esk2_0,esk4_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])])])])).
cnf(c_0_7,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_waybel_0(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,negated_conjecture,(l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13,negated_conjecture,(m1_subset_1(esk6_3(esk1_0,esk2_0,X1),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])]), c_0_10]), c_0_11])).
cnf(c_0_14,negated_conjecture,(m1_subset_1(esk5_4(X1,esk2_0,X2,esk6_3(esk1_0,esk2_0,X3)),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X3)|v2_struct_0(X1)|~r2_waybel_0(X1,esk2_0,X2)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_10])).
cnf(c_0_15,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(k2_waybel_0(X1,X2,esk5_4(X1,X2,X3,X4)),X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk2_0,X1,esk6_3(esk1_0,esk2_0,X2)),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X2)|~r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_17,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,esk3_0)|r2_waybel_0(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk7_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk7_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_19,negated_conjecture,(r2_hidden(k2_waybel_0(X1,esk2_0,esk5_4(X1,esk2_0,X2,esk6_3(esk1_0,esk2_0,X3))),X2)|r2_waybel_0(esk1_0,esk2_0,X3)|v2_struct_0(X1)|~r2_waybel_0(X1,esk2_0,X2)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_13]), c_0_10])).
cnf(c_0_20,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_hidden(k2_waybel_0(X1,X2,X4),X3)|~r1_orders_2(X2,esk6_3(X1,X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_21,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X1)|r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23,negated_conjecture,(r1_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,X1,esk6_3(esk1_0,esk2_0,X2))),X1)|r2_waybel_0(esk1_0,esk2_0,X2)|~r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_25,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_orders_2(X2,X4,esk5_4(X1,X2,X3,X4))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_26,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(X2,esk2_0,X3)|r1_waybel_0(esk1_0,esk2_0,esk3_0)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X2,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),X3)|~r1_orders_2(esk2_0,esk6_3(X2,esk2_0,X3),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))|~l1_waybel_0(esk2_0,X2)|~l1_struct_0(X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_10])).
cnf(c_0_27,negated_conjecture,(r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),esk3_0)|r2_waybel_0(esk1_0,esk2_0,X1)|r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_17])).
cnf(c_0_29,negated_conjecture,(r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,X1),esk5_4(X2,esk2_0,X3,esk6_3(esk1_0,esk2_0,X1)))|r2_waybel_0(esk1_0,esk2_0,X1)|v2_struct_0(X2)|~r2_waybel_0(X2,esk2_0,X3)|~l1_waybel_0(esk2_0,X2)|~l1_struct_0(X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_13]), c_0_10])).
cnf(c_0_30,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(esk1_0,esk2_0,X2)|r1_waybel_0(esk1_0,esk2_0,esk3_0)|~r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),X2)|~r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,X2),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_31,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),esk4_0)|r2_waybel_0(esk1_0,esk2_0,X1)|r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,negated_conjecture,(r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,X1),esk5_4(esk1_0,esk2_0,X2,esk6_3(esk1_0,esk2_0,X1)))|r2_waybel_0(esk1_0,esk2_0,X1)|~r2_waybel_0(esk1_0,esk2_0,X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_33,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,esk4_0)|r2_waybel_0(esk1_0,esk2_0,X1)|r1_waybel_0(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,esk4_0),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,negated_conjecture,(r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,X1),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))|r2_waybel_0(esk1_0,esk2_0,X1)|r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_17])).
fof(c_0_35, plain, (![X6]:![X7]:![X8]:![X10]:![X8]:![X11]:(((((m1_subset_1(esk8_3(X6,X7,X8),u1_struct_0(X7))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~m1_subset_1(X10,u1_struct_0(X7))|(~r1_orders_2(X7,esk8_3(X6,X7,X8),X10)|r2_hidden(k2_waybel_0(X6,X7,X10),X8)))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))&(((((m1_subset_1(esk9_4(X6,X7,X8,X11),u1_struct_0(X7))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X11,esk9_4(X6,X7,X8,X11))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~r2_hidden(k2_waybel_0(X6,X7,esk9_4(X6,X7,X8,X11)),X8)|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_waybel_0])])])])])])])])).
cnf(c_0_36,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,esk3_0)|~r2_waybel_0(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_37,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,esk4_0)|r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_39,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|m1_subset_1(esk9_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41,negated_conjecture,(m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_8]), c_0_9])]), c_0_10]), c_0_11])).
cnf(c_0_42,negated_conjecture,(m1_subset_1(esk9_4(X1,esk2_0,X2,esk8_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))|r1_waybel_0(X1,esk2_0,X2)|v2_struct_0(X1)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_10])).
cnf(c_0_43,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(k2_waybel_0(X1,X2,X4),X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)|~r1_orders_2(X2,esk8_3(X1,X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44,negated_conjecture,(m1_subset_1(esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))|r1_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_45,negated_conjecture,(r2_hidden(k2_waybel_0(X1,esk2_0,esk9_4(esk1_0,esk2_0,X2,esk8_3(esk1_0,esk2_0,esk3_0))),X3)|r1_waybel_0(esk1_0,esk2_0,X2)|v2_struct_0(X1)|~r1_orders_2(esk2_0,esk8_3(X1,esk2_0,X3),esk9_4(esk1_0,esk2_0,X2,esk8_3(esk1_0,esk2_0,esk3_0)))|~r1_waybel_0(X1,esk2_0,X3)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_10])).
cnf(c_0_46,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|r1_orders_2(X2,X4,esk9_4(X1,X2,X3,X4))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0))),X2)|r1_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk8_3(esk1_0,esk2_0,X2),esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0)))|~r1_waybel_0(esk1_0,esk2_0,X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_48,negated_conjecture,(r1_orders_2(esk2_0,esk8_3(esk1_0,esk2_0,esk3_0),esk9_4(X1,esk2_0,X2,esk8_3(esk1_0,esk2_0,esk3_0)))|r1_waybel_0(X1,esk2_0,X2)|v2_struct_0(X1)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_41]), c_0_10])).
cnf(c_0_49,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(k2_waybel_0(X1,X2,esk9_4(X1,X2,X3,X4)),X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_50,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0))),esk3_0)|r1_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk8_3(esk1_0,esk2_0,esk3_0),esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_47, c_0_39])).
cnf(c_0_51,negated_conjecture,(r1_orders_2(esk2_0,esk8_3(esk1_0,esk2_0,esk3_0),esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0)))|r1_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_52,negated_conjecture,(r1_waybel_0(X1,esk2_0,X2)|v2_struct_0(X1)|~r2_hidden(k2_waybel_0(X1,esk2_0,esk9_4(X1,esk2_0,X2,esk8_3(esk1_0,esk2_0,esk3_0))),X2)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_41]), c_0_10])).
cnf(c_0_53,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0))),esk3_0)|r1_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_54,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,X1)|~r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0))),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_55,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk9_4(esk1_0,esk2_0,X1,esk8_3(esk1_0,esk2_0,esk3_0))),esk4_0)|r1_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_27, c_0_53])).
cnf(c_0_56,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,esk3_0)|~r1_waybel_0(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_57,negated_conjecture,(r1_waybel_0(esk1_0,esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_58,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57])])).
cnf(c_0_59,negated_conjecture,(m1_subset_1(esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_58])).
cnf(c_0_60,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(X2,esk2_0,X3)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X2,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),X3)|~r1_orders_2(esk2_0,esk6_3(X2,esk2_0,X3),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))|~l1_waybel_0(esk2_0,X2)|~l1_struct_0(X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_59]), c_0_10])).
cnf(c_0_61,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),esk3_0)|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_24, c_0_58])).
cnf(c_0_62,negated_conjecture,(~r1_waybel_0(esk1_0,esk2_0,esk4_0)|~r2_waybel_0(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_63,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(esk1_0,esk2_0,X2)|~r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),X2)|~r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,X2),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_8]), c_0_9])]), c_0_11])).
cnf(c_0_64,negated_conjecture,(r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1))),esk4_0)|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_27, c_0_61])).
cnf(c_0_65,negated_conjecture,(~r2_waybel_0(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_57])])).
cnf(c_0_66,negated_conjecture,(r2_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,esk4_0),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_67,negated_conjecture,(r1_orders_2(esk2_0,esk6_3(esk1_0,esk2_0,X1),esk5_4(esk1_0,esk2_0,esk3_0,esk6_3(esk1_0,esk2_0,X1)))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_32, c_0_58])).
cnf(c_0_68,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 60
# Proof object formula steps           : 9
# Proof object conjectures             : 52
# Proof object clause conjectures      : 49
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 328
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 328
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 43
# Generated clauses                    : 588
# ...of the previous two non-trivial   : 579
# Contextual simplify-reflections      : 0
# Paramodulations                      : 588
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 283
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 251
# Current number of unprocessed clauses: 202
# ...number of literals in the above   : 1311
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 3312
# Rec. Clause-clause subsumption calls : 640
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 346
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 205
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 32378

# -------------------------------------------------
# User time                : 0.040 s
# System time              : 0.000 s
# Total time               : 0.040 s
# Maximum resident set size: 45316 pages
