# No SInE strategy applied
# Auto-Mode selected heuristic H_____011_C18_F1_PI_SE_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_wellord1, conjecture, (![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r2_hidden(X2,k1_relat_1(X1))&(?[X3]:(r2_hidden(X3,k1_relat_1(X1))&~(r2_hidden(k4_tarski(X3,X2),X1)))&![X3]:~((r2_hidden(X3,k1_relat_1(X1))&(r2_hidden(k4_tarski(X2,X3),X1)&![X4]:~((r2_hidden(X4,k1_relat_1(X1))&(r2_hidden(k4_tarski(X2,X4),X1)&(~(X4=X2)&~(r2_hidden(k4_tarski(X3,X4),X1))))))))))))))), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', t8_wellord1)).
fof(d5_tarski, axiom, (![X1]:![X2]:k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', d5_tarski)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(s1_xboole_0__e5_15__wellord1, axiom, (![X1]:![X2]:(v1_relat_1(X1)=>?[X3]:![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,k1_relat_1(X1))&~(r2_hidden(k4_tarski(X4,X2),X1)))))), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', s1_xboole_0__e5_15__wellord1)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', d3_tarski)).
fof(t6_wellord1, axiom, (![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r1_tarski(X2,k1_relat_1(X1))&(~(X2=k1_xboole_0)&![X3]:~((r2_hidden(X3,X2)&![X4]:(r2_hidden(X4,X2)=>r2_hidden(k4_tarski(X3,X4),X1)))))))))), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', t6_wellord1)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', t3_subset)).
fof(t5_subset, axiom, (![X1]:![X2]:![X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3))))), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', t5_subset)).
fof(dt_o_0_0_xboole_0, axiom, (v1_xboole_0(o_0_0_xboole_0)), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', dt_o_0_0_xboole_0)).
fof(d2_xboole_0, axiom, (k1_xboole_0=o_0_0_xboole_0), file('out/t8_wellord1__73_premises__rerun__rerun.E_input', d2_xboole_0)).
fof(c_0_10, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r2_hidden(X2,k1_relat_1(X1))&(?[X3]:(r2_hidden(X3,k1_relat_1(X1))&~(r2_hidden(k4_tarski(X3,X2),X1)))&![X3]:~((r2_hidden(X3,k1_relat_1(X1))&(r2_hidden(k4_tarski(X2,X3),X1)&![X4]:~((r2_hidden(X4,k1_relat_1(X1))&(r2_hidden(k4_tarski(X2,X4),X1)&(~(X4=X2)&~(r2_hidden(k4_tarski(X3,X4),X1)))))))))))))))), inference(assume_negation,[status(cth)],[t8_wellord1])).
fof(c_0_11, negated_conjecture, (![X8]:(v1_relat_1(esk1_0)&(v2_wellord1(esk1_0)&(r2_hidden(esk2_0,k1_relat_1(esk1_0))&((r2_hidden(esk3_0,k1_relat_1(esk1_0))&~r2_hidden(k4_tarski(esk3_0,esk2_0),esk1_0))&(((r2_hidden(esk4_1(X8),k1_relat_1(esk1_0))|~r2_hidden(k4_tarski(esk2_0,X8),esk1_0))|~r2_hidden(X8,k1_relat_1(esk1_0)))&(((r2_hidden(k4_tarski(esk2_0,esk4_1(X8)),esk1_0)|~r2_hidden(k4_tarski(esk2_0,X8),esk1_0))|~r2_hidden(X8,k1_relat_1(esk1_0)))&(((esk4_1(X8)!=esk2_0|~r2_hidden(k4_tarski(esk2_0,X8),esk1_0))|~r2_hidden(X8,k1_relat_1(esk1_0)))&((~r2_hidden(k4_tarski(X8,esk4_1(X8)),esk1_0)|~r2_hidden(k4_tarski(esk2_0,X8),esk1_0))|~r2_hidden(X8,k1_relat_1(esk1_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])])])).
fof(c_0_12, plain, (![X3]:![X4]:k4_tarski(X3,X4)=k2_tarski(k2_tarski(X3,X4),k1_tarski(X3))), inference(variable_rename,[status(thm)],[d5_tarski])).
cnf(c_0_13,negated_conjecture,(~r2_hidden(X1,k1_relat_1(esk1_0))|~r2_hidden(k4_tarski(esk2_0,X1),esk1_0)|~r2_hidden(k4_tarski(X1,esk4_1(X1)),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14,plain,(k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_16, plain, (![X5]:![X6]:![X8]:![X8]:((((r2_hidden(X8,k1_relat_1(X5))|~r2_hidden(X8,esk5_2(X5,X6)))|~v1_relat_1(X5))&((~r2_hidden(k4_tarski(X8,X6),X5)|~r2_hidden(X8,esk5_2(X5,X6)))|~v1_relat_1(X5)))&(((~r2_hidden(X8,k1_relat_1(X5))|r2_hidden(k4_tarski(X8,X6),X5))|r2_hidden(X8,esk5_2(X5,X6)))|~v1_relat_1(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s1_xboole_0__e5_15__wellord1])])])])])])])])).
cnf(c_0_17,negated_conjecture,(r2_hidden(k4_tarski(esk2_0,esk4_1(X1)),esk1_0)|~r2_hidden(X1,k1_relat_1(esk1_0))|~r2_hidden(k4_tarski(esk2_0,X1),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_18, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk6_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk6_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
cnf(c_0_19,negated_conjecture,(~r2_hidden(X1,k1_relat_1(esk1_0))|~r2_hidden(k2_tarski(k2_tarski(esk2_0,X1),k1_tarski(esk2_0)),esk1_0)|~r2_hidden(k2_tarski(k2_tarski(X1,esk4_1(X1)),k1_tarski(X1)),esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_14])).
cnf(c_0_20,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(r2_hidden(X2,esk5_2(X1,X3))|r2_hidden(k4_tarski(X2,X3),X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,negated_conjecture,(r2_hidden(k2_tarski(k2_tarski(esk2_0,esk4_1(X1)),k1_tarski(esk2_0)),esk1_0)|~r2_hidden(X1,k1_relat_1(esk1_0))|~r2_hidden(k2_tarski(k2_tarski(esk2_0,X1),k1_tarski(esk2_0)),esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_14]), c_0_14])).
cnf(c_0_23,plain,(r2_hidden(X2,k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,esk5_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,plain,(r1_tarski(X1,X2)|r2_hidden(esk6_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, (![X5]:![X6]:![X8]:(((((r2_hidden(esk7_2(X5,X6),X6)|X6=k1_xboole_0)|~r1_tarski(X6,k1_relat_1(X5)))|~v2_wellord1(X5))|~v1_relat_1(X5))&(((((~r2_hidden(X8,X6)|r2_hidden(k4_tarski(esk7_2(X5,X6),X8),X5))|X6=k1_xboole_0)|~r1_tarski(X6,k1_relat_1(X5)))|~v2_wellord1(X5))|~v1_relat_1(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_wellord1])])])])])])])])).
cnf(c_0_26,negated_conjecture,(~r2_hidden(k2_tarski(k2_tarski(esk2_0,X1),k1_tarski(esk2_0)),esk1_0)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk4_1(X1))),esk1_0)|~r2_hidden(X1,k1_relat_1(esk1_0))), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27,negated_conjecture,(r2_hidden(esk2_0,k1_relat_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28,plain,(r2_hidden(X2,esk5_2(X1,X3))|r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(rw,[status(thm)],[c_0_21, c_0_14])).
cnf(c_0_29,negated_conjecture,(v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_30,negated_conjecture,(r2_hidden(k2_tarski(k1_tarski(esk2_0),k2_tarski(esk2_0,esk4_1(X1))),esk1_0)|~r2_hidden(k2_tarski(k2_tarski(esk2_0,X1),k1_tarski(esk2_0)),esk1_0)|~r2_hidden(X1,k1_relat_1(esk1_0))), inference(rw,[status(thm)],[c_0_22, c_0_20])).
fof(c_0_31, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_32,plain,(r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_33,negated_conjecture,(~r2_hidden(k4_tarski(esk3_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_34,negated_conjecture,(r2_hidden(esk3_0,k1_relat_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35,plain,(r1_tarski(esk5_2(X1,X2),X3)|r2_hidden(esk6_2(esk5_2(X1,X2),X3),k1_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_36,plain,(X2=k1_xboole_0|r2_hidden(k4_tarski(esk7_2(X1,X2),X3),X1)|~v1_relat_1(X1)|~v2_wellord1(X1)|~r1_tarski(X2,k1_relat_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37,negated_conjecture,(~r2_hidden(k2_tarski(k1_tarski(esk2_0),k2_tarski(esk2_0,esk4_1(esk2_0))),esk1_0)|~r2_hidden(k2_tarski(k2_tarski(esk2_0,esk2_0),k1_tarski(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_38,negated_conjecture,(r2_hidden(k2_tarski(k2_tarski(esk2_0,X1),k1_tarski(esk2_0)),esk1_0)|r2_hidden(esk2_0,esk5_2(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_27]), c_0_29])])).
cnf(c_0_39,negated_conjecture,(r2_hidden(k2_tarski(k1_tarski(esk2_0),k2_tarski(esk2_0,esk4_1(esk2_0))),esk1_0)|~r2_hidden(k2_tarski(k2_tarski(esk2_0,esk2_0),k1_tarski(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_27])).
fof(c_0_40, plain, (![X4]:![X5]:![X6]:(~r2_hidden(X4,X5)|(~m1_subset_1(X5,k1_zfmisc_1(X6))|~v1_xboole_0(X6)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])])])).
cnf(c_0_41,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42,plain,(r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_32, c_0_24])).
cnf(c_0_43,negated_conjecture,(~r2_hidden(k2_tarski(k2_tarski(esk3_0,esk2_0),k1_tarski(esk3_0)),esk1_0)), inference(rw,[status(thm)],[c_0_33, c_0_14])).
cnf(c_0_44,negated_conjecture,(r2_hidden(k2_tarski(k2_tarski(esk3_0,X1),k1_tarski(esk3_0)),esk1_0)|r2_hidden(esk3_0,esk5_2(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_34]), c_0_29])])).
cnf(c_0_45,plain,(~v1_relat_1(X1)|~r2_hidden(X2,esk5_2(X1,X3))|~r2_hidden(k4_tarski(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_46,plain,(X2=k1_xboole_0|r2_hidden(esk7_2(X1,X2),X2)|~v1_relat_1(X1)|~v2_wellord1(X1)|~r1_tarski(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_47,plain,(r1_tarski(esk5_2(X1,X2),k1_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_35])).
cnf(c_0_48,plain,(X2=k1_xboole_0|r2_hidden(k2_tarski(k2_tarski(esk7_2(X1,X2),X3),k1_tarski(esk7_2(X1,X2))),X1)|~v1_relat_1(X1)|~v2_wellord1(X1)|~r2_hidden(X3,X2)|~r1_tarski(X2,k1_relat_1(X1))), inference(rw,[status(thm)],[c_0_36, c_0_14])).
cnf(c_0_49,negated_conjecture,(r2_hidden(esk2_0,esk5_2(esk1_0,esk2_0))|~r2_hidden(k2_tarski(k1_tarski(esk2_0),k2_tarski(esk2_0,esk4_1(esk2_0))),esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_50,negated_conjecture,(r2_hidden(k2_tarski(k1_tarski(esk2_0),k2_tarski(esk2_0,esk4_1(esk2_0))),esk1_0)|r2_hidden(esk2_0,esk5_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_39, c_0_38])).
cnf(c_0_51,plain,(~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52,plain,(m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_53,negated_conjecture,(~r2_hidden(k2_tarski(k2_tarski(esk2_0,esk3_0),k1_tarski(esk3_0)),esk1_0)), inference(rw,[status(thm)],[c_0_43, c_0_20])).
cnf(c_0_54,negated_conjecture,(r2_hidden(k2_tarski(k2_tarski(X1,esk3_0),k1_tarski(esk3_0)),esk1_0)|r2_hidden(esk3_0,esk5_2(esk1_0,X1))), inference(spm,[status(thm)],[c_0_44, c_0_20])).
cnf(c_0_55,plain,(~v1_relat_1(X1)|~r2_hidden(X2,esk5_2(X1,X3))|~r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),X1)), inference(rw,[status(thm)],[c_0_45, c_0_14])).
cnf(c_0_56,plain,(esk5_2(X1,X2)=k1_xboole_0|r2_hidden(esk7_2(X1,esk5_2(X1,X2)),esk5_2(X1,X2))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_57,plain,(esk5_2(X1,X2)=k1_xboole_0|r2_hidden(k2_tarski(k2_tarski(esk7_2(X1,esk5_2(X1,X2)),X3),k1_tarski(esk7_2(X1,esk5_2(X1,X2)))),X1)|~r2_hidden(X3,esk5_2(X1,X2))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_48, c_0_47])).
cnf(c_0_58,negated_conjecture,(r2_hidden(esk2_0,esk5_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59,negated_conjecture,(v2_wellord1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_60,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61,negated_conjecture,(r2_hidden(esk3_0,esk5_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62,plain,(esk5_2(X1,X2)=k1_xboole_0|~r2_hidden(k2_tarski(k2_tarski(X2,esk7_2(X1,esk5_2(X1,X2))),k1_tarski(esk7_2(X1,esk5_2(X1,X2)))),X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_20])).
cnf(c_0_63,negated_conjecture,(esk5_2(esk1_0,esk2_0)=k1_xboole_0|r2_hidden(k2_tarski(k2_tarski(esk2_0,esk7_2(esk1_0,esk5_2(esk1_0,esk2_0))),k1_tarski(esk7_2(esk1_0,esk5_2(esk1_0,esk2_0)))),esk1_0)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_29])]), c_0_20])).
cnf(c_0_64,plain,(v1_xboole_0(o_0_0_xboole_0)), inference(split_conjunct,[status(thm)],[dt_o_0_0_xboole_0])).
cnf(c_0_65,plain,(k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
cnf(c_0_66,negated_conjecture,(~v1_xboole_0(esk5_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_67,negated_conjecture,(esk5_2(esk1_0,esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_59]), c_0_29])])).
cnf(c_0_68,plain,(v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_69,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67]), c_0_68])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 70
# Proof object clause steps            : 51
# Proof object formula steps           : 19
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 23
# Processed clauses                    : 379
# ...of these trivial                  : 1
# ...subsumed                          : 186
# ...remaining for further processing  : 192
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 13
# Generated clauses                    : 574
# ...of the previous two non-trivial   : 529
# Contextual simplify-reflections      : 0
# Paramodulations                      : 574
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 164
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 149
# Current number of unprocessed clauses: 142
# ...number of literals in the above   : 720
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 3349
# Rec. Clause-clause subsumption calls : 1363
# Non-unit clause-clause subsumptions  : 183
# Unit Clause-clause subsumption calls : 128
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17345

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.004 s
# Total time               : 0.032 s
# Maximum resident set size: 45368 pages
