# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(r2_hidden(X4,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k5_waybel_9(X1,X2,X3))))<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X3,X5)&X4=k2_waybel_0(X1,X2,X5)))))))), file('out/t8_yellow19__5_premises__rerun.E_input', t7_yellow19)).
fof(d11_waybel_0, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))))), file('out/t8_yellow19__5_premises__rerun.E_input', d11_waybel_0)).
fof(d2_yellow19, axiom, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(m1_yellow19(X3,X1,X2)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&X3=k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k4_waybel_9(X1,X2,X4))))))))), file('out/t8_yellow19__5_premises__rerun.E_input', d2_yellow19)).
fof(dt_m1_yellow19, axiom, (![X1]:![X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&l1_waybel_0(X2,X1)))=>![X3]:(m1_yellow19(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('out/t8_yellow19__5_premises__rerun.E_input', dt_m1_yellow19)).
fof(redefinition_k5_waybel_9, axiom, (![X1]:![X2]:![X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))&m1_subset_1(X3,u1_struct_0(X2))))=>k5_waybel_9(X1,X2,X3)=k4_waybel_9(X1,X2,X3))), file('out/t8_yellow19__5_premises__rerun.E_input', redefinition_k5_waybel_9)).
fof(t8_yellow19, conjecture, (![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_yellow19(X3,X1,X2)=>r1_waybel_0(X1,X2,X3))))), file('out/t8_yellow19__5_premises__rerun.E_input', t8_yellow19)).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X9]:![X9]:![X11]:((((((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X7))|~r2_hidden(X9,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,k5_waybel_9(X6,X7,X8)))))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X8,esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,k5_waybel_9(X6,X7,X8)))))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((X9=k2_waybel_0(X6,X7,esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,k5_waybel_9(X6,X7,X8)))))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(v2_struct_0(X6)|~l1_struct_0(X6)))))&(((((~m1_subset_1(X11,u1_struct_0(X7))|(~r1_orders_2(X7,X8,X11)|X9!=k2_waybel_0(X6,X7,X11)))|r2_hidden(X9,k2_relset_1(u1_struct_0(X6),u1_waybel_0(X6,k5_waybel_9(X6,X7,X8)))))|~m1_subset_1(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6)))))|(v2_struct_0(X6)|~l1_struct_0(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_yellow19])])])])])])])])).
fof(c_0_7, plain, (![X6]:![X7]:![X8]:![X10]:![X8]:![X11]:(((((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X7))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~m1_subset_1(X10,u1_struct_0(X7))|(~r1_orders_2(X7,esk2_3(X6,X7,X8),X10)|r2_hidden(k2_waybel_0(X6,X7,X10),X8)))|~r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))&(((((m1_subset_1(esk3_4(X6,X7,X8,X11),u1_struct_0(X7))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(((((r1_orders_2(X7,X11,esk3_4(X6,X7,X8,X11))|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6)))&((((~r2_hidden(k2_waybel_0(X6,X7,esk3_4(X6,X7,X8,X11)),X8)|~m1_subset_1(X11,u1_struct_0(X7)))|r1_waybel_0(X6,X7,X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6)))|(v2_struct_0(X6)|~l1_struct_0(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_waybel_0])])])])])])])])).
cnf(c_0_8,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X4,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k5_waybel_9(X1,X2,X3))))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X2))|X4!=k2_waybel_0(X1,X2,X5)|~r1_orders_2(X2,X3,X5)|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X9]:((((((m1_subset_1(esk4_3(X5,X6,X7),u1_struct_0(X6))|~m1_yellow19(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X6)|~l1_waybel_0(X6,X5)))|(v2_struct_0(X5)|~l1_struct_0(X5)))&((((X7=k2_relset_1(u1_struct_0(X5),u1_waybel_0(X5,k4_waybel_9(X5,X6,esk4_3(X5,X6,X7))))|~m1_yellow19(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X6)|~l1_waybel_0(X6,X5)))|(v2_struct_0(X5)|~l1_struct_0(X5))))&(((((~m1_subset_1(X9,u1_struct_0(X6))|X7!=k2_relset_1(u1_struct_0(X5),u1_waybel_0(X5,k4_waybel_9(X5,X6,X9))))|m1_yellow19(X7,X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5))))|(v2_struct_0(X6)|~l1_waybel_0(X6,X5)))|(v2_struct_0(X5)|~l1_struct_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_yellow19])])])])])])])])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|~l1_struct_0(X4))|(v2_struct_0(X5)|~l1_waybel_0(X5,X4)))|(~m1_yellow19(X6,X4,X5)|m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X4)))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_yellow19])])])])])])).
cnf(c_0_11,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(k2_waybel_0(X1,X2,esk3_4(X1,X2,X3,X4)),X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(r2_hidden(k2_waybel_0(X1,X2,X3),k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k5_waybel_9(X1,X2,X4))))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_orders_2(X2,X4,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_struct_0(X1)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|m1_subset_1(esk3_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14,plain,(v2_struct_0(X1)|v2_struct_0(X2)|X3=k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k4_waybel_9(X1,X2,esk4_3(X1,X2,X3))))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_yellow19(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X3)|v2_struct_0(X2)|~m1_yellow19(X1,X2,X3)|~l1_waybel_0(X3,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(((v2_struct_0(X4)|~l1_struct_0(X4))|((v2_struct_0(X5)|(~v4_orders_2(X5)|(~v7_waybel_0(X5)|~l1_waybel_0(X5,X4))))|~m1_subset_1(X6,u1_struct_0(X5))))|k5_waybel_9(X4,X5,X6)=k4_waybel_9(X4,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k5_waybel_9])])])).
cnf(c_0_17,plain,(v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_yellow19(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,plain,(r1_waybel_0(X1,X2,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k5_waybel_9(X1,X2,X3))))|v2_struct_0(X1)|v2_struct_0(X2)|~r1_orders_2(X2,X3,esk3_4(X1,X2,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k5_waybel_9(X1,X2,X3))),X4))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
cnf(c_0_19,plain,(v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|r1_orders_2(X2,X4,esk3_4(X1,X2,X3,X4))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20,plain,(k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k4_waybel_9(X1,X2,esk4_3(X1,X2,X3))))=X3|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow19(X3,X1,X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(csr,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21,plain,(k5_waybel_9(X1,X2,X3)=k4_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow19(X3,X1,X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(csr,[status(thm)],[c_0_17, c_0_15])).
fof(c_0_23, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_yellow19(X3,X1,X2)=>r1_waybel_0(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t8_yellow19])).
cnf(c_0_24,plain,(r1_waybel_0(X1,X2,k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k5_waybel_9(X1,X2,X3))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,plain,(k2_relset_1(u1_struct_0(X1),u1_waybel_0(X1,k5_waybel_9(X1,X2,esk4_3(X1,X2,X3))))=X3|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow19(X3,X1,X2)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
fof(c_0_26, negated_conjecture, (((~v2_struct_0(esk5_0)&l1_struct_0(esk5_0))&((~v2_struct_0(esk6_0)&(v4_orders_2(esk6_0)&(v7_waybel_0(esk6_0)&l1_waybel_0(esk6_0,esk5_0))))&(m1_yellow19(esk7_0,esk5_0,esk6_0)&~r1_waybel_0(esk5_0,esk6_0,esk7_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])])])).
cnf(c_0_27,plain,(r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow19(X3,X1,X2)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_22])).
cnf(c_0_28,negated_conjecture,(m1_yellow19(esk7_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29,negated_conjecture,(l1_waybel_0(esk6_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30,negated_conjecture,(v7_waybel_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31,negated_conjecture,(v4_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32,negated_conjecture,(l1_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,negated_conjecture,(~r1_waybel_0(esk5_0,esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,negated_conjecture,(~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35,negated_conjecture,(~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 24
# Proof object formula steps           : 13
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 74
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 67
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 58
# ...of the previous two non-trivial   : 51
# Contextual simplify-reflections      : 26
# Paramodulations                      : 56
# Factorizations                       : 0
# Equation resolutions                 : 2
# Current number of processed clauses  : 42
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 304
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 1267
# Rec. Clause-clause subsumption calls : 35
# Non-unit clause-clause subsumptions  : 34
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5935

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 43280 pages
