# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', d5_xboole_0)).
fof(t91_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1))), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', t91_relat_1)).
fof(t4_boole, axiom, (![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', t4_boole)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', d3_tarski)).
fof(commutativity_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', commutativity_k3_xboole_0)).
fof(t48_xboole_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', t48_xboole_1)).
fof(t90_relat_1, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>k1_relat_1(k7_relat_1(X2,X1))=k3_xboole_0(k1_relat_1(X2),X1))), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', t90_relat_1)).
fof(t2_tarski, axiom, (![X1]:![X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2)), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', t2_tarski)).
fof(t3_boole, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('out/t91_relat_1__10_premises__rerun__rerun.E_input', t3_boole)).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk1_3(X5,X6,X7),X7)|(~r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X6)))|X7=k4_xboole_0(X5,X6))&(((r2_hidden(esk1_3(X5,X6,X7),X5)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6))&((~r2_hidden(esk1_3(X5,X6,X7),X6)|r2_hidden(esk1_3(X5,X6,X7),X7))|X7=k4_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])])).
fof(c_0_10, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1)))), inference(assume_negation,[status(cth)],[t91_relat_1])).
cnf(c_0_11,plain,(X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, plain, (![X2]:k4_xboole_0(k1_xboole_0,X2)=k1_xboole_0), inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk4_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk4_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_14, negated_conjecture, ((v1_relat_1(esk3_0)&(r1_tarski(esk2_0,k1_relat_1(esk3_0))&k1_relat_1(k7_relat_1(esk3_0,esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_15,plain,(~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(r1_tarski(X1,X2)|r2_hidden(esk4_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(r1_tarski(esk2_0,k1_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(r2_hidden(X4,X2)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21,plain,(~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|~r2_hidden(esk4_2(k4_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_15, c_0_17])).
cnf(c_0_23,negated_conjecture,(r2_hidden(X1,k1_relat_1(esk3_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(r1_tarski(k1_xboole_0,X1)|~r2_hidden(esk4_2(k1_xboole_0,X1),X2)), inference(spm,[status(thm)],[c_0_21, c_0_17])).
fof(c_0_26, plain, (![X3]:![X4]:k3_xboole_0(X3,X4)=k3_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_27, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_28, plain, (![X3]:![X4]:(~v1_relat_1(X4)|k1_relat_1(k7_relat_1(X4,X3))=k3_xboole_0(k1_relat_1(X4),X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t90_relat_1])])).
cnf(c_0_29,negated_conjecture,(r1_tarski(k4_xboole_0(X1,k1_relat_1(esk3_0)),X2)|~r2_hidden(esk4_2(k4_xboole_0(X1,k1_relat_1(esk3_0)),X2),esk2_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30,plain,(r1_tarski(k4_xboole_0(X1,X2),X3)|r2_hidden(esk4_2(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_24, c_0_17])).
cnf(c_0_31,plain,(r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_25, c_0_17])).
fof(c_0_32, plain, (![X4]:![X5]:(((~r2_hidden(esk5_2(X4,X5),X4)|~r2_hidden(esk5_2(X4,X5),X5))|X4=X5)&((r2_hidden(esk5_2(X4,X5),X4)|r2_hidden(esk5_2(X4,X5),X5))|X4=X5))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])])])).
cnf(c_0_33,plain,(k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35,plain,(k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36,negated_conjecture,(r1_tarski(k4_xboole_0(esk2_0,k1_relat_1(esk3_0)),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37,plain,(~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_31]), c_0_21])).
cnf(c_0_38,plain,(X1=X2|r2_hidden(esk5_2(X1,X2),X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_34])).
cnf(c_0_40,plain,(k1_relat_1(k7_relat_1(X1,X2))=k4_xboole_0(k1_relat_1(X1),k4_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_35, c_0_34])).
cnf(c_0_41,negated_conjecture,(r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(esk2_0,k1_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_18, c_0_36])).
cnf(c_0_42,plain,(k1_xboole_0=X1|r2_hidden(esk5_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43,negated_conjecture,(k1_relat_1(k7_relat_1(esk3_0,esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_44,plain,(k1_relat_1(k7_relat_1(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,k1_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_46,negated_conjecture,(k4_xboole_0(esk2_0,k1_relat_1(esk3_0))=k1_xboole_0|r2_hidden(esk5_2(k1_xboole_0,k4_xboole_0(esk2_0,k1_relat_1(esk3_0))),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_47, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_48,negated_conjecture,(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k1_relat_1(esk3_0)))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_49,negated_conjecture,(k4_xboole_0(esk2_0,k1_relat_1(esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_46])).
cnf(c_0_50,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_51,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49]), c_0_50])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 33
# Proof object formula steps           : 19
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 18
# Processed clauses                    : 145
# ...of these trivial                  : 5
# ...subsumed                          : 53
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 430
# ...of the previous two non-trivial   : 373
# Contextual simplify-reflections      : 21
# Paramodulations                      : 423
# Factorizations                       : 4
# Equation resolutions                 : 3
# Current number of processed clauses  : 61
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 239
# ...number of literals in the above   : 705
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 380
# Rec. Clause-clause subsumption calls : 323
# Non-unit clause-clause subsumptions  : 42
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6754

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 43040 pages
