# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, (![X1]:![X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), file('out/t91_relat_1__7_premises__rerun.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, (![X1]:![X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), file('out/t91_relat_1__7_premises__rerun.E_input', t70_enumset1)).
fof(t91_relat_1, conjecture, (![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1))), file('out/t91_relat_1__7_premises__rerun.E_input', t91_relat_1)).
fof(d4_xboole_0, axiom, (![X1]:![X2]:![X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2))))), file('out/t91_relat_1__7_premises__rerun.E_input', d4_xboole_0)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t91_relat_1__7_premises__rerun.E_input', t71_enumset1)).
fof(idempotence_k3_xboole_0, axiom, (![X1]:![X2]:k3_xboole_0(X1,X1)=X1), file('out/t91_relat_1__7_premises__rerun.E_input', idempotence_k3_xboole_0)).
fof(d3_tarski, axiom, (![X1]:![X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2)))), file('out/t91_relat_1__7_premises__rerun.E_input', d3_tarski)).
fof(t86_relat_1, axiom, (![X1]:![X2]:![X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k1_relat_1(k7_relat_1(X3,X2)))<=>(r2_hidden(X1,X2)&r2_hidden(X1,k1_relat_1(X3)))))), file('out/t91_relat_1__7_premises__rerun.E_input', t86_relat_1)).
fof(c_0_8, plain, (![X3]:![X4]:k1_setfam_1(k2_tarski(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_9, plain, (![X3]:![X4]:k1_enumset1(X3,X3,X4)=k2_tarski(X3,X4)), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_10, negated_conjecture, (~(![X1]:![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1)))), inference(assume_negation,[status(cth)],[t91_relat_1])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6))&((r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(X8,X5)|~r2_hidden(X8,X6))|r2_hidden(X8,X7))|X7!=k3_xboole_0(X5,X6)))&(((~r2_hidden(esk3_3(X5,X6,X7),X7)|(~r2_hidden(esk3_3(X5,X6,X7),X5)|~r2_hidden(esk3_3(X5,X6,X7),X6)))|X7=k3_xboole_0(X5,X6))&(((r2_hidden(esk3_3(X5,X6,X7),X5)|r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6))&((r2_hidden(esk3_3(X5,X6,X7),X6)|r2_hidden(esk3_3(X5,X6,X7),X7))|X7=k3_xboole_0(X5,X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])])).
cnf(c_0_12,plain,(k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_15, negated_conjecture, ((v1_relat_1(esk2_0)&(r1_tarski(esk1_0,k1_relat_1(esk2_0))&k1_relat_1(k7_relat_1(esk2_0,esk1_0))!=esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_16,plain,(X1=k3_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X1)|r2_hidden(esk3_3(X2,X3,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, (![X3]:k3_xboole_0(X3,X3)=X3), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_20,negated_conjecture,(k1_relat_1(k7_relat_1(esk2_0,esk1_0))!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(X1=k1_setfam_1(k2_enumset1(X2,X2,X2,X3))|r2_hidden(esk3_3(X2,X3,X1),X3)|r2_hidden(esk3_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
cnf(c_0_22,plain,(k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(r2_hidden(esk3_3(X1,X2,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),k1_relat_1(k7_relat_1(esk2_0,esk1_0)))|r2_hidden(esk3_3(X1,X2,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),X2)|k1_setfam_1(k2_enumset1(X1,X1,X1,X2))!=esk1_0), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24,plain,(k1_setfam_1(k2_enumset1(X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_17]), c_0_18])).
fof(c_0_25, plain, (![X4]:![X5]:![X6]:![X4]:![X5]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk4_2(X4,X5),X4)|r1_tarski(X4,X5))&(~r2_hidden(esk4_2(X4,X5),X5)|r1_tarski(X4,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_26, plain, (![X4]:![X5]:![X6]:((((r2_hidden(X4,X5)|~r2_hidden(X4,k1_relat_1(k7_relat_1(X6,X5))))|~v1_relat_1(X6))&((r2_hidden(X4,k1_relat_1(X6))|~r2_hidden(X4,k1_relat_1(k7_relat_1(X6,X5))))|~v1_relat_1(X6)))&(((~r2_hidden(X4,X5)|~r2_hidden(X4,k1_relat_1(X6)))|r2_hidden(X4,k1_relat_1(k7_relat_1(X6,X5))))|~v1_relat_1(X6)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_relat_1])])])).
cnf(c_0_27,negated_conjecture,(r2_hidden(esk3_3(X1,X1,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),k1_relat_1(k7_relat_1(esk2_0,esk1_0)))|r2_hidden(esk3_3(X1,X1,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),X1)|X1!=esk1_0), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29,negated_conjecture,(r1_tarski(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30,plain,(r2_hidden(X2,X3)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(k7_relat_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31,negated_conjecture,(r2_hidden(esk3_3(esk1_0,esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),k1_relat_1(k7_relat_1(esk2_0,esk1_0)))|r2_hidden(esk3_3(esk1_0,esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),esk1_0)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_32,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_33,negated_conjecture,(r2_hidden(X1,k1_relat_1(esk2_0))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34,negated_conjecture,(r2_hidden(esk3_3(esk1_0,esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_35,plain,(X1=k3_xboole_0(X2,X3)|~r2_hidden(esk3_3(X2,X3,X1),X3)|~r2_hidden(esk3_3(X2,X3,X1),X2)|~r2_hidden(esk3_3(X2,X3,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_36,plain,(r2_hidden(X2,k1_relat_1(k7_relat_1(X1,X3)))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37,negated_conjecture,(r2_hidden(esk3_3(esk1_0,esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38,plain,(X1=k1_setfam_1(k2_enumset1(X2,X2,X2,X3))|~r2_hidden(esk3_3(X2,X3,X1),X3)|~r2_hidden(esk3_3(X2,X3,X1),X2)|~r2_hidden(esk3_3(X2,X3,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_17]), c_0_18])).
cnf(c_0_39,negated_conjecture,(r2_hidden(esk3_3(esk1_0,esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),k1_relat_1(k7_relat_1(esk2_0,X1)))|~r2_hidden(esk3_3(esk1_0,esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_32])])).
cnf(c_0_40,negated_conjecture,(~r2_hidden(esk3_3(esk1_0,esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0))),k1_relat_1(k7_relat_1(esk2_0,esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_34]), c_0_24]), c_0_34])]), c_0_20])).
cnf(c_0_41,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_34]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 25
# Proof object formula steps           : 17
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 16
# Processed clauses                    : 30
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 29
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 177
# ...of the previous two non-trivial   : 155
# Contextual simplify-reflections      : 0
# Paramodulations                      : 162
# Factorizations                       : 2
# Equation resolutions                 : 13
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 133
# ...number of literals in the above   : 575
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 63
# Rec. Clause-clause subsumption calls : 54
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3970

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 42952 pages
