# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t98_enumset1, conjecture, (![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), file('out/t98_enumset1__2_premises__rerun.E_input', t98_enumset1)).
fof(t71_enumset1, axiom, (![X1]:![X2]:![X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), file('out/t98_enumset1__2_premises__rerun.E_input', t71_enumset1)).
fof(d1_enumset1, axiom, (![X1]:![X2]:![X3]:![X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3))))))), file('out/t98_enumset1__2_premises__rerun.E_input', d1_enumset1)).
fof(c_0_3, negated_conjecture, (~(![X1]:![X2]:![X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2))), inference(assume_negation,[status(cth)],[t98_enumset1])).
fof(c_0_4, negated_conjecture, (k1_enumset1(esk1_0,esk2_0,esk3_0)!=k1_enumset1(esk1_0,esk3_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
fof(c_0_5, plain, (![X4]:![X5]:![X6]:k2_enumset1(X4,X4,X5,X6)=k1_enumset1(X4,X5,X6)), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_6, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X10]:![X6]:![X7]:![X8]:![X9]:((((~r2_hidden(X10,X9)|(X10=X6|(X10=X7|X10=X8)))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X6|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&(((X10!=X7|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8))&((X10!=X8|r2_hidden(X10,X9))|X9!=k1_enumset1(X6,X7,X8)))))&((((esk4_4(X6,X7,X8,X9)!=X6|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&(((esk4_4(X6,X7,X8,X9)!=X7|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))&((esk4_4(X6,X7,X8,X9)!=X8|~r2_hidden(esk4_4(X6,X7,X8,X9),X9))|X9=k1_enumset1(X6,X7,X8))))&((r2_hidden(esk4_4(X6,X7,X8,X9),X9)|(esk4_4(X6,X7,X8,X9)=X6|(esk4_4(X6,X7,X8,X9)=X7|esk4_4(X6,X7,X8,X9)=X8)))|X9=k1_enumset1(X6,X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])])).
cnf(c_0_7,negated_conjecture,(k1_enumset1(esk1_0,esk2_0,esk3_0)!=k1_enumset1(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(X1=k1_enumset1(X2,X3,X4)|esk4_4(X2,X3,X4,X1)=X4|esk4_4(X2,X3,X4,X1)=X3|esk4_4(X2,X3,X4,X1)=X2|r2_hidden(esk4_4(X2,X3,X4,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(X5=X4|X5=X3|X5=X2|X1!=k1_enumset1(X2,X3,X4)|~r2_hidden(X5,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(k2_enumset1(esk1_0,esk1_0,esk3_0,esk2_0)!=k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_8])).
cnf(c_0_12,plain,(X1=k2_enumset1(X2,X2,X3,X4)|esk4_4(X2,X3,X4,X1)=X4|esk4_4(X2,X3,X4,X1)=X3|esk4_4(X2,X3,X4,X1)=X2|r2_hidden(esk4_4(X2,X3,X4,X1),X1)), inference(rw,[status(thm)],[c_0_9, c_0_8])).
cnf(c_0_13,plain,(X4=X5|X3=X5|X2=X5|X1!=k2_enumset1(X2,X2,X3,X4)|~r2_hidden(X5,X1)), inference(rw,[status(thm)],[c_0_10, c_0_8])).
cnf(c_0_14,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X3), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,plain,(X1=k1_enumset1(X2,X3,X4)|~r2_hidden(esk4_4(X2,X3,X4,X1),X1)|esk4_4(X2,X3,X4,X1)!=X4), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,negated_conjecture,(esk4_4(X1,X2,X3,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=X1|esk4_4(X1,X2,X3,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=X2|esk4_4(X1,X2,X3,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=X3|r2_hidden(esk4_4(X1,X2,X3,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0)),k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))|k2_enumset1(X1,X1,X2,X3)!=k2_enumset1(esk1_0,esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17,plain,(X1=X2|X3=X2|X4=X2|~r2_hidden(X2,k2_enumset1(X4,X4,X3,X1))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(r2_hidden(X5,X1)|X3!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_14, c_0_8])).
cnf(c_0_19,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X4), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_20,plain,(X1=k1_enumset1(X2,X3,X4)|~r2_hidden(esk4_4(X2,X3,X4,X1),X1)|esk4_4(X2,X3,X4,X1)!=X3), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21,plain,(X1=k2_enumset1(X2,X2,X3,X4)|esk4_4(X2,X3,X4,X1)!=X4|~r2_hidden(esk4_4(X2,X3,X4,X1),X1)), inference(rw,[status(thm)],[c_0_15, c_0_8])).
cnf(c_0_22,negated_conjecture,(esk4_4(esk1_0,esk3_0,esk2_0,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=esk2_0|esk4_4(esk1_0,esk3_0,esk2_0,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=esk3_0|esk4_4(esk1_0,esk3_0,esk2_0,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=esk1_0), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_16]), c_0_17])).
cnf(c_0_23,plain,(r2_hidden(X1,k2_enumset1(X2,X2,X1,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_18])])).
cnf(c_0_24,plain,(r2_hidden(X5,X1)|X4!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_19, c_0_8])).
cnf(c_0_25,plain,(r2_hidden(X5,X1)|X1!=k1_enumset1(X2,X3,X4)|X5!=X2), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26,plain,(X1=k1_enumset1(X2,X3,X4)|~r2_hidden(esk4_4(X2,X3,X4,X1),X1)|esk4_4(X2,X3,X4,X1)!=X2), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_27,plain,(X1=k2_enumset1(X2,X2,X3,X4)|esk4_4(X2,X3,X4,X1)!=X3|~r2_hidden(esk4_4(X2,X3,X4,X1),X1)), inference(rw,[status(thm)],[c_0_20, c_0_8])).
cnf(c_0_28,negated_conjecture,(esk4_4(esk1_0,esk3_0,esk2_0,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=esk1_0|esk4_4(esk1_0,esk3_0,esk2_0,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])]), c_0_11])).
cnf(c_0_29,plain,(r2_hidden(X1,k2_enumset1(X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_24])])).
cnf(c_0_30,plain,(r2_hidden(X5,X1)|X2!=X5|X1!=k2_enumset1(X2,X2,X3,X4)), inference(rw,[status(thm)],[c_0_25, c_0_8])).
cnf(c_0_31,plain,(X1=k2_enumset1(X2,X2,X3,X4)|esk4_4(X2,X3,X4,X1)!=X2|~r2_hidden(esk4_4(X2,X3,X4,X1),X1)), inference(rw,[status(thm)],[c_0_26, c_0_8])).
cnf(c_0_32,negated_conjecture,(esk4_4(esk1_0,esk3_0,esk2_0,k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0))=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])]), c_0_11])).
cnf(c_0_33,plain,(r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_30])])).
cnf(c_0_34,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])]), c_0_11]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 28
# Proof object formula steps           : 7
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 9
# Processed clauses                    : 26
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 26
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 1
# Generated clauses                    : 46
# ...of the previous two non-trivial   : 37
# Contextual simplify-reflections      : 1
# Paramodulations                      : 34
# Factorizations                       : 3
# Equation resolutions                 : 12
# Current number of processed clauses  : 11
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 6
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 126
# Current number of archived formulas  : 0
# Current number of archived clauses   : 12
# Clause-clause subsumption calls (NU) : 17
# Rec. Clause-clause subsumption calls : 16
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1413

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 45240 pages
