# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_mcart_1, conjecture, (![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3]:![X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4)))))))), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', t9_mcart_1)).
fof(d5_tarski, axiom, (![X1]:![X2]:k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', d5_tarski)).
fof(t7_boole, axiom, (![X1]:![X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', t7_boole)).
fof(t3_xboole_0, axiom, (![X1]:![X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2))))), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', t3_xboole_0)).
fof(d2_tarski, axiom, (![X1]:![X2]:![X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2)))), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', d2_tarski)).
fof(commutativity_k2_tarski, axiom, (![X1]:![X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1)), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', commutativity_k2_tarski)).
fof(t2_xregular, axiom, (![X1]:(~(v1_xboole_0(X1))=>?[X2]:(r2_hidden(X2,X1)&![X3]:(r2_hidden(X3,X2)=>r1_xboole_0(X3,X1))))), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', t2_xregular)).
fof(t6_boole, axiom, (![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0)), file('out/t9_mcart_1__8_premises__rerun__rerun.E_input', t6_boole)).
fof(c_0_8, negated_conjecture, (~(![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3]:![X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))))), inference(assume_negation,[status(cth)],[t9_mcart_1])).
fof(c_0_9, negated_conjecture, (![X6]:(esk1_0!=k1_xboole_0&(((r2_hidden(esk2_1(X6),esk1_0)|r2_hidden(esk3_1(X6),esk1_0))|~r2_hidden(X6,esk1_0))&(X6=k4_tarski(esk2_1(X6),esk3_1(X6))|~r2_hidden(X6,esk1_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])])])).
fof(c_0_10, plain, (![X3]:![X4]:k4_tarski(X3,X4)=k2_tarski(k2_tarski(X3,X4),k1_tarski(X3))), inference(variable_rename,[status(thm)],[d5_tarski])).
fof(c_0_11, plain, (![X3]:![X4]:(~r2_hidden(X3,X4)|~v1_xboole_0(X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_12, plain, (![X4]:![X5]:![X4]:![X5]:![X7]:(((r2_hidden(esk4_2(X4,X5),X4)|r1_xboole_0(X4,X5))&(r2_hidden(esk4_2(X4,X5),X5)|r1_xboole_0(X4,X5)))&((~r2_hidden(X7,X4)|~r2_hidden(X7,X5))|~r1_xboole_0(X4,X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])])).
fof(c_0_13, plain, (![X5]:![X6]:![X7]:![X8]:![X8]:![X5]:![X6]:![X7]:((((~r2_hidden(X8,X7)|(X8=X5|X8=X6))|X7!=k2_tarski(X5,X6))&(((X8!=X5|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))&((X8!=X6|r2_hidden(X8,X7))|X7!=k2_tarski(X5,X6))))&((((esk5_3(X5,X6,X7)!=X5|~r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6))&((esk5_3(X5,X6,X7)!=X6|~r2_hidden(esk5_3(X5,X6,X7),X7))|X7=k2_tarski(X5,X6)))&((r2_hidden(esk5_3(X5,X6,X7),X7)|(esk5_3(X5,X6,X7)=X5|esk5_3(X5,X6,X7)=X6))|X7=k2_tarski(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])])).
cnf(c_0_14,negated_conjecture,(X1=k4_tarski(esk2_1(X1),esk3_1(X1))|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, plain, (![X3]:![X4]:k2_tarski(X3,X4)=k2_tarski(X4,X3)), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_17, plain, (![X4]:![X6]:((r2_hidden(esk6_1(X4),X4)|v1_xboole_0(X4))&((~r2_hidden(X6,esk6_1(X4))|r1_xboole_0(X6,X4))|v1_xboole_0(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t2_xregular])])])])])])])])).
cnf(c_0_18,plain,(~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,(r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(r2_hidden(X4,X1)|X1!=k2_tarski(X2,X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,negated_conjecture,(X1=k2_tarski(k2_tarski(esk2_1(X1),esk3_1(X1)),k1_tarski(esk2_1(X1)))|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22,plain,(k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(v1_xboole_0(X1)|r1_xboole_0(X2,X1)|~r2_hidden(X2,esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,plain,(r1_xboole_0(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,plain,(r2_hidden(X1,k2_tarski(X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_20])])).
cnf(c_0_26,negated_conjecture,(k2_tarski(k1_tarski(esk2_1(X1)),k2_tarski(esk2_1(X1),esk3_1(X1)))=X1|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27,plain,(r1_xboole_0(X1,X2)|~r2_hidden(X1,esk6_1(X2))), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28,negated_conjecture,(r2_hidden(k2_tarski(esk2_1(X1),esk3_1(X1)),X1)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29,plain,(~r1_xboole_0(X1,X2)|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30,negated_conjecture,(r1_xboole_0(k2_tarski(esk2_1(esk6_1(X1)),esk3_1(esk6_1(X1))),X1)|~r2_hidden(esk6_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31,plain,(r2_hidden(X4,X1)|X1!=k2_tarski(X2,X3)|X4!=X2), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32,negated_conjecture,(~r2_hidden(X1,k2_tarski(esk2_1(esk6_1(X2)),esk3_1(esk6_1(X2))))|~r2_hidden(esk6_1(X2),esk1_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33,plain,(r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
cnf(c_0_34,negated_conjecture,(~r2_hidden(esk3_1(esk6_1(X1)),X1)|~r2_hidden(esk6_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_25])).
cnf(c_0_35,negated_conjecture,(r2_hidden(esk3_1(X1),esk1_0)|r2_hidden(esk2_1(X1),esk1_0)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36,negated_conjecture,(~r2_hidden(esk2_1(esk6_1(X1)),X1)|~r2_hidden(esk6_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37,negated_conjecture,(r2_hidden(esk2_1(esk6_1(esk1_0)),esk1_0)|~r2_hidden(esk6_1(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_38, plain, (![X2]:(~v1_xboole_0(X2)|X2=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_39,negated_conjecture,(~r2_hidden(esk6_1(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40,plain,(v1_xboole_0(X1)|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_41,plain,(X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42,negated_conjecture,(v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_43,negated_conjecture,(esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_44,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 28
# Proof object formula steps           : 17
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 17
# Processed clauses                    : 406
# ...of these trivial                  : 4
# ...subsumed                          : 205
# ...remaining for further processing  : 197
# Other redundant clauses eliminated   : 29
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 1171
# ...of the previous two non-trivial   : 1093
# Contextual simplify-reflections      : 90
# Paramodulations                      : 1122
# Factorizations                       : 22
# Equation resolutions                 : 29
# Current number of processed clauses  : 176
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 163
# Current number of unprocessed clauses: 714
# ...number of literals in the above   : 3317
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 7885
# Rec. Clause-clause subsumption calls : 2671
# Non-unit clause-clause subsumptions  : 262
# Unit Clause-clause subsumption calls : 91
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20876

# -------------------------------------------------
# User time                : 0.040 s
# System time              : 0.004 s
# Total time               : 0.044 s
# Maximum resident set size: 45276 pages
