# No SInE strategy applied
# Auto-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.012 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_yellow_0, conjecture, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4))))))), file('out/t9_yellow_0__4_premises__rerun.E_input', t9_yellow_0)).
fof(d9_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3))))))), file('out/t9_yellow_0__4_premises__rerun.E_input', d9_lattice3)).
fof(t3_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('out/t9_yellow_0__4_premises__rerun.E_input', t3_subset)).
fof(l3_subset_1, axiom, (![X1]:![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1)))), file('out/t9_yellow_0__4_premises__rerun.E_input', l3_subset_1)).
fof(d8_lattice3, axiom, (![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4))))))), file('out/t9_yellow_0__4_premises__rerun.E_input', d8_lattice3)).
fof(c_0_5, negated_conjecture, (~(![X1]:(l1_orders_2(X1)=>![X2]:![X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))))), inference(assume_negation,[status(cth)],[t9_yellow_0])).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:((((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk5_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk5_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,esk5_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])])])).
fof(c_0_7, negated_conjecture, ((l1_orders_2(esk1_0)&(r1_tarski(esk2_0,esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(((r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0))&(~r2_lattice3(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)))&((r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0))&(~r2_lattice3(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_8,plain,(r2_lattice3(X1,X3,X2)|m1_subset_1(esk5_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,negated_conjecture,(m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (![X3]:![X4]:![X3]:![X4]:((~m1_subset_1(X3,k1_zfmisc_1(X4))|r1_tarski(X3,X4))&(~r1_tarski(X3,X4)|m1_subset_1(X3,k1_zfmisc_1(X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_12,plain,(r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,negated_conjecture,(r2_lattice3(esk1_0,X1,esk4_0)|m1_subset_1(esk5_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(~m1_subset_1(X5,k1_zfmisc_1(X4))|(~r2_hidden(X6,X5)|r2_hidden(X6,X4)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])])])).
cnf(c_0_15,plain,(m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(esk1_0,X1,esk4_0),X2)|r2_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk5_3(esk1_0,X1,esk4_0),X3)|~r2_lattice3(esk1_0,X3,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_10])])).
cnf(c_0_18,plain,(r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,negated_conjecture,(m1_subset_1(esk2_0,k1_zfmisc_1(esk3_0))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20,plain,(r2_lattice3(X1,X3,X2)|r2_hidden(esk5_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_21, plain, (![X5]:![X6]:![X7]:![X8]:((((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8))))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((m1_subset_1(esk6_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&((((r2_hidden(esk6_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5))&(((~r1_orders_2(X5,X7,esk6_3(X5,X6,X7))|r1_lattice3(X5,X6,X7))|~m1_subset_1(X7,u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])])])).
cnf(c_0_22,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(esk1_0,X1,esk4_0),esk4_0)|r2_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk5_3(esk1_0,X1,esk4_0),X2)|~r2_lattice3(esk1_0,X2,esk4_0)), inference(spm,[status(thm)],[c_0_17, c_0_9])).
cnf(c_0_23,negated_conjecture,(r1_lattice3(esk1_0,esk3_0,esk4_0)|r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24,negated_conjecture,(r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,negated_conjecture,(r2_hidden(esk5_3(esk1_0,X1,esk4_0),X1)|r2_lattice3(esk1_0,X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_9]), c_0_10])])).
cnf(c_0_26,plain,(r1_lattice3(X1,X3,X2)|m1_subset_1(esk6_3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,plain,(r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk5_3(X1,X3,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_28,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(esk1_0,X1,esk4_0),esk4_0)|r2_lattice3(esk1_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk5_3(esk1_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29,negated_conjecture,(r2_hidden(esk5_3(esk1_0,esk2_0,esk4_0),esk3_0)|r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30,plain,(r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31,negated_conjecture,(r1_lattice3(esk1_0,X1,esk4_0)|m1_subset_1(esk6_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_9]), c_0_10])])).
cnf(c_0_32,negated_conjecture,(r2_lattice3(esk1_0,X1,esk4_0)|~r1_orders_2(esk1_0,esk5_3(esk1_0,X1,esk4_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_9]), c_0_10])])).
cnf(c_0_33,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(esk1_0,esk2_0,esk4_0),esk4_0)|r2_lattice3(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34,negated_conjecture,(r1_orders_2(esk1_0,X1,esk6_3(esk1_0,X2,esk4_0))|r1_lattice3(esk1_0,X2,esk4_0)|~r2_hidden(esk6_3(esk1_0,X2,esk4_0),X3)|~r1_lattice3(esk1_0,X3,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_10])])).
cnf(c_0_35,negated_conjecture,(r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_36,negated_conjecture,(r2_lattice3(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37,plain,(r1_lattice3(X1,X3,X2)|r2_hidden(esk6_3(X1,X3,X2),X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk1_0,X1,esk4_0))|r1_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk6_3(esk1_0,X1,esk4_0),X2)|~r1_lattice3(esk1_0,X2,esk4_0)), inference(spm,[status(thm)],[c_0_34, c_0_9])).
cnf(c_0_39,negated_conjecture,(r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40,negated_conjecture,(r2_hidden(esk6_3(esk1_0,X1,esk4_0),X1)|r1_lattice3(esk1_0,X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_9]), c_0_10])])).
cnf(c_0_41,plain,(r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_42,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk1_0,X1,esk4_0))|r1_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk6_3(esk1_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43,negated_conjecture,(r2_hidden(esk6_3(esk1_0,esk2_0,esk4_0),esk3_0)|r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_24, c_0_40])).
cnf(c_0_44,negated_conjecture,(r1_lattice3(esk1_0,X1,esk4_0)|~r1_orders_2(esk1_0,esk4_0,esk6_3(esk1_0,X1,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_9]), c_0_10])])).
cnf(c_0_45,negated_conjecture,(r1_orders_2(esk1_0,esk4_0,esk6_3(esk1_0,esk2_0,esk4_0))|r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_47,negated_conjecture,(r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_48,negated_conjecture,(~r1_lattice3(esk1_0,esk2_0,esk4_0)|~r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_49,negated_conjecture,(r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47])])).
cnf(c_0_50,negated_conjecture,(~r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_47])])).
cnf(c_0_51,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(esk1_0,X1,esk4_0),esk4_0)|r2_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk5_3(esk1_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_49])).
cnf(c_0_52,negated_conjecture,(r2_hidden(esk5_3(esk1_0,esk2_0,esk4_0),esk3_0)), inference(sr,[status(thm)],[c_0_29, c_0_50])).
cnf(c_0_53,negated_conjecture,(r1_orders_2(esk1_0,esk5_3(esk1_0,esk2_0,esk4_0),esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_50])).
cnf(c_0_54,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_53]), c_0_50]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 44
# Proof object formula steps           : 11
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 69
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 11
# Generated clauses                    : 107
# ...of the previous two non-trivial   : 102
# Contextual simplify-reflections      : 0
# Paramodulations                      : 106
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 55
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 44
# Current number of unprocessed clauses: 48
# ...number of literals in the above   : 224
# Current number of archived formulas  : 0
# Current number of archived clauses   : 13
# Clause-clause subsumption calls (NU) : 166
# Rec. Clause-clause subsumption calls : 45
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 81
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4477

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45228 pages
