# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.012 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d13_pralg_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_partfun1(X2,X1)&v2_pralg_1(X2)))))=>![X3]:((v1_relat_1(X3)&(v4_relat_1(X3,X1)&(v1_funct_1(X3)&v1_partfun1(X3,X1))))=>(X3=k12_pralg_1(X1,X2)<=>![X4]:~((r2_hidden(X4,X1)&![X5]:(l1_struct_0(X5)=>~((X5=k1_funct_1(X2,X4)&k1_funct_1(X3,X4)=u1_struct_0(X5)))))))))), file('out/t9_yellow_6__4_premises__rerun.E_input', d13_pralg_1)).
fof(t9_yellow_6, conjecture, (![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_partfun1(X2,X1)&v2_pralg_1(X2)))))=>![X3]:(m1_subset_1(X3,X1)=>k1_funct_1(k12_pralg_1(X1,X2),X3)=u1_struct_0(k10_pralg_1(X1,X2,X3)))))), file('out/t9_yellow_6__4_premises__rerun.E_input', t9_yellow_6)).
fof(dt_k12_pralg_1, axiom, (![X1]:![X2]:((v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_partfun1(X2,X1)&v2_pralg_1(X2)))))=>(v1_relat_1(k12_pralg_1(X1,X2))&(v4_relat_1(k12_pralg_1(X1,X2),X1)&(v1_funct_1(k12_pralg_1(X1,X2))&v1_partfun1(k12_pralg_1(X1,X2),X1)))))), file('out/t9_yellow_6__4_premises__rerun.E_input', dt_k12_pralg_1)).
fof(t2_subset, axiom, (![X1]:![X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('out/t9_yellow_6__4_premises__rerun.E_input', t2_subset)).
fof(redefinition_k10_pralg_1, axiom, (![X1]:![X2]:![X3]:((~(v1_xboole_0(X1))&((v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_partfun1(X2,X1)&v2_pralg_1(X2)))))&m1_subset_1(X3,X1)))=>k10_pralg_1(X1,X2,X3)=k1_funct_1(X2,X3))), file('out/t9_yellow_6__4_premises__rerun.E_input', redefinition_k10_pralg_1)).
fof(c_0_5, plain, (![X6]:![X7]:![X8]:![X9]:![X12]:((((((l1_struct_0(esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X6))|X8!=k12_pralg_1(X6,X7))|(~v1_relat_1(X8)|(~v4_relat_1(X8,X6)|(~v1_funct_1(X8)|~v1_partfun1(X8,X6)))))|(~v1_relat_1(X7)|(~v4_relat_1(X7,X6)|(~v1_funct_1(X7)|(~v1_partfun1(X7,X6)|~v2_pralg_1(X7))))))&(((((esk1_4(X6,X7,X8,X9)=k1_funct_1(X7,X9)|~r2_hidden(X9,X6))|X8!=k12_pralg_1(X6,X7))|(~v1_relat_1(X8)|(~v4_relat_1(X8,X6)|(~v1_funct_1(X8)|~v1_partfun1(X8,X6)))))|(~v1_relat_1(X7)|(~v4_relat_1(X7,X6)|(~v1_funct_1(X7)|(~v1_partfun1(X7,X6)|~v2_pralg_1(X7))))))&((((k1_funct_1(X8,X9)=u1_struct_0(esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X6))|X8!=k12_pralg_1(X6,X7))|(~v1_relat_1(X8)|(~v4_relat_1(X8,X6)|(~v1_funct_1(X8)|~v1_partfun1(X8,X6)))))|(~v1_relat_1(X7)|(~v4_relat_1(X7,X6)|(~v1_funct_1(X7)|(~v1_partfun1(X7,X6)|~v2_pralg_1(X7))))))))&((((r2_hidden(esk2_3(X6,X7,X8),X6)|X8=k12_pralg_1(X6,X7))|(~v1_relat_1(X8)|(~v4_relat_1(X8,X6)|(~v1_funct_1(X8)|~v1_partfun1(X8,X6)))))|(~v1_relat_1(X7)|(~v4_relat_1(X7,X6)|(~v1_funct_1(X7)|(~v1_partfun1(X7,X6)|~v2_pralg_1(X7))))))&((((~l1_struct_0(X12)|(X12!=k1_funct_1(X7,esk2_3(X6,X7,X8))|k1_funct_1(X8,esk2_3(X6,X7,X8))!=u1_struct_0(X12)))|X8=k12_pralg_1(X6,X7))|(~v1_relat_1(X8)|(~v4_relat_1(X8,X6)|(~v1_funct_1(X8)|~v1_partfun1(X8,X6)))))|(~v1_relat_1(X7)|(~v4_relat_1(X7,X6)|(~v1_funct_1(X7)|(~v1_partfun1(X7,X6)|~v2_pralg_1(X7))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_pralg_1])])])])])])])).
fof(c_0_6, negated_conjecture, (~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_partfun1(X2,X1)&v2_pralg_1(X2)))))=>![X3]:(m1_subset_1(X3,X1)=>k1_funct_1(k12_pralg_1(X1,X2),X3)=u1_struct_0(k10_pralg_1(X1,X2,X3))))))), inference(assume_negation,[status(cth)],[t9_yellow_6])).
cnf(c_0_7,plain,(k1_funct_1(X3,X4)=u1_struct_0(esk1_4(X2,X1,X3,X4))|~v2_pralg_1(X1)|~v1_partfun1(X1,X2)|~v1_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_partfun1(X3,X2)|~v1_funct_1(X3)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|X3!=k12_pralg_1(X2,X1)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(esk1_4(X2,X1,X3,X4)=k1_funct_1(X1,X4)|~v2_pralg_1(X1)|~v1_partfun1(X1,X2)|~v1_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_partfun1(X3,X2)|~v1_funct_1(X3)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|X3!=k12_pralg_1(X2,X1)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_9, plain, (![X3]:![X4]:((v1_relat_1(k12_pralg_1(X3,X4))|(~v1_relat_1(X4)|(~v4_relat_1(X4,X3)|(~v1_funct_1(X4)|(~v1_partfun1(X4,X3)|~v2_pralg_1(X4))))))&((v4_relat_1(k12_pralg_1(X3,X4),X3)|(~v1_relat_1(X4)|(~v4_relat_1(X4,X3)|(~v1_funct_1(X4)|(~v1_partfun1(X4,X3)|~v2_pralg_1(X4))))))&((v1_funct_1(k12_pralg_1(X3,X4))|(~v1_relat_1(X4)|(~v4_relat_1(X4,X3)|(~v1_funct_1(X4)|(~v1_partfun1(X4,X3)|~v2_pralg_1(X4))))))&(v1_partfun1(k12_pralg_1(X3,X4),X3)|(~v1_relat_1(X4)|(~v4_relat_1(X4,X3)|(~v1_funct_1(X4)|(~v1_partfun1(X4,X3)|~v2_pralg_1(X4)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k12_pralg_1])])])).
fof(c_0_10, plain, (![X3]:![X4]:(~m1_subset_1(X3,X4)|(v1_xboole_0(X4)|r2_hidden(X3,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_11, negated_conjecture, ((~v1_xboole_0(esk3_0)&((v1_relat_1(esk4_0)&(v4_relat_1(esk4_0,esk3_0)&(v1_funct_1(esk4_0)&(v1_partfun1(esk4_0,esk3_0)&v2_pralg_1(esk4_0)))))&(m1_subset_1(esk5_0,esk3_0)&k1_funct_1(k12_pralg_1(esk3_0,esk4_0),esk5_0)!=u1_struct_0(k10_pralg_1(esk3_0,esk4_0,esk5_0)))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
fof(c_0_12, plain, (![X4]:![X5]:![X6]:((v1_xboole_0(X4)|((~v1_relat_1(X5)|(~v4_relat_1(X5,X4)|(~v1_funct_1(X5)|(~v1_partfun1(X5,X4)|~v2_pralg_1(X5)))))|~m1_subset_1(X6,X4)))|k10_pralg_1(X4,X5,X6)=k1_funct_1(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k10_pralg_1])])])).
cnf(c_0_13,plain,(u1_struct_0(k1_funct_1(X1,X2))=k1_funct_1(X3,X2)|X3!=k12_pralg_1(X4,X1)|~r2_hidden(X2,X4)|~v2_pralg_1(X1)|~v1_partfun1(X3,X4)|~v1_partfun1(X1,X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v4_relat_1(X3,X4)|~v4_relat_1(X1,X4)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_14,plain,(v1_relat_1(k12_pralg_1(X2,X1))|~v2_pralg_1(X1)|~v1_partfun1(X1,X2)|~v1_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(v4_relat_1(k12_pralg_1(X2,X1),X2)|~v2_pralg_1(X1)|~v1_partfun1(X1,X2)|~v1_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16,plain,(v1_funct_1(k12_pralg_1(X2,X1))|~v2_pralg_1(X1)|~v1_partfun1(X1,X2)|~v1_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,plain,(v1_partfun1(k12_pralg_1(X2,X1),X2)|~v2_pralg_1(X1)|~v1_partfun1(X1,X2)|~v1_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,plain,(r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,negated_conjecture,(m1_subset_1(esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,negated_conjecture,(~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21,negated_conjecture,(k1_funct_1(k12_pralg_1(esk3_0,esk4_0),esk5_0)!=u1_struct_0(k10_pralg_1(esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22,plain,(k10_pralg_1(X1,X2,X3)=k1_funct_1(X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~v2_pralg_1(X2)|~v1_partfun1(X2,X1)|~v1_funct_1(X2)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(v2_pralg_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24,negated_conjecture,(v1_partfun1(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,negated_conjecture,(v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_26,negated_conjecture,(v4_relat_1(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_27,negated_conjecture,(v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28,plain,(u1_struct_0(k1_funct_1(X1,X2))=k1_funct_1(k12_pralg_1(X3,X1),X2)|~r2_hidden(X2,X3)|~v2_pralg_1(X1)|~v1_partfun1(X1,X3)|~v1_funct_1(X1)|~v4_relat_1(X1,X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17])).
cnf(c_0_29,negated_conjecture,(r2_hidden(esk5_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_30,negated_conjecture,(u1_struct_0(k1_funct_1(esk4_0,esk5_0))!=k1_funct_1(k12_pralg_1(esk3_0,esk4_0),esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_19]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])]), c_0_20])).
cnf(c_0_31,negated_conjecture,(u1_struct_0(k1_funct_1(X1,esk5_0))=k1_funct_1(k12_pralg_1(esk3_0,X1),esk5_0)|~v2_pralg_1(X1)|~v1_partfun1(X1,esk3_0)|~v1_funct_1(X1)|~v4_relat_1(X1,esk3_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 22
# Proof object formula steps           : 11
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 12
# ...of the previous two non-trivial   : 11
# Contextual simplify-reflections      : 8
# Paramodulations                      : 9
# Factorizations                       : 0
# Equation resolutions                 : 3
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 44
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 373
# Rec. Clause-clause subsumption calls : 9
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2591

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.004 s
# Total time               : 0.016 s
# Maximum resident set size: 45300 pages
